<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC397+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC397+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC397+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC397%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC397+1.p">SWC397+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n128.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:47 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :   77 ( 974 expanded)
%              Number of clauses        :   51 ( 373 expanded)
%              Number of leaves         :   13 ( 250 expanded)
%              Depth                    :   15
%              Number of atoms          :  295 (8837 expanded)
%              Number of equality atoms :   87 (2982 expanded)
%              Maximal formula depth    :   21 (   5 average)
%              Maximal clause size      :   19 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ! [X5] :
                        ( ssItem(X5)
                       => ( ~ memberP(X1,X5)
                          | memberP(X2,X5) ) )
                    | ( nil != X3
                      & nil = X4 )
                    | ( ! [X6] :
                          ( ssItem(X6)
                         => ! [X7] :
                              ( ssList(X7)
                             => ( app(cons(X6,nil),X7) != X3
                                | app(X7,cons(X6,nil)) != X4 ) ) )
                      & neq(X4,nil) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax27)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(app(X2,X3),X1)
              &lt;=> ( memberP(X2,X1)
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax36)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X3,X2) = app(X1,X2)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax79)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => hd(cons(X2,X1)) = X2 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax23)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( memberP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(X3,cons(X2,X4)) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax3)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssItem(hd(X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax22)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(cons(X2,X3),X1)
              &lt;=> ( X1 = X2
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax37)).

<A NAME="c_0_13"></A>fof(c_0_13,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ! [X5] :
                          ( ssItem(X5)
                         => ( ~ memberP(X1,X5)
                            | memberP(X2,X5) ) )
                      | ( nil != X3
                        & nil = X4 )
                      | ( ! [X6] :
                            ( ssItem(X6)
                           => ! [X7] :
                                ( ssList(X7)
                               => ( app(cons(X6,nil),X7) != X3
                                  | app(X7,cons(X6,nil)) != X4 ) ) )
                        & neq(X4,nil) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_14"></A>fof(c_0_14,negated_conjecture,
    ( ssList(esk48_0)
    & ssList(esk49_0)
    & ssList(esk50_0)
    & ssList(esk51_0)
    & esk49_0 = esk51_0
    & esk48_0 = esk50_0
    & ssItem(esk52_0)
    & memberP(esk48_0,esk52_0)
    & ~ memberP(esk49_0,esk52_0)
    & ( nil = esk50_0
      | nil != esk51_0 )
    & ( ssItem(esk53_0)
      | ~ neq(esk51_0,nil) )
    & ( ssList(esk54_0)
      | ~ neq(esk51_0,nil) )
    & ( app(cons(esk53_0,nil),esk54_0) = esk50_0
      | ~ neq(esk51_0,nil) )
    & ( app(esk54_0,cons(esk53_0,nil)) = esk51_0
      | ~ neq(esk51_0,nil) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).

<A NAME="c_0_15"></A>fof(c_0_15,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( app(cons(esk53_0,nil),esk54_0) = esk50_0
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_22"></A>fof(c_0_22,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( ssItem(esk53_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( ssList(esk54_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_25"></A>fof(c_0_25,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssItem(X6)
      | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( app(cons(esk53_0,nil),esk54_0) = esk50_0
    | esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_19]),c_0_20]),c_0_21])])).

<A NAME="c_0_27"></A>cnf(c_0_27,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( esk51_0 = nil
    | ssItem(esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_19]),c_0_20]),c_0_21])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( esk51_0 = nil
    | ssList(esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_19]),c_0_20]),c_0_21])])).

<A NAME="c_0_30"></A>cnf(c_0_30,plain,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( cons(esk53_0,esk54_0) = esk50_0
    | esk51_0 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_27]),c_0_28]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( app(esk54_0,cons(esk53_0,nil)) = esk51_0
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_33"></A>fof(c_0_33,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(app(X5,X6),X4)
        | memberP(X5,X4)
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X5,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( memberP(esk48_0,esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( cons(esk53_0,app(esk54_0,X1)) = app(esk50_0,X1)
    | esk51_0 = nil
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31]),c_0_28]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( app(esk54_0,cons(esk53_0,nil)) = esk51_0
    | esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_19]),c_0_20]),c_0_21])])).

<A NAME="c_0_39"></A>fof(c_0_39,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).

<A NAME="c_0_40"></A>cnf(c_0_40,plain,
    ( memberP(app(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(X2,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( memberP(esk50_0,esk52_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( ssItem(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( app(esk50_0,cons(esk53_0,nil)) = cons(esk53_0,esk51_0)
    | esk51_0 = nil
    | ~ ssList(cons(esk53_0,nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_46"></A>fof(c_0_46,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X6,X5) != app(X4,X5)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).

<A NAME="c_0_47"></A>fof(c_0_47,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).

<A NAME="c_0_48"></A>fof(c_0_48,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | hd(cons(X4,X3)) = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_49"></A>fof(c_0_49,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk3_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ssList(esk4_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6))) = X5
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(X9,cons(X6,X10)) != X5
        | memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( memberP(app(esk50_0,X1),esk52_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_41]),c_0_42]),c_0_43])])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( app(esk50_0,cons(esk53_0,nil)) = cons(esk53_0,esk51_0)
    | esk51_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_45]),c_0_20])]),<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,plain,
    ( X1 = X2
    | app(X1,X3) != app(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_54"></A>fof(c_0_54,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssItem(hd(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])).

<A NAME="c_0_55"></A>cnf(c_0_55,plain,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( memberP(X1,X2)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | app(X3,cons(X2,X4)) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_57"></A>fof(c_0_57,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(cons(X5,X6),X4)
        | X4 = X5
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( X4 != X5
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( esk51_0 = nil
    | memberP(cons(esk53_0,esk51_0),esk52_0)
    | ~ ssList(cons(esk53_0,nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,<A HREF="#c_0_51">c_0_51</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( ~ memberP(esk49_0,esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_60"></A>cnf(c_0_60,plain,
    ( X1 = nil
    | app(X1,X2) != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_53]),c_0_20])])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( nil = esk50_0
    | nil != esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_62"></A>cnf(c_0_62,plain,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( hd(esk50_0) = esk53_0
    | esk51_0 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_31]),c_0_28]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_64"></A>cnf(c_0_64,plain,
    ( memberP(app(X1,cons(X2,X3)),X2)
    | ~ ssList(app(X1,cons(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( memberP(X3,X1)
    | X1 = X2
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ memberP(cons(X2,X3),X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,negated_conjecture,
    ( esk51_0 = nil
    | memberP(cons(esk53_0,esk51_0),esk52_0) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_45]),c_0_20])]),<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( ~ memberP(esk51_0,esk52_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( esk50_0 = nil
    | cons(esk53_0,esk51_0) != cons(esk53_0,nil)
    | ~ ssList(cons(esk53_0,nil)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>,c_0_51]),c_0_42])]),<A HREF="#c_0_61">c_0_61</A>])).

<A NAME="c_0_69"></A>cnf(c_0_69,negated_conjecture,
    ( esk50_0 = nil
    | ssItem(esk53_0) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_63]),c_0_42])]),<A HREF="#c_0_61">c_0_61</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( esk51_0 = nil
    | memberP(esk51_0,esk53_0) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,c_0_38]),c_0_21]),c_0_20])]),c_0_28]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( esk51_0 = nil
    | esk53_0 = esk52_0 ),
    inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_66]),c_0_21]),c_0_43])]),c_0_67]),<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_72"></A>cnf(c_0_72,negated_conjecture,
    ( esk50_0 = nil
    | cons(esk53_0,esk51_0) != cons(esk53_0,nil) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_45]),c_0_20])]),<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,negated_conjecture,
    ( esk51_0 = nil ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>,c_0_71]),<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_73])])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( ~ memberP(nil,esk52_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,<A HREF="#c_0_73">c_0_73</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_74]),<A HREF="#c_0_75">c_0_75</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.04  % Problem    : SWC397+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n128.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:43:37 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.44  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.44  # and selection function SelectCQPrecWNTNp.
% 0.06/0.44  #
% 0.06/0.44  # Preprocessing time       : 0.014 s
% 0.06/0.44  # Presaturation interreduction done
% 0.06/0.44  
% 0.06/0.44  # Proof found!
% 0.06/0.44  # SZS status Theorem
% 0.06/0.44  # SZS output start CNFRefutation
% 0.06/0.44  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssItem(X5)=>(~(memberP(X1,X5))|memberP(X2,X5))))|(nil!=X3&nil=X4))|(![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>(app(cons(X6,nil),X7)!=X3|app(X7,cons(X6,nil))!=X4)))&neq(X4,nil)))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.06/0.44  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.06/0.44  fof(c_0_2, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.44  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.06/0.44  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax27)).
% 0.06/0.44  fof(c_0_5, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(memberP(app(X2,X3),X1)&lt;=>(memberP(X2,X1)|memberP(X3,X1))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax36)).
% 0.06/0.44  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.06/0.44  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax79)).
% 0.06/0.44  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
% 0.06/0.44  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax23)).
% 0.06/0.44  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>(memberP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(X3,cons(X2,X4))=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax3)).
% 0.06/0.44  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax22)).
% 0.06/0.44  fof(c_0_12, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)&lt;=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax37)).
% 0.06/0.44  fof(c_0_13, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssItem(X5)=>(~memberP(X1,X5)|memberP(X2,X5))))|(nil!=X3&nil=X4))|(![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>(app(cons(X6,nil),X7)!=X3|app(X7,cons(X6,nil))!=X4)))&neq(X4,nil))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.44  fof(c_0_14, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(ssItem(esk52_0)&(memberP(esk48_0,esk52_0)&~memberP(esk49_0,esk52_0))))&(nil=esk50_0|nil!=esk51_0))&((ssItem(esk53_0)|~neq(esk51_0,nil))&((ssList(esk54_0)|~neq(esk51_0,nil))&((app(cons(esk53_0,nil),esk54_0)=esk50_0|~neq(esk51_0,nil))&(app(esk54_0,cons(esk53_0,nil))=esk51_0|~neq(esk51_0,nil))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).
% 0.06/0.44  fof(c_0_15, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
% 0.06/0.44  cnf(c_0_16, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  cnf(c_0_17, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  cnf(c_0_18, negated_conjecture, (app(cons(esk53_0,nil),esk54_0)=esk50_0|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  cnf(c_0_19, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
% 0.06/0.44  cnf(c_0_20, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.06/0.44  cnf(c_0_21, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_16, c_0_17])).
% 0.06/0.44  fof(c_0_22, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
% 0.06/0.44  cnf(c_0_23, negated_conjecture, (ssItem(esk53_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  cnf(c_0_24, negated_conjecture, (ssList(esk54_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  fof(c_0_25, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 0.06/0.44  cnf(c_0_26, negated_conjecture, (app(cons(esk53_0,nil),esk54_0)=esk50_0|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21])])).
% 0.06/0.44  cnf(c_0_27, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.44  cnf(c_0_28, negated_conjecture, (esk51_0=nil|ssItem(esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_19]), c_0_20]), c_0_21])])).
% 0.06/0.44  cnf(c_0_29, negated_conjecture, (esk51_0=nil|ssList(esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_19]), c_0_20]), c_0_21])])).
% 0.06/0.44  cnf(c_0_30, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
% 0.06/0.44  cnf(c_0_31, negated_conjecture, (cons(esk53_0,esk54_0)=esk50_0|esk51_0=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29])).
% 0.06/0.44  cnf(c_0_32, negated_conjecture, (app(esk54_0,cons(esk53_0,nil))=esk51_0|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  fof(c_0_33, plain, (![X4]:![X5]:![X6]:(((((~memberP(app(X5,X6),X4)|(memberP(X5,X4)|memberP(X6,X4)))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&(((((~memberP(X5,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
% 0.06/0.44  cnf(c_0_34, negated_conjecture, (memberP(esk48_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  cnf(c_0_35, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  cnf(c_0_36, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  cnf(c_0_37, negated_conjecture, (cons(esk53_0,app(esk54_0,X1))=app(esk50_0,X1)|esk51_0=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_28]), c_0_29])).
% 0.06/0.44  cnf(c_0_38, negated_conjecture, (app(esk54_0,cons(esk53_0,nil))=esk51_0|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_19]), c_0_20]), c_0_21])])).
% 0.06/0.44  fof(c_0_39, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
% 0.06/0.44  cnf(c_0_40, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
% 0.06/0.44  cnf(c_0_41, negated_conjecture, (memberP(esk50_0,esk52_0)), inference(rw,[status(thm)],[c_0_34, c_0_35])).
% 0.06/0.44  cnf(c_0_42, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_36, c_0_35])).
% 0.06/0.44  cnf(c_0_43, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  cnf(c_0_44, negated_conjecture, (app(esk50_0,cons(esk53_0,nil))=cons(esk53_0,esk51_0)|esk51_0=nil|~ssList(cons(esk53_0,nil))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
% 0.06/0.44  cnf(c_0_45, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.44  fof(c_0_46, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
% 0.06/0.44  fof(c_0_47, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
% 0.06/0.44  fof(c_0_48, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 0.06/0.44  fof(c_0_49, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk3_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&((((ssList(esk4_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&(((app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6)))=X5|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(X9,cons(X6,X10))!=X5))|memberP(X5,X6))|~ssItem(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).
% 0.06/0.44  cnf(c_0_50, negated_conjecture, (memberP(app(esk50_0,X1),esk52_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_43])])).
% 0.06/0.44  cnf(c_0_51, negated_conjecture, (app(esk50_0,cons(esk53_0,nil))=cons(esk53_0,esk51_0)|esk51_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_20])]), c_0_28])).
% 0.06/0.44  cnf(c_0_52, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
% 0.06/0.44  cnf(c_0_53, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
% 0.06/0.44  fof(c_0_54, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])).
% 0.06/0.44  cnf(c_0_55, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
% 0.06/0.44  cnf(c_0_56, plain, (memberP(X1,X2)|~ssList(X1)|~ssItem(X2)|app(X3,cons(X2,X4))!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_49])).
% 0.06/0.44  fof(c_0_57, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
% 0.06/0.44  cnf(c_0_58, negated_conjecture, (esk51_0=nil|memberP(cons(esk53_0,esk51_0),esk52_0)|~ssList(cons(esk53_0,nil))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
% 0.06/0.44  cnf(c_0_59, negated_conjecture, (~memberP(esk49_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  cnf(c_0_60, plain, (X1=nil|app(X1,X2)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_20])])).
% 0.06/0.44  cnf(c_0_61, negated_conjecture, (nil=esk50_0|nil!=esk51_0), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  cnf(c_0_62, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
% 0.06/0.44  cnf(c_0_63, negated_conjecture, (hd(esk50_0)=esk53_0|esk51_0=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_31]), c_0_28]), c_0_29])).
% 0.06/0.44  cnf(c_0_64, plain, (memberP(app(X1,cons(X2,X3)),X2)|~ssList(app(X1,cons(X2,X3)))|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(er,[status(thm)],[c_0_56])).
% 0.06/0.44  cnf(c_0_65, plain, (memberP(X3,X1)|X1=X2|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~memberP(cons(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
% 0.06/0.44  cnf(c_0_66, negated_conjecture, (esk51_0=nil|memberP(cons(esk53_0,esk51_0),esk52_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_45]), c_0_20])]), c_0_28])).
% 0.06/0.44  cnf(c_0_67, negated_conjecture, (~memberP(esk51_0,esk52_0)), inference(rw,[status(thm)],[c_0_59, c_0_17])).
% 0.06/0.44  cnf(c_0_68, negated_conjecture, (esk50_0=nil|cons(esk53_0,esk51_0)!=cons(esk53_0,nil)|~ssList(cons(esk53_0,nil))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_51]), c_0_42])]), c_0_61])).
% 0.06/0.44  cnf(c_0_69, negated_conjecture, (esk50_0=nil|ssItem(esk53_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_42])]), c_0_61])).
% 0.06/0.44  cnf(c_0_70, negated_conjecture, (esk51_0=nil|memberP(esk51_0,esk53_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_38]), c_0_21]), c_0_20])]), c_0_28]), c_0_29])).
% 0.06/0.44  cnf(c_0_71, negated_conjecture, (esk51_0=nil|esk53_0=esk52_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_21]), c_0_43])]), c_0_67]), c_0_28])).
% 0.06/0.44  cnf(c_0_72, negated_conjecture, (esk50_0=nil|cons(esk53_0,esk51_0)!=cons(esk53_0,nil)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_45]), c_0_20])]), c_0_69])).
% 0.06/0.44  cnf(c_0_73, negated_conjecture, (esk51_0=nil), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_67])).
% 0.06/0.44  cnf(c_0_74, negated_conjecture, (esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_73])])).
% 0.06/0.44  cnf(c_0_75, negated_conjecture, (~memberP(nil,esk52_0)), inference(rw,[status(thm)],[c_0_67, c_0_73])).
% 0.06/0.44  cnf(c_0_76, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_74]), c_0_75]), ['proof']).
% 0.06/0.44  # SZS output end CNFRefutation
% 0.06/0.44  # Training examples: 0 positive, 0 negative
% 0.06/0.45  
% 0.06/0.45  # -------------------------------------------------
% 0.06/0.45  # User time                : 0.208 s
% 0.06/0.45  # System time              : 0.010 s
% 0.06/0.45  # Total time               : 0.218 s
% 0.06/0.45  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
