<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC057-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC057-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC057-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC057-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC057-1.p">SWC057-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n158.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:28 EDT 2016

% Result     : Unsatisfiable 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of clauses        :   40 ( 381 expanded)
%              Number of leaves         :   19 ( 235 expanded)
%              Depth                    :    7
%              Number of atoms          :   97 ( 717 expanded)
%              Number of equality atoms :   28 ( 249 expanded)
%              Maximal clause size      :    6 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( nil = sk2
    | ~ ssList(X1)
    | ~ neq(X1,nil)
    | ~ segmentP(sk2,X1)
    | ~ segmentP(sk1,X1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( nil = sk4
    | neq(sk3,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( nil = sk3
    | frontsegP(sk4,sk3) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_14)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( segmentP(X1,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause57)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( nil = X1
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause118)).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( nil != sk1
    | ~ ssList(X1)
    | ~ neq(X1,nil)
    | ~ segmentP(sk2,X1)
    | ~ segmentP(sk1,X1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( neq(sk2,nil)
    | nil != sk1 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( segmentP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause56)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( segmentP(nil,X1)
    | nil != X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause79)).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( sk2 = nil
    | ~ neq(X1,nil)
    | ~ segmentP(sk3,X1)
    | ~ segmentP(sk2,X1)
    | ~ ssList(X1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( sk2 = nil
    | neq(sk3,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( sk3 = nil
    | frontsegP(sk2,sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>,<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( sk2 = nil
    | ~ segmentP(sk3,sk3)
    | ~ segmentP(sk2,sk3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_20]),c_0_21])])).

<A NAME="c_0_25"></A>cnf(c_0_25,plain,
    ( segmentP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_7]),c_0_8])]),<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( app(sk3,skaf45(sk2,sk3)) = sk2
    | sk3 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_23]),c_0_21]),c_0_11])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( sk2 = nil
    | ~ segmentP(sk2,sk3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_12]),c_0_21])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( sk3 = nil
    | segmentP(sk2,sk3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26]),c_0_13]),c_0_21])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( sk3 = nil
    | sk2 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_26]),c_0_13]),c_0_21])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( sk3 != nil
    | ~ neq(X1,nil)
    | ~ segmentP(sk3,X1)
    | ~ segmentP(sk2,X1)
    | ~ ssList(X1) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_1]),<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( sk3 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_28]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( neq(sk2,nil)
    | sk3 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( ~ neq(X1,nil)
    | ~ segmentP(nil,X1)
    | ~ segmentP(sk2,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31]),c_0_31])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( neq(sk2,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_31])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( sk2 = nil
    | ~ segmentP(sk2,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( ~ segmentP(nil,sk2)
    | ~ segmentP(sk2,sk2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_34]),c_0_11])])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( sk2 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_17]),c_0_11])])).

<A NAME="c_0_38"></A>cnf(c_0_38,plain,
    ( segmentP(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_18]),c_0_8])])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_37]),c_0_38]),c_0_37]),c_0_37]),c_0_38])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC057-1 : TPTP v6.4.0. Released v2.4.0.
% 0.02/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n158.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:34:52 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.38  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.06/0.38  # and selection function SelectCQArNTEqFirst.
% 0.06/0.38  #
% 0.06/0.38  # Preprocessing time       : 0.009 s
% 0.06/0.38  # Presaturation interreduction done
% 0.06/0.38  
% 0.06/0.38  # Proof found!
% 0.06/0.38  # SZS status Unsatisfiable
% 0.06/0.38  # SZS output start CNFRefutation
% 0.06/0.38  cnf(c_0_0, negated_conjecture, (nil=sk2|~ssList(X1)|~neq(X1,nil)|~segmentP(sk2,X1)|~segmentP(sk1,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
% 0.06/0.38  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 0.06/0.38  cnf(c_0_2, negated_conjecture, (nil=sk4|neq(sk3,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_11)).
% 0.06/0.38  cnf(c_0_3, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 0.06/0.38  cnf(c_0_4, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 0.06/0.38  cnf(c_0_5, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause173)).
% 0.06/0.38  cnf(c_0_6, negated_conjecture, (nil=sk3|frontsegP(sk4,sk3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_14)).
% 0.06/0.38  cnf(c_0_7, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause74)).
% 0.06/0.38  cnf(c_0_8, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.06/0.38  cnf(c_0_9, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.06/0.38  cnf(c_0_10, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause132)).
% 0.06/0.38  cnf(c_0_11, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
% 0.06/0.38  cnf(c_0_12, axiom, (segmentP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause57)).
% 0.06/0.38  cnf(c_0_13, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause51)).
% 0.06/0.38  cnf(c_0_14, axiom, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause118)).
% 0.06/0.38  cnf(c_0_15, negated_conjecture, (nil!=sk1|~ssList(X1)|~neq(X1,nil)|~segmentP(sk2,X1)|~segmentP(sk1,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
% 0.06/0.38  cnf(c_0_16, negated_conjecture, (neq(sk2,nil)|nil!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
% 0.06/0.38  cnf(c_0_17, axiom, (segmentP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause56)).
% 0.06/0.38  cnf(c_0_18, axiom, (segmentP(nil,X1)|nil!=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause79)).
% 0.06/0.38  cnf(c_0_19, negated_conjecture, (sk2=nil|~neq(X1,nil)|~segmentP(sk3,X1)|~segmentP(sk2,X1)|~ssList(X1)), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 0.06/0.38  cnf(c_0_20, negated_conjecture, (sk2=nil|neq(sk3,nil)), inference(rw,[status(thm)],[c_0_2, c_0_3])).
% 0.06/0.38  cnf(c_0_21, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_4, c_0_1])).
% 0.06/0.38  cnf(c_0_22, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(er,[status(thm)],[c_0_5])).
% 0.06/0.38  cnf(c_0_23, negated_conjecture, (sk3=nil|frontsegP(sk2,sk3)), inference(rw,[status(thm)],[c_0_6, c_0_3])).
% 0.06/0.38  cnf(c_0_24, negated_conjecture, (sk2=nil|~segmentP(sk3,sk3)|~segmentP(sk2,sk3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
% 0.06/0.38  cnf(c_0_25, plain, (segmentP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_7]), c_0_8])]), c_0_9])).
% 0.06/0.38  cnf(c_0_26, negated_conjecture, (app(sk3,skaf45(sk2,sk3))=sk2|sk3=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_23]), c_0_21]), c_0_11])])).
% 0.06/0.38  cnf(c_0_27, negated_conjecture, (sk2=nil|~segmentP(sk2,sk3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_12]), c_0_21])])).
% 0.06/0.38  cnf(c_0_28, negated_conjecture, (sk3=nil|segmentP(sk2,sk3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_13]), c_0_21])])).
% 0.06/0.38  cnf(c_0_29, negated_conjecture, (sk3=nil|sk2!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_26]), c_0_13]), c_0_21])])).
% 0.06/0.38  cnf(c_0_30, negated_conjecture, (sk3!=nil|~neq(X1,nil)|~segmentP(sk3,X1)|~segmentP(sk2,X1)|~ssList(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_1]), c_0_1])).
% 0.06/0.38  cnf(c_0_31, negated_conjecture, (sk3=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
% 0.06/0.38  cnf(c_0_32, negated_conjecture, (neq(sk2,nil)|sk3!=nil), inference(rw,[status(thm)],[c_0_16, c_0_1])).
% 0.06/0.38  cnf(c_0_33, negated_conjecture, (~neq(X1,nil)|~segmentP(nil,X1)|~segmentP(sk2,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31]), c_0_31])])).
% 0.06/0.38  cnf(c_0_34, negated_conjecture, (neq(sk2,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_31])])).
% 0.06/0.38  cnf(c_0_35, negated_conjecture, (sk2=nil|~segmentP(sk2,nil)), inference(rw,[status(thm)],[c_0_27, c_0_31])).
% 0.06/0.38  cnf(c_0_36, negated_conjecture, (~segmentP(nil,sk2)|~segmentP(sk2,sk2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_11])])).
% 0.06/0.38  cnf(c_0_37, negated_conjecture, (sk2=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_17]), c_0_11])])).
% 0.06/0.38  cnf(c_0_38, plain, (segmentP(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_18]), c_0_8])])).
% 0.06/0.38  cnf(c_0_39, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_37]), c_0_37]), c_0_38])]), ['proof']).
% 0.06/0.38  # SZS output end CNFRefutation
% 0.06/0.38  # Training examples: 0 positive, 0 negative
% 0.06/0.38  
% 0.06/0.38  # -------------------------------------------------
% 0.06/0.38  # User time                : 0.150 s
% 0.06/0.38  # System time              : 0.008 s
% 0.06/0.38  # Total time               : 0.158 s
% 0.06/0.38  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
