<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC059+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC059+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC059+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC059%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC059+1.p">SWC059+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n180.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:28 EDT 2016

% Result     : Theorem 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of formulae       :   33 ( 212 expanded)
%              Number of clauses        :   25 (  72 expanded)
%              Number of leaves         :    4 (  49 expanded)
%              Depth                    :    9
%              Number of atoms          :  119 (2812 expanded)
%              Number of equality atoms :   36 ( 911 expanded)
%              Maximal formula depth    :   20 (   4 average)
%              Maximal clause size      :   26 (   2 average)
%              Maximal term depth       :    1 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ( nil != X3
                      & nil = X4 )
                    | ( neq(X4,nil)
                      & ( ~ neq(X3,nil)
                        | ~ segmentP(X4,X3) ) )
                    | ( ( nil != X2
                        | nil = X1 )
                      & ( ~ neq(X2,nil)
                        | ? [X5] :
                            ( ssList(X5)
                            & neq(X5,nil)
                            & segmentP(X2,X5)
                            & segmentP(X1,X5) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,nil) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax57)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax55)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_4"></A>fof(c_0_4,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ( nil != X3
                        & nil = X4 )
                      | ( neq(X4,nil)
                        & ( ~ neq(X3,nil)
                          | ~ segmentP(X4,X3) ) )
                      | ( ( nil != X2
                          | nil = X1 )
                        & ( ~ neq(X2,nil)
                          | ? [X5] :
                              ( ssList(X5)
                              & neq(X5,nil)
                              & segmentP(X2,X5)
                              & segmentP(X1,X5) ) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_5"></A>fof(c_0_5,negated_conjecture,(
    ! [X10] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ( nil = esk50_0
        | nil != esk51_0 )
      & ( neq(esk50_0,nil)
        | ~ neq(esk51_0,nil) )
      & ( segmentP(esk51_0,esk50_0)
        | ~ neq(esk51_0,nil) )
      & ( neq(esk49_0,nil)
        | nil = esk49_0 )
      & ( ~ ssList(X10)
        | ~ neq(X10,nil)
        | ~ segmentP(esk49_0,X10)
        | ~ segmentP(esk48_0,X10)
        | nil = esk49_0 )
      & ( neq(esk49_0,nil)
        | nil != esk48_0 )
      & ( ~ ssList(X10)
        | ~ neq(X10,nil)
        | ~ segmentP(esk49_0,X10)
        | ~ segmentP(esk48_0,X10)
        | nil != esk48_0 ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( neq(esk50_0,nil)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( segmentP(esk51_0,esk50_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_10"></A>fof(c_0_10,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).

<A NAME="c_0_11"></A>fof(c_0_11,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( neq(esk48_0,nil)
    | ~ neq(esk49_0,nil) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>,c_0_7]),<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( nil = esk49_0
    | neq(esk49_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( segmentP(esk49_0,esk48_0)
    | ~ neq(esk49_0,nil) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_8]),c_0_7]),<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( nil != esk48_0
    | ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,plain,
    ( segmentP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( nil = esk50_0
    | nil != esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( nil = esk49_0
    | ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,plain,
    ( segmentP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( esk49_0 = nil
    | neq(esk48_0,nil) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( esk49_0 = nil
    | segmentP(esk49_0,esk48_0) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( esk48_0 != nil
    | ~ segmentP(esk49_0,nil)
    | ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_16]),c_0_17]),c_0_18])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( esk48_0 = nil
    | esk49_0 != nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_7]),<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( esk49_0 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_21]),c_0_18])]),c_0_22]),<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( esk48_0 != nil
    | ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_16]),c_0_25])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( esk48_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_27])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( neq(esk49_0,nil)
    | nil != esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_29])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_27]),c_0_29])]),<A HREF="#c_0_31">c_0_31</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC059+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n180.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:34:39 CDT 2016
% 0.02/0.22  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.02/0.25  # and selection function SelectCQPrecWNTNp.
% 0.02/0.25  #
% 0.02/0.25  # Preprocessing time       : 0.014 s
% 0.02/0.25  # Presaturation interreduction done
% 0.02/0.25  
% 0.02/0.25  # Proof found!
% 0.02/0.25  # SZS status Theorem
% 0.02/0.25  # SZS output start CNFRefutation
% 0.02/0.25  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|(nil!=X3&nil=X4))|(neq(X4,nil)&(~(neq(X3,nil))|~(segmentP(X4,X3)))))|((nil!=X2|nil=X1)&(~(neq(X2,nil))|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.02/0.25  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>segmentP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax57)).
% 0.02/0.25  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>segmentP(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax55)).
% 0.02/0.25  fof(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.02/0.25  fof(c_0_4, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|(nil!=X3&nil=X4))|(neq(X4,nil)&(~neq(X3,nil)|~segmentP(X4,X3))))|((nil!=X2|nil=X1)&(~neq(X2,nil)|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.02/0.25  fof(c_0_5, negated_conjecture, (![X10]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(nil=esk50_0|nil!=esk51_0))&((neq(esk50_0,nil)|~neq(esk51_0,nil))&(segmentP(esk51_0,esk50_0)|~neq(esk51_0,nil))))&(((neq(esk49_0,nil)|nil=esk49_0)&((((~ssList(X10)|~neq(X10,nil))|~segmentP(esk49_0,X10))|~segmentP(esk48_0,X10))|nil=esk49_0))&((neq(esk49_0,nil)|nil!=esk48_0)&((((~ssList(X10)|~neq(X10,nil))|~segmentP(esk49_0,X10))|~segmentP(esk48_0,X10))|nil!=esk48_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
% 0.02/0.25  cnf(c_0_6, negated_conjecture, (neq(esk50_0,nil)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.25  cnf(c_0_7, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.25  cnf(c_0_8, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.25  cnf(c_0_9, negated_conjecture, (segmentP(esk51_0,esk50_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.25  fof(c_0_10, plain, (![X2]:(~ssList(X2)|segmentP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
% 0.02/0.25  fof(c_0_11, plain, (![X2]:(~ssList(X2)|segmentP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
% 0.02/0.25  cnf(c_0_12, negated_conjecture, (neq(esk48_0,nil)|~neq(esk49_0,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_7]), c_0_8])).
% 0.02/0.25  cnf(c_0_13, negated_conjecture, (nil=esk49_0|neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.25  cnf(c_0_14, negated_conjecture, (segmentP(esk49_0,esk48_0)|~neq(esk49_0,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_8]), c_0_7]), c_0_8])).
% 0.02/0.25  cnf(c_0_15, negated_conjecture, (nil!=esk48_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.25  cnf(c_0_16, plain, (segmentP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.25  cnf(c_0_17, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.02/0.25  cnf(c_0_18, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.25  cnf(c_0_19, negated_conjecture, (nil=esk50_0|nil!=esk51_0), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.25  cnf(c_0_20, negated_conjecture, (nil=esk49_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.25  cnf(c_0_21, plain, (segmentP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.02/0.25  cnf(c_0_22, negated_conjecture, (esk49_0=nil|neq(esk48_0,nil)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
% 0.02/0.25  cnf(c_0_23, negated_conjecture, (esk49_0=nil|segmentP(esk49_0,esk48_0)), inference(spm,[status(thm)],[c_0_14, c_0_13])).
% 0.02/0.25  cnf(c_0_24, negated_conjecture, (esk48_0!=nil|~segmentP(esk49_0,nil)|~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])])).
% 0.02/0.25  cnf(c_0_25, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.25  cnf(c_0_26, negated_conjecture, (esk48_0=nil|esk49_0!=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_7]), c_0_8])).
% 0.02/0.25  cnf(c_0_27, negated_conjecture, (esk49_0=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_18])]), c_0_22]), c_0_23])).
% 0.02/0.25  cnf(c_0_28, negated_conjecture, (esk48_0!=nil|~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_16]), c_0_25])])).
% 0.02/0.25  cnf(c_0_29, negated_conjecture, (esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27])])).
% 0.02/0.25  cnf(c_0_30, negated_conjecture, (neq(esk49_0,nil)|nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.25  cnf(c_0_31, negated_conjecture, (~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_29])])).
% 0.02/0.25  cnf(c_0_32, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_27]), c_0_29])]), c_0_31]), ['proof']).
% 0.02/0.25  # SZS output end CNFRefutation
% 0.02/0.25  # Training examples: 0 positive, 0 negative
% 0.02/0.25  
% 0.02/0.25  # -------------------------------------------------
% 0.02/0.25  # User time                : 0.024 s
% 0.02/0.25  # System time              : 0.004 s
% 0.02/0.25  # Total time               : 0.028 s
% 0.02/0.25  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
