<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC093+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC093+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC093+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC093%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC093+1.p">SWC093+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n180.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:37 EDT 2016

% Result     : Theorem 55.45s
% Output     : CNFRefutation 55.45s
% Verified   : 
% Statistics : Number of formulae       :  102 (1714 expanded)
%              Number of clauses        :   68 ( 673 expanded)
%              Number of leaves         :   17 ( 394 expanded)
%              Depth                    :   22
%              Number of atoms          :  334 (16888 expanded)
%              Number of equality atoms :   58 (4626 expanded)
%              Maximal formula depth    :   20 (   4 average)
%              Maximal clause size      :   19 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ~ neq(X2,nil)
                    | ? [X5] :
                        ( ssList(X5)
                        & neq(X5,nil)
                        & segmentP(X2,X5)
                        & segmentP(X1,X5) )
                    | ( nil != X3
                      & nil = X4 )
                    | ( ! [X6] :
                          ( ssItem(X6)
                         => ( cons(X6,nil) != X3
                            | ~ memberP(X4,X6) ) )
                      & neq(X4,nil) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( frontsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X2,X3) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax5)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( frontsegP(X1,X2)
               => frontsegP(app(X1,X3),X2) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax43)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,nil) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax45)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( segmentP(X1,X2)
                   => segmentP(app(app(X3,X1),X4),X2) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax56)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,X1) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax55)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( memberP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(X3,cons(X2,X4)) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax3)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( ( segmentP(X1,X2)
                  & segmentP(X2,X3) )
               => segmentP(X1,X3) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax53)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) != X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax18)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_17"></A>fof(c_0_17,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ~ neq(X2,nil)
                      | ? [X5] :
                          ( ssList(X5)
                          & neq(X5,nil)
                          & segmentP(X2,X5)
                          & segmentP(X1,X5) )
                      | ( nil != X3
                        & nil = X4 )
                      | ( ! [X6] :
                            ( ssItem(X6)
                           => ( cons(X6,nil) != X3
                              | ~ memberP(X4,X6) ) )
                        & neq(X4,nil) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_18"></A>fof(c_0_18,negated_conjecture,(
    ! [X11] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & neq(esk49_0,nil)
      & ( ~ ssList(X11)
        | ~ neq(X11,nil)
        | ~ segmentP(esk49_0,X11)
        | ~ segmentP(esk48_0,X11) )
      & ( nil = esk50_0
        | nil != esk51_0 )
      & ( ssItem(esk52_0)
        | ~ neq(esk51_0,nil) )
      & ( cons(esk52_0,nil) = esk50_0
        | ~ neq(esk51_0,nil) )
      & ( memberP(esk51_0,esk52_0)
        | ~ neq(esk51_0,nil) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( neq(esk49_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_21"></A>fof(c_0_21,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk6_2(X4,X5))
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(X5,esk6_2(X4,X5)) = X4
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X5,X7) != X4
        | frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

<A NAME="c_0_22"></A>fof(c_0_22,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_23"></A>fof(c_0_23,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( cons(esk52_0,nil) = esk50_0
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( neq(esk51_0,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( ssItem(esk52_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,plain,
    ( frontsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X2,X3) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( cons(esk52_0,nil) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_25])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( ssItem(esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_25])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_27]),<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( app(esk50_0,X1) = cons(esk52_0,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_30]),c_0_31])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,<A HREF="#c_0_33">c_0_33</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( frontsegP(cons(esk52_0,X1),esk50_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_35]),c_0_36])])).

<A NAME="c_0_38"></A>cnf(c_0_38,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_39"></A>cnf(c_0_39,plain,
    ( app(X2,esk6_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( frontsegP(esk50_0,esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_30]),c_0_38])])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( ssList(esk6_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_42"></A>fof(c_0_42,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( app(esk50_0,esk6_2(esk50_0,esk50_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_40]),c_0_36])])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( ssList(esk6_2(esk50_0,esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_40]),c_0_36])])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_46"></A>fof(c_0_46,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ frontsegP(X4,X5)
      | frontsegP(app(X4,X6),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).

<A NAME="c_0_47"></A>fof(c_0_47,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( cons(esk52_0,esk6_2(esk50_0,esk50_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_43]),c_0_44])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( tl(esk50_0) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_30]),c_0_38]),c_0_31])])).

<A NAME="c_0_50"></A>cnf(c_0_50,plain,
    ( frontsegP(app(X1,X2),X3)
    | ~ frontsegP(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,plain,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( esk6_2(esk50_0,esk50_0) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_48]),c_0_49]),c_0_44]),c_0_31])])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_51]),c_0_38])])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( app(esk50_0,nil) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,<A HREF="#c_0_52">c_0_52</A>])).

<A NAME="c_0_55"></A>fof(c_0_55,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( frontsegP(esk50_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_54]),c_0_38]),c_0_36])])).

<A NAME="c_0_57"></A>fof(c_0_57,plain,(
    ! [X5,X6,X7,X8] :
      ( ~ ssList(X5)
      | ~ ssList(X6)
      | ~ ssList(X7)
      | ~ ssList(X8)
      | ~ segmentP(X5,X6)
      | segmentP(app(app(X7,X5),X8),X6) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_58"></A>fof(c_0_58,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).

<A NAME="c_0_59"></A>cnf(c_0_59,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( app(nil,esk6_2(esk50_0,nil)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_56]),c_0_38]),c_0_36])])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( ssList(esk6_2(esk50_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_56]),c_0_38]),c_0_36])])).

<A NAME="c_0_62"></A>cnf(c_0_62,plain,
    ( segmentP(app(app(X1,X2),X3),X4)
    | ~ segmentP(X2,X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X4)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_63"></A>cnf(c_0_63,plain,
    ( segmentP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( esk6_2(esk50_0,nil) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,c_0_60]),c_0_61])])).

<A NAME="c_0_65"></A>fof(c_0_65,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])).

<A NAME="c_0_66"></A>fof(c_0_66,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk3_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ssList(esk4_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6))) = X5
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(X9,cons(X6,X10)) != X5
        | memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( memberP(esk51_0,esk52_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( app(nil,esk50_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>,<A HREF="#c_0_64">c_0_64</A>])).

<A NAME="c_0_71"></A>cnf(c_0_71,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( app(esk3_2(X1,X2),cons(X2,esk4_2(X1,X2))) = X1
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ memberP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,negated_conjecture,
    ( memberP(esk51_0,esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,c_0_25])])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_75"></A>cnf(c_0_75,plain,
    ( ssList(esk3_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ memberP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>])).

<A NAME="c_0_76"></A>fof(c_0_76,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ segmentP(X4,X5)
      | ~ segmentP(X5,X6)
      | segmentP(X4,X6) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).

<A NAME="c_0_77"></A>cnf(c_0_77,negated_conjecture,
    ( segmentP(app(esk50_0,X1),esk50_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>,c_0_70]),c_0_36]),c_0_38])])).

<A NAME="c_0_78"></A>cnf(c_0_78,plain,
    ( segmentP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>,c_0_71]),c_0_38])]),<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( app(esk3_2(esk51_0,esk52_0),cons(esk52_0,esk4_2(esk51_0,esk52_0))) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_73]),c_0_74]),c_0_31])])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( ssList(esk3_2(esk51_0,esk52_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_73]),c_0_74]),c_0_31])])).

<A NAME="c_0_81"></A>fof(c_0_81,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).

<A NAME="c_0_82"></A>cnf(c_0_82,plain,
    ( ssList(esk4_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ memberP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,plain,
    ( segmentP(X1,X2)
    | ~ segmentP(X3,X2)
    | ~ segmentP(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( segmentP(cons(esk52_0,X1),esk50_0)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( ssList(cons(esk52_0,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_35]),c_0_36])])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( segmentP(esk51_0,cons(esk52_0,esk4_2(esk51_0,esk52_0)))
    | ~ ssList(cons(esk52_0,esk4_2(esk51_0,esk52_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_79]),c_0_80])])).

<A NAME="c_0_88"></A>cnf(c_0_88,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( ssList(esk4_2(esk51_0,esk52_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,c_0_73]),c_0_74]),c_0_31])])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( ~ segmentP(esk50_0,X1)
    | ~ segmentP(esk51_0,X1)
    | ~ ssList(X1)
    | ~ neq(X1,nil) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_33]),<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_91"></A>cnf(c_0_91,negated_conjecture,
    ( segmentP(X1,esk50_0)
    | ~ segmentP(X1,cons(esk52_0,X2))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,c_0_85]),c_0_36])]),<A HREF="#c_0_86">c_0_86</A>])).

<A NAME="c_0_92"></A>cnf(c_0_92,negated_conjecture,
    ( segmentP(esk51_0,cons(esk52_0,esk4_2(esk51_0,esk52_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,c_0_88]),c_0_89]),c_0_31])])).

<A NAME="c_0_93"></A>fof(c_0_93,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) != X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).

<A NAME="c_0_94"></A>cnf(c_0_94,negated_conjecture,
    ( ~ segmentP(esk51_0,esk50_0)
    | ~ neq(esk50_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,c_0_63]),c_0_36])])).

<A NAME="c_0_95"></A>cnf(c_0_95,negated_conjecture,
    ( segmentP(esk51_0,esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_92]),c_0_74]),c_0_89])])).

<A NAME="c_0_96"></A>fof(c_0_96,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])).

<A NAME="c_0_97"></A>cnf(c_0_97,plain,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>])).

<A NAME="c_0_98"></A>cnf(c_0_98,negated_conjecture,
    ( ~ neq(esk50_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>,c_0_95])])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,negated_conjecture,
    ( esk50_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_30]),c_0_38]),c_0_31])])).

<A NAME="c_0_101"></A>cnf(c_0_101,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_98">c_0_98</A>,c_0_99]),c_0_38]),c_0_36])]),<A HREF="#c_0_100">c_0_100</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC093+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n180.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:35:38 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 55.45/55.64  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 55.45/55.64  # and selection function SelectCQPrecWNTNp.
% 55.45/55.64  #
% 55.45/55.64  # Preprocessing time       : 0.014 s
% 55.45/55.64  # Presaturation interreduction done
% 55.45/55.64  
% 55.45/55.64  # Proof found!
% 55.45/55.64  # SZS status Theorem
% 55.45/55.64  # SZS output start CNFRefutation
% 55.45/55.64  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~(neq(X2,nil)))|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))|(nil!=X3&nil=X4))|(![X6]:(ssItem(X6)=>(cons(X6,nil)!=X3|~(memberP(X4,X6))))&neq(X4,nil)))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 55.45/55.64  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax5)).
% 55.45/55.64  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax26)).
% 55.45/55.64  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax81)).
% 55.45/55.64  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 55.45/55.64  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax25)).
% 55.45/55.64  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(frontsegP(X1,X2)=>frontsegP(app(X1,X3),X2)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax43)).
% 55.45/55.64  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax45)).
% 55.45/55.64  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
% 55.45/55.64  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(segmentP(X1,X2)=>segmentP(app(app(X3,X1),X4),X2))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax56)).
% 55.45/55.64  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>segmentP(X1,X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax55)).
% 55.45/55.64  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax84)).
% 55.45/55.64  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>(memberP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(X3,cons(X2,X4))=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax3)).
% 55.45/55.64  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>((segmentP(X1,X2)&segmentP(X2,X3))=>segmentP(X1,X3)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax53)).
% 55.45/55.64  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
% 55.45/55.64  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax18)).
% 55.45/55.64  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax15)).
% 55.45/55.64  fof(c_0_17, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~neq(X2,nil))|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))|(nil!=X3&nil=X4))|(![X6]:(ssItem(X6)=>(cons(X6,nil)!=X3|~memberP(X4,X6)))&neq(X4,nil))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 55.45/55.64  fof(c_0_18, negated_conjecture, (![X11]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((((esk49_0=esk51_0&esk48_0=esk50_0)&neq(esk49_0,nil))&(((~ssList(X11)|~neq(X11,nil))|~segmentP(esk49_0,X11))|~segmentP(esk48_0,X11)))&(nil=esk50_0|nil!=esk51_0))&((ssItem(esk52_0)|~neq(esk51_0,nil))&((cons(esk52_0,nil)=esk50_0|~neq(esk51_0,nil))&(memberP(esk51_0,esk52_0)|~neq(esk51_0,nil)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])])).
% 55.45/55.64  cnf(c_0_19, negated_conjecture, (neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_18])).
% 55.45/55.64  cnf(c_0_20, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_18])).
% 55.45/55.64  fof(c_0_21, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
% 55.45/55.64  fof(c_0_22, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 55.45/55.64  fof(c_0_23, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
% 55.45/55.64  cnf(c_0_24, negated_conjecture, (cons(esk52_0,nil)=esk50_0|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_18])).
% 55.45/55.64  cnf(c_0_25, negated_conjecture, (neq(esk51_0,nil)), inference(rw,[status(thm)],[c_0_19, c_0_20])).
% 55.45/55.64  cnf(c_0_26, negated_conjecture, (ssItem(esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_18])).
% 55.45/55.64  cnf(c_0_27, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
% 55.45/55.64  cnf(c_0_28, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 55.45/55.64  cnf(c_0_29, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 55.45/55.64  cnf(c_0_30, negated_conjecture, (cons(esk52_0,nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25])])).
% 55.45/55.64  cnf(c_0_31, negated_conjecture, (ssItem(esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_25])])).
% 55.45/55.64  cnf(c_0_32, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
% 55.45/55.64  cnf(c_0_33, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_18])).
% 55.45/55.64  cnf(c_0_34, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_27]), c_0_28])).
% 55.45/55.64  cnf(c_0_35, negated_conjecture, (app(esk50_0,X1)=cons(esk52_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
% 55.45/55.64  cnf(c_0_36, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_32, c_0_33])).
% 55.45/55.64  cnf(c_0_37, negated_conjecture, (frontsegP(cons(esk52_0,X1),esk50_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
% 55.45/55.64  cnf(c_0_38, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 55.45/55.64  cnf(c_0_39, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
% 55.45/55.64  cnf(c_0_40, negated_conjecture, (frontsegP(esk50_0,esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_30]), c_0_38])])).
% 55.45/55.64  cnf(c_0_41, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
% 55.45/55.64  fof(c_0_42, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 55.45/55.64  cnf(c_0_43, negated_conjecture, (app(esk50_0,esk6_2(esk50_0,esk50_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_36])])).
% 55.45/55.64  cnf(c_0_44, negated_conjecture, (ssList(esk6_2(esk50_0,esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_40]), c_0_36])])).
% 55.45/55.64  cnf(c_0_45, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
% 55.45/55.64  fof(c_0_46, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~frontsegP(X4,X5)|frontsegP(app(X4,X6),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
% 55.45/55.64  fof(c_0_47, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
% 55.45/55.64  cnf(c_0_48, negated_conjecture, (cons(esk52_0,esk6_2(esk50_0,esk50_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_43]), c_0_44])])).
% 55.45/55.64  cnf(c_0_49, negated_conjecture, (tl(esk50_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_30]), c_0_38]), c_0_31])])).
% 55.45/55.64  cnf(c_0_50, plain, (frontsegP(app(X1,X2),X3)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
% 55.45/55.64  cnf(c_0_51, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
% 55.45/55.64  cnf(c_0_52, negated_conjecture, (esk6_2(esk50_0,esk50_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_48]), c_0_49]), c_0_44]), c_0_31])])).
% 55.45/55.64  cnf(c_0_53, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_38])])).
% 55.45/55.64  cnf(c_0_54, negated_conjecture, (app(esk50_0,nil)=esk50_0), inference(rw,[status(thm)],[c_0_43, c_0_52])).
% 55.45/55.64  fof(c_0_55, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
% 55.45/55.64  cnf(c_0_56, negated_conjecture, (frontsegP(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_38]), c_0_36])])).
% 55.45/55.64  fof(c_0_57, plain, (![X5]:![X6]:![X7]:![X8]:(~ssList(X5)|(~ssList(X6)|(~ssList(X7)|(~ssList(X8)|(~segmentP(X5,X6)|segmentP(app(app(X7,X5),X8),X6))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 55.45/55.64  fof(c_0_58, plain, (![X2]:(~ssList(X2)|segmentP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
% 55.45/55.64  cnf(c_0_59, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
% 55.45/55.64  cnf(c_0_60, negated_conjecture, (app(nil,esk6_2(esk50_0,nil))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_56]), c_0_38]), c_0_36])])).
% 55.45/55.64  cnf(c_0_61, negated_conjecture, (ssList(esk6_2(esk50_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_56]), c_0_38]), c_0_36])])).
% 55.45/55.64  cnf(c_0_62, plain, (segmentP(app(app(X1,X2),X3),X4)|~segmentP(X2,X4)|~ssList(X3)|~ssList(X1)|~ssList(X4)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
% 55.45/55.64  cnf(c_0_63, plain, (segmentP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
% 55.45/55.64  cnf(c_0_64, negated_conjecture, (esk6_2(esk50_0,nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])])).
% 55.45/55.64  fof(c_0_65, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])).
% 55.45/55.64  fof(c_0_66, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk3_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&((((ssList(esk4_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&(((app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6)))=X5|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(X9,cons(X6,X10))!=X5))|memberP(X5,X6))|~ssItem(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).
% 55.45/55.64  cnf(c_0_67, negated_conjecture, (memberP(esk51_0,esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_18])).
% 55.45/55.64  cnf(c_0_68, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
% 55.45/55.64  cnf(c_0_69, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
% 55.45/55.64  cnf(c_0_70, negated_conjecture, (app(nil,esk50_0)=esk50_0), inference(rw,[status(thm)],[c_0_60, c_0_64])).
% 55.45/55.64  cnf(c_0_71, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
% 55.45/55.64  cnf(c_0_72, plain, (app(esk3_2(X1,X2),cons(X2,esk4_2(X1,X2)))=X1|~ssList(X1)|~ssItem(X2)|~memberP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
% 55.45/55.64  cnf(c_0_73, negated_conjecture, (memberP(esk51_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_25])])).
% 55.45/55.64  cnf(c_0_74, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_68, c_0_20])).
% 55.45/55.64  cnf(c_0_75, plain, (ssList(esk3_2(X1,X2))|~ssList(X1)|~ssItem(X2)|~memberP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
% 55.45/55.64  fof(c_0_76, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|((~segmentP(X4,X5)|~segmentP(X5,X6))|segmentP(X4,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
% 55.45/55.64  cnf(c_0_77, negated_conjecture, (segmentP(app(esk50_0,X1),esk50_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_36]), c_0_38])])).
% 55.45/55.64  cnf(c_0_78, plain, (segmentP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_71]), c_0_38])]), c_0_28])).
% 55.45/55.64  cnf(c_0_79, negated_conjecture, (app(esk3_2(esk51_0,esk52_0),cons(esk52_0,esk4_2(esk51_0,esk52_0)))=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74]), c_0_31])])).
% 55.45/55.64  cnf(c_0_80, negated_conjecture, (ssList(esk3_2(esk51_0,esk52_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_73]), c_0_74]), c_0_31])])).
% 55.45/55.64  fof(c_0_81, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
% 55.45/55.64  cnf(c_0_82, plain, (ssList(esk4_2(X1,X2))|~ssList(X1)|~ssItem(X2)|~memberP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
% 55.45/55.64  cnf(c_0_83, negated_conjecture, (~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
% 55.45/55.64  cnf(c_0_84, plain, (segmentP(X1,X2)|~segmentP(X3,X2)|~segmentP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
% 55.45/55.64  cnf(c_0_85, negated_conjecture, (segmentP(cons(esk52_0,X1),esk50_0)|~ssList(X1)), inference(spm,[status(thm)],[c_0_77, c_0_35])).
% 55.45/55.64  cnf(c_0_86, negated_conjecture, (ssList(cons(esk52_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_35]), c_0_36])])).
% 55.45/55.64  cnf(c_0_87, negated_conjecture, (segmentP(esk51_0,cons(esk52_0,esk4_2(esk51_0,esk52_0)))|~ssList(cons(esk52_0,esk4_2(esk51_0,esk52_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80])])).
% 55.45/55.64  cnf(c_0_88, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
% 55.45/55.64  cnf(c_0_89, negated_conjecture, (ssList(esk4_2(esk51_0,esk52_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_73]), c_0_74]), c_0_31])])).
% 55.45/55.65  cnf(c_0_90, negated_conjecture, (~segmentP(esk50_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_33]), c_0_20])).
% 55.45/55.65  cnf(c_0_91, negated_conjecture, (segmentP(X1,esk50_0)|~segmentP(X1,cons(esk52_0,X2))|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_36])]), c_0_86])).
% 55.45/55.65  cnf(c_0_92, negated_conjecture, (segmentP(esk51_0,cons(esk52_0,esk4_2(esk51_0,esk52_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_89]), c_0_31])])).
% 55.45/55.65  fof(c_0_93, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
% 55.45/55.65  cnf(c_0_94, negated_conjecture, (~segmentP(esk51_0,esk50_0)|~neq(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_63]), c_0_36])])).
% 55.45/55.65  cnf(c_0_95, negated_conjecture, (segmentP(esk51_0,esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_74]), c_0_89])])).
% 55.45/55.65  fof(c_0_96, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])).
% 55.45/55.65  cnf(c_0_97, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_93])).
% 55.45/55.65  cnf(c_0_98, negated_conjecture, (~neq(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_95])])).
% 55.45/55.65  cnf(c_0_99, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_96])).
% 55.45/55.65  cnf(c_0_100, negated_conjecture, (esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_30]), c_0_38]), c_0_31])])).
% 55.45/55.65  cnf(c_0_101, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_38]), c_0_36])]), c_0_100]), ['proof']).
% 55.45/55.65  # SZS output end CNFRefutation
% 55.45/55.65  # Training examples: 0 positive, 0 negative
% 55.55/55.77  
% 55.55/55.77  # -------------------------------------------------
% 55.55/55.77  # User time                : 54.272 s
% 55.55/55.77  # System time              : 1.305 s
% 55.55/55.77  # Total time               : 55.577 s
% 55.55/55.77  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
