<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC026+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC026+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC026+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC026%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC026+1.p">SWC026+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n122.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:21 EDT 2016

% Result     : Theorem 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of formulae       :   13 (  54 expanded)
%              Number of clauses        :   10 (  18 expanded)
%              Number of leaves         :    1 (  12 expanded)
%              Depth                    :    7
%              Number of atoms          :   61 ( 578 expanded)
%              Number of equality atoms :   48 ( 433 expanded)
%              Maximal formula depth    :   17 (   5 average)
%              Maximal clause size      :   18 (   2 average)
%              Maximal term depth       :    1 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ( nil != X4
                      & nil = X3 )
                    | ( nil != X3
                      & nil = X4 )
                    | ( ( nil != X2
                        | nil = X1 )
                      & ( nil != X1
                        | nil = X2 ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ( nil != X4
                        & nil = X3 )
                      | ( nil != X3
                        & nil = X4 )
                      | ( ( nil != X2
                          | nil = X1 )
                        & ( nil != X1
                          | nil = X2 ) ) ) ) ) ) ) ),
    inference(assume_negation,[status(cth)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_2"></A>fof(c_0_2,negated_conjecture,
    ( ssList(esk48_0)
    & ssList(esk49_0)
    & ssList(esk50_0)
    & ssList(esk51_0)
    & esk49_0 = esk51_0
    & esk48_0 = esk50_0
    & ( nil = esk51_0
      | nil != esk50_0 )
    & ( nil = esk50_0
      | nil != esk51_0 )
    & ( nil = esk48_0
      | nil = esk49_0 )
    & ( nil != esk49_0
      | nil = esk49_0 )
    & ( nil = esk48_0
      | nil != esk48_0 )
    & ( nil != esk49_0
      | nil != esk48_0 ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( nil != esk48_0
    | nil != esk49_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( nil = esk51_0
    | nil != esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( esk48_0 != nil
    | esk51_0 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>,<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( nil = esk50_0
    | nil != esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( esk48_0 != nil ),
    inference(csr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,c_0_6]),<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( nil = esk49_0
    | nil = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( esk51_0 != nil ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_6]),<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_4]),c_0_9]),<A HREF="#c_0_11">c_0_11</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC026+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n122.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:34:23 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.02/0.25  # and selection function SelectCQPrecWNTNp.
% 0.02/0.25  #
% 0.02/0.25  # Preprocessing time       : 0.014 s
% 0.02/0.25  # Presaturation interreduction done
% 0.02/0.25  
% 0.02/0.25  # Proof found!
% 0.02/0.25  # SZS status Theorem
% 0.02/0.25  # SZS output start CNFRefutation
% 0.02/0.25  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|(nil!=X4&nil=X3))|(nil!=X3&nil=X4))|((nil!=X2|nil=X1)&(nil!=X1|nil=X2)))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.02/0.25  fof(c_0_1, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|(nil!=X4&nil=X3))|(nil!=X3&nil=X4))|((nil!=X2|nil=X1)&(nil!=X1|nil=X2))))))))), inference(assume_negation,[status(cth)],[c_0_0])).
% 0.02/0.25  fof(c_0_2, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(nil=esk51_0|nil!=esk50_0))&(nil=esk50_0|nil!=esk51_0))&(((nil=esk48_0|nil=esk49_0)&(nil!=esk49_0|nil=esk49_0))&((nil=esk48_0|nil!=esk48_0)&(nil!=esk49_0|nil!=esk48_0))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
% 0.02/0.25  cnf(c_0_3, negated_conjecture, (nil!=esk48_0|nil!=esk49_0), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.02/0.25  cnf(c_0_4, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.02/0.25  cnf(c_0_5, negated_conjecture, (nil=esk51_0|nil!=esk50_0), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.02/0.25  cnf(c_0_6, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.02/0.25  cnf(c_0_7, negated_conjecture, (esk48_0!=nil|esk51_0!=nil), inference(rw,[status(thm)],[c_0_3, c_0_4])).
% 0.02/0.25  cnf(c_0_8, negated_conjecture, (nil=esk50_0|nil!=esk51_0), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.02/0.25  cnf(c_0_9, negated_conjecture, (esk48_0!=nil), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_6]), c_0_7])).
% 0.02/0.25  cnf(c_0_10, negated_conjecture, (nil=esk49_0|nil=esk48_0), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.02/0.25  cnf(c_0_11, negated_conjecture, (esk51_0!=nil), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_6]), c_0_9])).
% 0.02/0.25  cnf(c_0_12, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_4]), c_0_9]), c_0_11]), ['proof']).
% 0.02/0.25  # SZS output end CNFRefutation
% 0.02/0.25  # Training examples: 0 positive, 0 negative
% 0.02/0.25  
% 0.02/0.25  # -------------------------------------------------
% 0.02/0.25  # User time                : 0.015 s
% 0.02/0.25  # System time              : 0.002 s
% 0.02/0.25  # Total time               : 0.018 s
% 0.02/0.25  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
