<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC348+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC348+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC348+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC348%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC348+1.p">SWC348+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n141.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:35 EDT 2016

% Result     : Theorem 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of formulae       :   38 (  98 expanded)
%              Number of clauses        :   25 (  38 expanded)
%              Number of leaves         :    7 (  23 expanded)
%              Depth                    :    9
%              Number of atoms          :  116 ( 640 expanded)
%              Number of equality atoms :   23 ( 121 expanded)
%              Maximal formula depth    :   16 (   4 average)
%              Maximal clause size      :   11 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ~ segmentP(X4,X3)
                    | ( ~ singletonP(X3)
                      & neq(X4,nil) )
                    | ( segmentP(X2,X1)
                      & strictorderedP(X1) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => strictorderedP(cons(X1,nil)) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax68)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( singletonP(X1)
      &lt;=> ? [X2] :
            ( ssItem(X2)
            & cons(X2,nil) = X1 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax4)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( segmentP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax58)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    strictorderedP(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax69)).

<A NAME="c_0_7"></A>fof(c_0_7,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ~ segmentP(X4,X3)
                      | ( ~ singletonP(X3)
                        & neq(X4,nil) )
                      | ( segmentP(X2,X1)
                        & strictorderedP(X1) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_8"></A>fof(c_0_8,negated_conjecture,
    ( ssList(esk48_0)
    & ssList(esk49_0)
    & ssList(esk50_0)
    & ssList(esk51_0)
    & esk49_0 = esk51_0
    & esk48_0 = esk50_0
    & segmentP(esk51_0,esk50_0)
    & ( singletonP(esk50_0)
      | ~ neq(esk51_0,nil) )
    & ( ~ segmentP(esk49_0,esk48_0)
      | ~ strictorderedP(esk48_0) ) ),
    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).

<A NAME="c_0_9"></A>fof(c_0_9,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | strictorderedP(cons(X2,nil)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).

<A NAME="c_0_10"></A>fof(c_0_10,plain,(
    ! [X3,X5] :
      ( ( ssItem(esk5_1(X3))
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( cons(esk5_1(X3),nil) = X3
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( ~ ssItem(X5)
        | cons(X5,nil) != X3
        | singletonP(X3)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( singletonP(esk50_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_13"></A>fof(c_0_13,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ~ strictorderedP(esk48_0)
    | ~ segmentP(esk49_0,esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( segmentP(esk51_0,esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,plain,
    ( strictorderedP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,plain,
    ( cons(esk5_1(X1),nil) = X1
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,plain,
    ( ssItem(esk5_1(X1))
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( singletonP(esk48_0)
    | ~ neq(esk51_0,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( ~ strictorderedP(esk48_0)
    | ~ segmentP(esk51_0,esk48_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( segmentP(esk51_0,esk48_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,plain,
    ( strictorderedP(X1)
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_19]),<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( esk51_0 = nil
    | singletonP(esk48_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_22]),c_0_23]),c_0_24])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( ~ strictorderedP(esk48_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26])])).

<A NAME="c_0_31"></A>fof(c_0_31,plain,(
    ! [X2] :
      ( ( ~ segmentP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | segmentP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( esk51_0 = nil ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_28]),c_0_29])]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ segmentP(nil,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( segmentP(nil,esk48_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( esk48_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_34]),c_0_29])])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( strictorderedP(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_35]),c_0_36])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC348+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n141.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:41:53 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.26  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.03/0.26  # and selection function SelectCQPrecWNTNp.
% 0.03/0.26  #
% 0.03/0.26  # Preprocessing time       : 0.014 s
% 0.03/0.26  # Presaturation interreduction done
% 0.03/0.26  
% 0.03/0.26  # Proof found!
% 0.03/0.26  # SZS status Theorem
% 0.03/0.26  # SZS output start CNFRefutation
% 0.03/0.26  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|~(segmentP(X4,X3)))|(~(singletonP(X3))&neq(X4,nil)))|(segmentP(X2,X1)&strictorderedP(X1)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.03/0.26  fof(c_0_1, axiom, (![X1]:(ssItem(X1)=>strictorderedP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax68)).
% 0.03/0.26  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)&lt;=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax4)).
% 0.03/0.26  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.03/0.26  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.03/0.26  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>(segmentP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax58)).
% 0.03/0.26  fof(c_0_6, axiom, (strictorderedP(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax69)).
% 0.03/0.26  fof(c_0_7, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|~segmentP(X4,X3))|(~singletonP(X3)&neq(X4,nil)))|(segmentP(X2,X1)&strictorderedP(X1))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.03/0.26  fof(c_0_8, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&segmentP(esk51_0,esk50_0))&(singletonP(esk50_0)|~neq(esk51_0,nil)))&(~segmentP(esk49_0,esk48_0)|~strictorderedP(esk48_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
% 0.03/0.26  fof(c_0_9, plain, (![X2]:(~ssItem(X2)|strictorderedP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
% 0.03/0.26  fof(c_0_10, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
% 0.03/0.26  cnf(c_0_11, negated_conjecture, (singletonP(esk50_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.03/0.26  cnf(c_0_12, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.03/0.26  fof(c_0_13, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
% 0.03/0.26  cnf(c_0_14, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.03/0.26  cnf(c_0_15, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.03/0.26  cnf(c_0_16, negated_conjecture, (~strictorderedP(esk48_0)|~segmentP(esk49_0,esk48_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.03/0.26  cnf(c_0_17, negated_conjecture, (segmentP(esk51_0,esk50_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.03/0.26  cnf(c_0_18, plain, (strictorderedP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.03/0.26  cnf(c_0_19, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  cnf(c_0_20, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  cnf(c_0_21, negated_conjecture, (singletonP(esk48_0)|~neq(esk51_0,nil)), inference(rw,[status(thm)],[c_0_11, c_0_12])).
% 0.03/0.26  cnf(c_0_22, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
% 0.03/0.26  cnf(c_0_23, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.03/0.26  cnf(c_0_24, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_14, c_0_15])).
% 0.03/0.26  cnf(c_0_25, negated_conjecture, (~strictorderedP(esk48_0)|~segmentP(esk51_0,esk48_0)), inference(rw,[status(thm)],[c_0_16, c_0_15])).
% 0.03/0.26  cnf(c_0_26, negated_conjecture, (segmentP(esk51_0,esk48_0)), inference(rw,[status(thm)],[c_0_17, c_0_12])).
% 0.03/0.26  cnf(c_0_27, plain, (strictorderedP(X1)|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
% 0.03/0.26  cnf(c_0_28, negated_conjecture, (esk51_0=nil|singletonP(esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24])])).
% 0.03/0.26  cnf(c_0_29, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.03/0.26  cnf(c_0_30, negated_conjecture, (~strictorderedP(esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_26])])).
% 0.03/0.26  fof(c_0_31, plain, (![X2]:(((~segmentP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|segmentP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 0.03/0.26  cnf(c_0_32, negated_conjecture, (esk51_0=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])]), c_0_30])).
% 0.03/0.26  cnf(c_0_33, plain, (nil=X1|~ssList(X1)|~segmentP(nil,X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
% 0.03/0.26  cnf(c_0_34, negated_conjecture, (segmentP(nil,esk48_0)), inference(rw,[status(thm)],[c_0_26, c_0_32])).
% 0.03/0.26  cnf(c_0_35, negated_conjecture, (esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_29])])).
% 0.03/0.26  cnf(c_0_36, plain, (strictorderedP(nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
% 0.03/0.26  cnf(c_0_37, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_35]), c_0_36])]), ['proof']).
% 0.03/0.26  # SZS output end CNFRefutation
% 0.03/0.26  # Training examples: 0 positive, 0 negative
% 0.03/0.26  
% 0.03/0.26  # -------------------------------------------------
% 0.03/0.26  # User time                : 0.026 s
% 0.03/0.26  # System time              : 0.004 s
% 0.03/0.26  # Total time               : 0.029 s
% 0.03/0.26  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
