<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC390-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC390-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC390-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC390-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC390-1.p">SWC390-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n098.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:45 EDT 2016

% Result     : Unsatisfiable 12.62s
% Output     : CNFRefutation 12.62s
% Verified   : 
% Statistics : Number of clauses        :  165 (7342 expanded)
%              Number of leaves         :   55 (4868 expanded)
%              Depth                    :   28
%              Number of atoms          :  448 (13916 expanded)
%              Number of equality atoms :  165 (5454 expanded)
%              Maximal clause size      :    7 (   3 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( memberP(sk4,sk5)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_3"></A>cnf(c_0_3,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_4"></A>cnf(c_0_4,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( ssItem(sk5)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
    | ~ memberP(X1,X2)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause169)).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( cons(sk5,nil) = sk3
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( ssList(skaf43(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause52)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( nil = sk3
    | nil != sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( X4 = X2
    | cons(X1,X2) != cons(X3,X4)
    | ~ ssItem(X3)
    | ~ ssItem(X1)
    | ~ ssList(X4)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause171)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause99)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( singletonP(X2)
    | cons(X1,nil) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause116)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause104)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause97)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( cons(skaf83(X1),skaf82(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause109)).

<A NAME="c_0_23"></A>cnf(c_0_23,axiom,
    ( ssList(skaf82(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause13)).

<A NAME="c_0_24"></A>cnf(c_0_24,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_25"></A>cnf(c_0_25,axiom,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause76)).

<A NAME="c_0_26"></A>cnf(c_0_26,axiom,
    ( frontsegP(X3,X1)
    | app(X1,X2) != X3
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause144)).

<A NAME="c_0_27"></A>cnf(c_0_27,axiom,
    ( X1 = X3
    | cons(X1,X2) != cons(X3,X4)
    | ~ ssItem(X3)
    | ~ ssItem(X1)
    | ~ ssList(X4)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause170)).

<A NAME="c_0_28"></A>cnf(c_0_28,axiom,
    ( cons(skaf44(X1),nil) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause101)).

<A NAME="c_0_29"></A>cnf(c_0_29,axiom,
    ( ssItem(skaf44(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause47)).

<A NAME="c_0_30"></A>cnf(c_0_30,axiom,
    ( ssItem(skaf83(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause12)).

<A NAME="c_0_31"></A>cnf(c_0_31,axiom,
    ( frontsegP(app(X1,X3),X2)
    | ~ frontsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause137)).

<A NAME="c_0_32"></A>cnf(c_0_32,axiom,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause60)).

<A NAME="c_0_33"></A>cnf(c_0_33,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_34"></A>cnf(c_0_34,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_35"></A>cnf(c_0_35,axiom,
    ( memberP(cons(X2,X3),X1)
    | X1 != X2
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause138)).

<A NAME="c_0_36"></A>cnf(c_0_36,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_37"></A>cnf(c_0_37,axiom,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause157)).

<A NAME="c_0_38"></A>cnf(c_0_38,axiom,
    ( frontsegP(X2,X4)
    | ~ frontsegP(cons(X1,X2),cons(X3,X4))
    | ~ ssList(X4)
    | ~ ssList(X2)
    | ~ ssItem(X3)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause174)).

<A NAME="c_0_39"></A>cnf(c_0_39,axiom,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause75)).

<A NAME="c_0_40"></A>cnf(c_0_40,axiom,
    ( app(X1,X2) = nil
    | nil != X1
    | nil != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause145)).

<A NAME="c_0_41"></A>cnf(c_0_41,axiom,
    ( frontsegP(nil,X1)
    | nil != X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause83)).

<A NAME="c_0_42"></A>cnf(c_0_42,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_43"></A>cnf(c_0_43,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_44"></A>cnf(c_0_44,axiom,
    ( frontsegP(X1,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause61)).

<A NAME="c_0_45"></A>cnf(c_0_45,axiom,
    ( nil = X2
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause119)).

<A NAME="c_0_46"></A>cnf(c_0_46,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_47"></A>cnf(c_0_47,axiom,
    ( cons(X1,X2) != nil
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause98)).

<A NAME="c_0_48"></A>cnf(c_0_48,axiom,
    ( X1 = X3
    | ~ frontsegP(cons(X1,X2),cons(X3,X4))
    | ~ ssList(X4)
    | ~ ssList(X2)
    | ~ ssItem(X3)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause176)).

<A NAME="c_0_49"></A>cnf(c_0_49,axiom,
    ( nil = X1
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause118)).

<A NAME="c_0_50"></A>cnf(c_0_50,axiom,
    ( X2 = X3
    | app(X1,X2) != app(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause150)).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( ~ segmentP(sk2,sk1)
    | ~ singletonP(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_52"></A>cnf(c_0_52,axiom,
    ( X1 != X2
    | ~ neq(X1,X2)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause117)).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( neq(sk2,nil)
    | ~ segmentP(sk2,sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_54"></A>cnf(c_0_54,axiom,
    ( segmentP(nil,X1)
    | nil != X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause79)).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( sk4 = nil
    | memberP(sk4,sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_3]),c_0_55]),c_0_4])])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( sk4 = nil
    | ssItem(sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,c_0_3]),c_0_55]),c_0_4])])).

<A NAME="c_0_58"></A>cnf(c_0_58,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( app(skaf42(sk4,sk5),cons(sk5,skaf43(sk5,sk4))) = sk4
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_56]),c_0_55])]),<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( cons(sk5,nil) = sk1
    | ~ neq(sk4,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( sk4 = nil
    | ~ ssList(app(sk4,cons(sk5,X1)))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(sk4,cons(sk5,X1))) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_59]),c_0_10]),c_0_11])]),<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( cons(sk5,nil) = sk1
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>,c_0_3]),c_0_55]),c_0_4])])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( sk4 = nil
    | ~ ssList(app(sk4,sk1))
    | ~ duplicatefreeP(app(sk4,sk1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_62]),c_0_4])])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( sk4 = nil
    | ssItem(X1)
    | ~ ssList(app(sk4,sk1)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,negated_conjecture,
    ( sk4 = nil
    | ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,c_0_13]),c_0_55]),c_0_14])])).

<A NAME="c_0_66"></A>cnf(c_0_66,negated_conjecture,
    ( tl(sk1) = nil
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_62]),c_0_4])]),<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( sk1 = nil
    | sk4 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( sk4 = nil
    | X1 = nil
    | cons(X2,X1) != sk1
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_62]),c_0_4])]),c_0_65]),<A HREF="#c_0_65">c_0_65</A>])).

<A NAME="c_0_69"></A>cnf(c_0_69,negated_conjecture,
    ( sk4 = nil
    | sk1 != nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_62]),c_0_4])]),<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1)
    | ~ ssList(cons(X1,nil)) ),
    inference(er,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( cons(hd(sk1),nil) = sk1
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_66]),c_0_14])]),<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_72"></A>cnf(c_0_72,negated_conjecture,
    ( hd(sk1) = sk5
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_62]),c_0_4])]),<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,negated_conjecture,
    ( skaf82(sk1) = nil
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_22]),c_0_23])])]),c_0_14])]),<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_74"></A>cnf(c_0_74,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>,c_0_24]),c_0_4])])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( cons(sk5,nil) = sk1
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>,c_0_72]),<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,negated_conjecture,
    ( sk1 = nil
    | ssItem(sk5) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_72]),c_0_14])]),<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_26]),<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_78"></A>cnf(c_0_78,plain,
    ( skaf44(cons(X1,X2)) = X1
    | ~ singletonP(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_28]),c_0_29]),c_0_4])])]),<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( cons(skaf83(sk1),nil) = sk1
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_73]),c_0_14])]),<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( sk1 = nil
    | singletonP(sk1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,c_0_75]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_81"></A>cnf(c_0_81,negated_conjecture,
    ( sk4 = nil
    | frontsegP(sk4,skaf42(sk4,sk5))
    | ~ ssList(cons(sk5,skaf43(sk5,sk4))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_59]),c_0_11])])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( skaf44(sk1) = skaf83(sk1)
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_79]),c_0_30]),c_0_4])]),<A HREF="#c_0_80">c_0_80</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( skaf44(sk1) = sk5
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_75]),c_0_4])]),c_0_76]),<A HREF="#c_0_80">c_0_80</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( sk4 = nil
    | frontsegP(sk4,skaf42(sk4,sk5)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_24]),c_0_10])]),<A HREF="#c_0_65">c_0_65</A>])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( skaf83(sk1) = sk5
    | sk1 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,<A HREF="#c_0_83">c_0_83</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_32]),c_0_4])])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( app(skaf42(sk4,sk5),skaf45(sk4,skaf42(sk4,sk5))) = sk4
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_84]),c_0_11]),c_0_55])])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( cons(sk5,skaf82(sk1)) = sk1
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_85]),c_0_14])])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( sk4 = nil
    | frontsegP(sk4,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>,c_0_87]),c_0_34]),c_0_11])])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( hd(sk1) = sk5
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_75]),c_0_4])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_91"></A>cnf(c_0_91,negated_conjecture,
    ( skaf82(sk1) = tl(sk1)
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_88]),c_0_23])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_92"></A>cnf(c_0_92,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_93"></A>cnf(c_0_93,negated_conjecture,
    ( app(nil,skaf45(sk4,nil)) = sk4
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_89]),c_0_4]),c_0_55])])).

<A NAME="c_0_94"></A>cnf(c_0_94,negated_conjecture,
    ( cons(sk5,tl(sk1)) = sk1
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_90]),c_0_14])])).

<A NAME="c_0_95"></A>cnf(c_0_95,negated_conjecture,
    ( sk1 = nil
    | ssList(tl(sk1)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,<A HREF="#c_0_91">c_0_91</A>])).

<A NAME="c_0_96"></A>cnf(c_0_96,negated_conjecture,
    ( sk1 = nil
    | memberP(sk1,sk5) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_75]),c_0_4])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,plain,
    ( app(nil,skaf45(X1,nil)) = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_32]),c_0_4])])).

<A NAME="c_0_98"></A>cnf(c_0_98,negated_conjecture,
    ( skaf45(sk4,nil) = sk4
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_93]),c_0_34])])).

<A NAME="c_0_99"></A>cnf(c_0_99,negated_conjecture,
    ( sk1 = nil
    | ~ ssList(app(app(X1,cons(sk5,X2)),sk1))
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(app(X1,cons(sk5,X2)),sk1)) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_94]),c_0_95]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,negated_conjecture,
    ( app(skaf42(sk1,sk5),cons(sk5,skaf43(sk5,sk1))) = sk1
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_96]),c_0_14])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_101"></A>cnf(c_0_101,plain,
    ( memberP(X1,skaf44(X1))
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_28]),c_0_29]),c_0_4])])).

<A NAME="c_0_102"></A>cnf(c_0_102,negated_conjecture,
    ( cons(sk5,app(nil,X1)) = app(sk1,X1)
    | sk1 = nil
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_75]),c_0_4])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_103"></A>cnf(c_0_103,negated_conjecture,
    ( app(nil,sk4) = sk4
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_98]),c_0_55])])).

<A NAME="c_0_104"></A>cnf(c_0_104,negated_conjecture,
    ( sk1 = nil
    | ~ ssList(app(sk1,sk1))
    | ~ duplicatefreeP(app(sk1,sk1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_100]),c_0_10]),c_0_11])])).

<A NAME="c_0_105"></A>cnf(c_0_105,plain,
    ( ~ singletonP(X1)
    | ~ ssList(app(app(X2,cons(skaf44(X1),X3)),X1))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(app(X2,cons(skaf44(X1),X3)),X1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_28]),c_0_29]),c_0_4])])).

<A NAME="c_0_106"></A>cnf(c_0_106,plain,
    ( app(skaf42(X1,skaf44(X1)),cons(skaf44(X1),skaf43(skaf44(X1),X1))) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_101]),c_0_29])])).

<A NAME="c_0_107"></A>cnf(c_0_107,negated_conjecture,
    ( app(sk1,sk4) = cons(sk5,sk4)
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>,c_0_103]),c_0_55])]),<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_108"></A>cnf(c_0_108,negated_conjecture,
    ( sk1 = nil
    | ssItem(X1)
    | ~ ssList(app(sk1,sk1)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_109"></A>cnf(c_0_109,plain,
    ( ~ singletonP(X1)
    | ~ ssList(app(X1,X1))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(X1,X1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_105">c_0_105</A>,c_0_106]),c_0_10]),c_0_11])])).

<A NAME="c_0_110"></A>cnf(c_0_110,plain,
    ( nil = X1
    | frontsegP(X2,tl(X1))
    | ~ frontsegP(cons(X3,X2),X1)
    | ~ ssItem(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_20]),c_0_39]),<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_111"></A>cnf(c_0_111,negated_conjecture,
    ( sk1 = nil
    | frontsegP(cons(sk5,sk4),sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_107]),c_0_55]),c_0_14])])).

<A NAME="c_0_112"></A>cnf(c_0_112,negated_conjecture,
    ( sk1 = nil
    | ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>,c_0_13]),c_0_14])])).

<A NAME="c_0_113"></A>cnf(c_0_113,plain,
    ( ssItem(X1)
    | ~ singletonP(X2)
    | ~ ssList(app(X2,X2))
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_114"></A>cnf(c_0_114,plain,
    ( ssList(cons(X1,app(X2,X3)))
    | ~ ssItem(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_37]),<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_40])]),c_0_4])])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( sk1 = nil
    | frontsegP(sk4,tl(sk1)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_110">c_0_110</A>,c_0_111]),c_0_55]),c_0_14])]),<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_117"></A>cnf(c_0_117,negated_conjecture,
    ( tl(sk1) = nil
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_75]),c_0_4])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_118"></A>cnf(c_0_118,plain,
    ( frontsegP(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_41]),c_0_4])])).

<A NAME="c_0_119"></A>cnf(c_0_119,plain,
    ( ssItem(X1)
    | ~ singletonP(X2)
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_113">c_0_113</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_120"></A>cnf(c_0_120,plain,
    ( ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_114">c_0_114</A>,c_0_115]),c_0_4])])).

<A NAME="c_0_121"></A>cnf(c_0_121,negated_conjecture,
    ( sk1 = nil
    | frontsegP(sk4,nil) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_116">c_0_116</A>,<A HREF="#c_0_117">c_0_117</A>])).

<A NAME="c_0_122"></A>cnf(c_0_122,plain,
    ( app(nil,skaf45(nil,nil)) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_118]),c_0_4])])).

<A NAME="c_0_123"></A>cnf(c_0_123,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_124"></A>cnf(c_0_124,plain,
    ( ssList(app(X1,app(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_43]),<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_125"></A>cnf(c_0_125,plain,
    ( ssItem(X1)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>,c_0_74]),<A HREF="#c_0_120">c_0_120</A>])).

<A NAME="c_0_126"></A>cnf(c_0_126,plain,
    ( frontsegP(app(X1,app(X2,X3)),app(X1,X2))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_43]),<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_127"></A>cnf(c_0_127,plain,
    ( app(X1,skaf45(X1,X1)) = X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_128"></A>cnf(c_0_128,plain,
    ( app(nil,app(X1,X2)) = app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_36]),c_0_4])])).

<A NAME="c_0_129"></A>cnf(c_0_129,negated_conjecture,
    ( app(nil,skaf45(sk4,nil)) = sk4
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_121]),c_0_4]),c_0_55])])).

<A NAME="c_0_130"></A>cnf(c_0_130,plain,
    ( skaf45(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_122]),c_0_34]),c_0_4])])).

<A NAME="c_0_131"></A>cnf(c_0_131,plain,
    ( segmentP(app(X1,app(X2,X3)),X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>,c_0_43]),<A HREF="#c_0_124">c_0_124</A>])).

<A NAME="c_0_132"></A>cnf(c_0_132,negated_conjecture,
    ( app(sk1,X1) = cons(sk5,X1)
    | sk4 = nil
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_62]),<A HREF="#c_0_65">c_0_65</A>])).

<A NAME="c_0_133"></A>cnf(c_0_133,plain,
    ( skaf83(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_22]),c_0_30]),c_0_23])])]),c_0_24]),<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_134"></A>cnf(c_0_134,plain,
    ( skaf83(X1) = X2
    | nil = X1
    | ~ frontsegP(X1,cons(X2,X3))
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_22]),c_0_30]),c_0_23])])).

<A NAME="c_0_135"></A>cnf(c_0_135,plain,
    ( ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_125">c_0_125</A>,<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_136"></A>cnf(c_0_136,plain,
    ( frontsegP(app(X1,X2),app(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_126">c_0_126</A>,c_0_127]),c_0_34])])).

<A NAME="c_0_137"></A>cnf(c_0_137,negated_conjecture,
    ( app(nil,cons(sk5,sk4)) = cons(sk5,sk4)
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_128">c_0_128</A>,c_0_107]),c_0_14]),c_0_55])])).

<A NAME="c_0_138"></A>cnf(c_0_138,negated_conjecture,
    ( sk1 = nil
    | ssList(cons(X1,sk4)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_114">c_0_114</A>,c_0_129]),c_0_34]),c_0_4])]),<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_139"></A>cnf(c_0_139,plain,
    ( cons(X1,app(X2,X3)) != nil
    | ~ ssItem(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_37]),c_0_24]),<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_140"></A>cnf(c_0_140,negated_conjecture,
    ( app(sk1,nil) = cons(sk5,nil)
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>,c_0_127]),c_0_130]),c_0_34]),c_0_4])])).

<A NAME="c_0_141"></A>cnf(c_0_141,negated_conjecture,
    ( cons(sk5,app(nil,X1)) = app(sk1,X1)
    | sk4 = nil
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_62]),c_0_4])]),<A HREF="#c_0_65">c_0_65</A>])).

<A NAME="c_0_142"></A>cnf(c_0_142,negated_conjecture,
    ( sk4 = nil
    | segmentP(app(X1,cons(sk5,X2)),sk1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>,c_0_132]),c_0_14])])).

<A NAME="c_0_143"></A>cnf(c_0_143,plain,
    ( skaf44(X1) = skaf83(X1)
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_133">c_0_133</A>,c_0_28]),c_0_29]),c_0_4])])).

<A NAME="c_0_144"></A>cnf(c_0_144,plain,
    ( skaf83(X1) = X2
    | nil = X1
    | ~ frontsegP(X1,cons(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_134">c_0_134</A>,c_0_135])])).

<A NAME="c_0_145"></A>cnf(c_0_145,negated_conjecture,
    ( sk1 = nil
    | frontsegP(cons(sk5,sk4),cons(sk5,sk4)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_136">c_0_136</A>,c_0_137]),c_0_4])]),<A HREF="#c_0_138">c_0_138</A>])).

<A NAME="c_0_146"></A>cnf(c_0_146,negated_conjecture,
    ( sk1 = nil
    | cons(X1,sk4) != nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_139">c_0_139</A>,c_0_129]),c_0_34]),c_0_4])]),<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_147"></A>cnf(c_0_147,negated_conjecture,
    ( sk1 = nil
    | X1 = nil
    | app(sk1,X1) != cons(sk5,nil)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_140]),c_0_4]),c_0_14])])).

<A NAME="c_0_148"></A>cnf(c_0_148,negated_conjecture,
    ( app(sk1,sk4) = cons(sk5,sk4)
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_141">c_0_141</A>,c_0_103]),c_0_55])])).

<A NAME="c_0_149"></A>cnf(c_0_149,negated_conjecture,
    ( ~ segmentP(sk4,sk1)
    | ~ singletonP(sk1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_150"></A>cnf(c_0_150,negated_conjecture,
    ( sk4 = nil
    | segmentP(sk4,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_142">c_0_142</A>,c_0_59]),c_0_10]),c_0_11])])).

<A NAME="c_0_151"></A>cnf(c_0_151,negated_conjecture,
    ( sk4 = nil
    | singletonP(sk1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,c_0_62]),<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_152"></A>cnf(c_0_152,plain,
    ( ~ neq(X1,X1)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>])).

<A NAME="c_0_153"></A>cnf(c_0_153,plain,
    ( cons(skaf83(X1),nil) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,<A HREF="#c_0_143">c_0_143</A>])).

<A NAME="c_0_154"></A>cnf(c_0_154,negated_conjecture,
    ( skaf83(cons(sk5,sk4)) = sk5
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_144">c_0_144</A>,c_0_145]),c_0_55])]),c_0_138]),<A HREF="#c_0_146">c_0_146</A>])).

<A NAME="c_0_155"></A>cnf(c_0_155,negated_conjecture,
    ( sk1 = nil
    | cons(sk5,sk4) != cons(sk5,nil) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_147">c_0_147</A>,c_0_148]),c_0_55])]),<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_156"></A>cnf(c_0_156,negated_conjecture,
    ( neq(sk4,nil)
    | ~ segmentP(sk4,sk1) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_1]),<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_157"></A>cnf(c_0_157,negated_conjecture,
    ( sk4 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_149">c_0_149</A>,c_0_150]),<A HREF="#c_0_151">c_0_151</A>])).

<A NAME="c_0_158"></A>cnf(c_0_158,plain,
    ( ~ neq(X1,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_152">c_0_152</A>,c_0_135])])).

<A NAME="c_0_159"></A>cnf(c_0_159,negated_conjecture,
    ( sk1 = nil
    | ~ singletonP(cons(sk5,sk4)) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_153">c_0_153</A>,c_0_154]),c_0_138]),<A HREF="#c_0_155">c_0_155</A>])).

<A NAME="c_0_160"></A>cnf(c_0_160,plain,
    ( singletonP(cons(X1,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,c_0_135])])).

<A NAME="c_0_161"></A>cnf(c_0_161,negated_conjecture,
    ( ~ segmentP(nil,sk1) ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_156">c_0_156</A>,c_0_157]),c_0_157]),<A HREF="#c_0_158">c_0_158</A>])).

<A NAME="c_0_162"></A>cnf(c_0_162,negated_conjecture,
    ( sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_159">c_0_159</A>,c_0_157]),c_0_160])])).

<A NAME="c_0_163"></A>cnf(c_0_163,plain,
    ( segmentP(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_54]),c_0_4])])).

<A NAME="c_0_164"></A>cnf(c_0_164,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_161">c_0_161</A>,c_0_162]),c_0_163])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC390-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n098.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:43:23 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 12.62/12.81  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 12.62/12.81  # and selection function SelectCQArNTEqFirst.
% 12.62/12.81  #
% 12.62/12.81  # Preprocessing time       : 0.010 s
% 12.62/12.81  # Presaturation interreduction done
% 12.62/12.81  
% 12.62/12.81  # Proof found!
% 12.62/12.81  # SZS status Unsatisfiable
% 12.62/12.81  # SZS output start CNFRefutation
% 12.62/12.81  cnf(c_0_0, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
% 12.62/12.81  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 12.62/12.81  cnf(c_0_2, negated_conjecture, (memberP(sk4,sk5)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
% 12.62/12.81  cnf(c_0_3, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause100)).
% 12.62/12.81  cnf(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 12.62/12.81  cnf(c_0_5, negated_conjecture, (ssItem(sk5)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
% 12.62/12.81  cnf(c_0_6, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause179)).
% 12.62/12.81  cnf(c_0_7, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause169)).
% 12.62/12.81  cnf(c_0_8, negated_conjecture, (cons(sk5,nil)=sk3|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
% 12.62/12.81  cnf(c_0_9, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 12.62/12.81  cnf(c_0_10, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause52)).
% 12.62/12.81  cnf(c_0_11, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause53)).
% 12.62/12.81  cnf(c_0_12, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause72)).
% 12.62/12.81  cnf(c_0_13, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
% 12.62/12.81  cnf(c_0_14, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 12.62/12.81  cnf(c_0_15, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause96)).
% 12.62/12.81  cnf(c_0_16, negated_conjecture, (nil=sk3|nil!=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 12.62/12.81  cnf(c_0_17, axiom, (X4=X2|cons(X1,X2)!=cons(X3,X4)|~ssItem(X3)|~ssItem(X1)|~ssList(X4)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause171)).
% 12.62/12.81  cnf(c_0_18, axiom, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause99)).
% 12.62/12.81  cnf(c_0_19, axiom, (singletonP(X2)|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause116)).
% 12.62/12.81  cnf(c_0_20, axiom, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause104)).
% 12.62/12.81  cnf(c_0_21, axiom, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause97)).
% 12.62/12.81  cnf(c_0_22, axiom, (cons(skaf83(X1),skaf82(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause109)).
% 12.62/12.81  cnf(c_0_23, axiom, (ssList(skaf82(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause13)).
% 12.62/12.81  cnf(c_0_24, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause86)).
% 12.62/12.81  cnf(c_0_25, axiom, (ssItem(hd(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause76)).
% 12.62/12.81  cnf(c_0_26, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause144)).
% 12.62/12.81  cnf(c_0_27, axiom, (X1=X3|cons(X1,X2)!=cons(X3,X4)|~ssItem(X3)|~ssItem(X1)|~ssList(X4)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause170)).
% 12.62/12.81  cnf(c_0_28, axiom, (cons(skaf44(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause101)).
% 12.62/12.81  cnf(c_0_29, axiom, (ssItem(skaf44(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause47)).
% 12.62/12.81  cnf(c_0_30, axiom, (ssItem(skaf83(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause12)).
% 12.62/12.81  cnf(c_0_31, axiom, (frontsegP(app(X1,X3),X2)|~frontsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause137)).
% 12.62/12.81  cnf(c_0_32, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause60)).
% 12.62/12.81  cnf(c_0_33, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause132)).
% 12.62/12.81  cnf(c_0_34, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause51)).
% 12.62/12.81  cnf(c_0_35, axiom, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause138)).
% 12.62/12.81  cnf(c_0_36, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause74)).
% 12.62/12.81  cnf(c_0_37, axiom, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause157)).
% 12.62/12.81  cnf(c_0_38, axiom, (frontsegP(X2,X4)|~frontsegP(cons(X1,X2),cons(X3,X4))|~ssList(X4)|~ssList(X2)|~ssItem(X3)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause174)).
% 12.62/12.81  cnf(c_0_39, axiom, (ssList(tl(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause75)).
% 12.62/12.81  cnf(c_0_40, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause145)).
% 12.62/12.81  cnf(c_0_41, axiom, (frontsegP(nil,X1)|nil!=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause83)).
% 12.62/12.82  cnf(c_0_42, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause173)).
% 12.62/12.82  cnf(c_0_43, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause149)).
% 12.62/12.82  cnf(c_0_44, axiom, (frontsegP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause61)).
% 12.62/12.82  cnf(c_0_45, axiom, (nil=X2|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause119)).
% 12.62/12.82  cnf(c_0_46, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause120)).
% 12.62/12.82  cnf(c_0_47, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause98)).
% 12.62/12.82  cnf(c_0_48, axiom, (X1=X3|~frontsegP(cons(X1,X2),cons(X3,X4))|~ssList(X4)|~ssList(X2)|~ssItem(X3)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause176)).
% 12.62/12.82  cnf(c_0_49, axiom, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause118)).
% 12.62/12.82  cnf(c_0_50, axiom, (X2=X3|app(X1,X2)!=app(X1,X3)|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause150)).
% 12.62/12.82  cnf(c_0_51, negated_conjecture, (~segmentP(sk2,sk1)|~singletonP(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_12)).
% 12.62/12.82  cnf(c_0_52, axiom, (X1!=X2|~neq(X1,X2)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause117)).
% 12.62/12.82  cnf(c_0_53, negated_conjecture, (neq(sk2,nil)|~segmentP(sk2,sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_11)).
% 12.62/12.82  cnf(c_0_54, axiom, (segmentP(nil,X1)|nil!=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause79)).
% 12.62/12.82  cnf(c_0_55, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 12.62/12.82  cnf(c_0_56, negated_conjecture, (sk4=nil|memberP(sk4,sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_3]), c_0_55]), c_0_4])])).
% 12.62/12.82  cnf(c_0_57, negated_conjecture, (sk4=nil|ssItem(sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_3]), c_0_55]), c_0_4])])).
% 12.62/12.82  cnf(c_0_58, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_6])).
% 12.62/12.82  cnf(c_0_59, negated_conjecture, (app(skaf42(sk4,sk5),cons(sk5,skaf43(sk5,sk4)))=sk4|sk4=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_56]), c_0_55])]), c_0_57])).
% 12.62/12.82  cnf(c_0_60, negated_conjecture, (cons(sk5,nil)=sk1|~neq(sk4,nil)), inference(rw,[status(thm)],[c_0_8, c_0_9])).
% 12.62/12.82  cnf(c_0_61, negated_conjecture, (sk4=nil|~ssList(app(sk4,cons(sk5,X1)))|~ssList(X1)|~duplicatefreeP(app(sk4,cons(sk5,X1)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_10]), c_0_11])]), c_0_57])).
% 12.62/12.82  cnf(c_0_62, negated_conjecture, (cons(sk5,nil)=sk1|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_3]), c_0_55]), c_0_4])])).
% 12.62/12.82  cnf(c_0_63, negated_conjecture, (sk4=nil|~ssList(app(sk4,sk1))|~duplicatefreeP(app(sk4,sk1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_4])])).
% 12.62/12.82  cnf(c_0_64, negated_conjecture, (sk4=nil|ssItem(X1)|~ssList(app(sk4,sk1))), inference(spm,[status(thm)],[c_0_63, c_0_12])).
% 12.62/12.82  cnf(c_0_65, negated_conjecture, (sk4=nil|ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_13]), c_0_55]), c_0_14])])).
% 12.62/12.82  cnf(c_0_66, negated_conjecture, (tl(sk1)=nil|sk4=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_62]), c_0_4])]), c_0_57])).
% 12.62/12.82  cnf(c_0_67, negated_conjecture, (sk1=nil|sk4!=nil), inference(rw,[status(thm)],[c_0_16, c_0_9])).
% 12.62/12.82  cnf(c_0_68, negated_conjecture, (sk4=nil|X1=nil|cons(X2,X1)!=sk1|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_62]), c_0_4])]), c_0_65]), c_0_65])).
% 12.62/12.82  cnf(c_0_69, negated_conjecture, (sk4=nil|sk1!=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_62]), c_0_4])]), c_0_57])).
% 12.62/12.82  cnf(c_0_70, plain, (singletonP(cons(X1,nil))|~ssItem(X1)|~ssList(cons(X1,nil))), inference(er,[status(thm)],[c_0_19])).
% 12.62/12.82  cnf(c_0_71, negated_conjecture, (cons(hd(sk1),nil)=sk1|sk1=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_66]), c_0_14])]), c_0_67])).
% 12.62/12.82  cnf(c_0_72, negated_conjecture, (hd(sk1)=sk5|sk4=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_62]), c_0_4])]), c_0_57])).
% 12.62/12.82  cnf(c_0_73, negated_conjecture, (skaf82(sk1)=nil|sk4=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_22]), c_0_23])])]), c_0_14])]), c_0_69])).
% 12.62/12.82  cnf(c_0_74, plain, (singletonP(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_24]), c_0_4])])).
% 12.62/12.82  cnf(c_0_75, negated_conjecture, (cons(sk5,nil)=sk1|sk1=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_67])).
% 12.62/12.82  cnf(c_0_76, negated_conjecture, (sk1=nil|ssItem(sk5)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_72]), c_0_14])]), c_0_67])).
% 12.62/12.82  cnf(c_0_77, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_26]), c_0_13])).
% 12.62/12.82  cnf(c_0_78, plain, (skaf44(cons(X1,X2))=X1|~singletonP(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_4])])]), c_0_24])).
% 12.62/12.82  cnf(c_0_79, negated_conjecture, (cons(skaf83(sk1),nil)=sk1|sk1=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_73]), c_0_14])]), c_0_67])).
% 12.62/12.82  cnf(c_0_80, negated_conjecture, (sk1=nil|singletonP(sk1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])).
% 12.62/12.82  cnf(c_0_81, negated_conjecture, (sk4=nil|frontsegP(sk4,skaf42(sk4,sk5))|~ssList(cons(sk5,skaf43(sk5,sk4)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_59]), c_0_11])])).
% 12.62/12.82  cnf(c_0_82, negated_conjecture, (skaf44(sk1)=skaf83(sk1)|sk1=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_30]), c_0_4])]), c_0_80])).
% 12.62/12.82  cnf(c_0_83, negated_conjecture, (skaf44(sk1)=sk5|sk1=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_75]), c_0_4])]), c_0_76]), c_0_80])).
% 12.62/12.82  cnf(c_0_84, negated_conjecture, (sk4=nil|frontsegP(sk4,skaf42(sk4,sk5))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_24]), c_0_10])]), c_0_65])).
% 12.62/12.82  cnf(c_0_85, negated_conjecture, (skaf83(sk1)=sk5|sk1=nil), inference(spm,[status(thm)],[c_0_82, c_0_83])).
% 12.62/12.82  cnf(c_0_86, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_4])])).
% 12.62/12.82  cnf(c_0_87, negated_conjecture, (app(skaf42(sk4,sk5),skaf45(sk4,skaf42(sk4,sk5)))=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_84]), c_0_11]), c_0_55])])).
% 12.62/12.82  cnf(c_0_88, negated_conjecture, (cons(sk5,skaf82(sk1))=sk1|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_85]), c_0_14])])).
% 12.62/12.82  cnf(c_0_89, negated_conjecture, (sk4=nil|frontsegP(sk4,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_34]), c_0_11])])).
% 12.62/12.82  cnf(c_0_90, negated_conjecture, (hd(sk1)=sk5|sk1=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_75]), c_0_4])]), c_0_76])).
% 12.62/12.82  cnf(c_0_91, negated_conjecture, (skaf82(sk1)=tl(sk1)|sk1=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_88]), c_0_23])]), c_0_76])).
% 12.62/12.82  cnf(c_0_92, plain, (memberP(cons(X1,X2),X1)|~ssItem(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_35])).
% 12.62/12.82  cnf(c_0_93, negated_conjecture, (app(nil,skaf45(sk4,nil))=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_89]), c_0_4]), c_0_55])])).
% 12.62/12.82  cnf(c_0_94, negated_conjecture, (cons(sk5,tl(sk1))=sk1|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_90]), c_0_14])])).
% 12.62/12.82  cnf(c_0_95, negated_conjecture, (sk1=nil|ssList(tl(sk1))), inference(spm,[status(thm)],[c_0_23, c_0_91])).
% 12.62/12.82  cnf(c_0_96, negated_conjecture, (sk1=nil|memberP(sk1,sk5)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_75]), c_0_4])]), c_0_76])).
% 12.62/12.82  cnf(c_0_97, plain, (app(nil,skaf45(X1,nil))=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_32]), c_0_4])])).
% 12.62/12.82  cnf(c_0_98, negated_conjecture, (skaf45(sk4,nil)=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_93]), c_0_34])])).
% 12.62/12.82  cnf(c_0_99, negated_conjecture, (sk1=nil|~ssList(app(app(X1,cons(sk5,X2)),sk1))|~ssList(X2)|~ssList(X1)|~duplicatefreeP(app(app(X1,cons(sk5,X2)),sk1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_94]), c_0_95]), c_0_76])).
% 12.62/12.82  cnf(c_0_100, negated_conjecture, (app(skaf42(sk1,sk5),cons(sk5,skaf43(sk5,sk1)))=sk1|sk1=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_96]), c_0_14])]), c_0_76])).
% 12.62/12.82  cnf(c_0_101, plain, (memberP(X1,skaf44(X1))|~singletonP(X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_28]), c_0_29]), c_0_4])])).
% 12.62/12.82  cnf(c_0_102, negated_conjecture, (cons(sk5,app(nil,X1))=app(sk1,X1)|sk1=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_75]), c_0_4])]), c_0_76])).
% 12.62/12.82  cnf(c_0_103, negated_conjecture, (app(nil,sk4)=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_55])])).
% 12.62/12.82  cnf(c_0_104, negated_conjecture, (sk1=nil|~ssList(app(sk1,sk1))|~duplicatefreeP(app(sk1,sk1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_10]), c_0_11])])).
% 12.62/12.82  cnf(c_0_105, plain, (~singletonP(X1)|~ssList(app(app(X2,cons(skaf44(X1),X3)),X1))|~ssList(X3)|~ssList(X2)|~ssList(X1)|~duplicatefreeP(app(app(X2,cons(skaf44(X1),X3)),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_28]), c_0_29]), c_0_4])])).
% 12.62/12.82  cnf(c_0_106, plain, (app(skaf42(X1,skaf44(X1)),cons(skaf44(X1),skaf43(skaf44(X1),X1)))=X1|~singletonP(X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_101]), c_0_29])])).
% 12.62/12.82  cnf(c_0_107, negated_conjecture, (app(sk1,sk4)=cons(sk5,sk4)|sk1=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_55])]), c_0_67])).
% 12.62/12.82  cnf(c_0_108, negated_conjecture, (sk1=nil|ssItem(X1)|~ssList(app(sk1,sk1))), inference(spm,[status(thm)],[c_0_104, c_0_12])).
% 12.62/12.82  cnf(c_0_109, plain, (~singletonP(X1)|~ssList(app(X1,X1))|~ssList(X1)|~duplicatefreeP(app(X1,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_10]), c_0_11])])).
% 12.62/12.82  cnf(c_0_110, plain, (nil=X1|frontsegP(X2,tl(X1))|~frontsegP(cons(X3,X2),X1)|~ssItem(X3)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_20]), c_0_39]), c_0_25])).
% 12.62/12.82  cnf(c_0_111, negated_conjecture, (sk1=nil|frontsegP(cons(sk5,sk4),sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_107]), c_0_55]), c_0_14])])).
% 12.62/12.82  cnf(c_0_112, negated_conjecture, (sk1=nil|ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_13]), c_0_14])])).
% 12.62/12.82  cnf(c_0_113, plain, (ssItem(X1)|~singletonP(X2)|~ssList(app(X2,X2))|~ssList(X2)), inference(spm,[status(thm)],[c_0_109, c_0_12])).
% 12.62/12.82  cnf(c_0_114, plain, (ssList(cons(X1,app(X2,X3)))|~ssItem(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_37]), c_0_24])).
% 12.62/12.82  cnf(c_0_115, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_40])]), c_0_4])])).
% 12.62/12.82  cnf(c_0_116, negated_conjecture, (sk1=nil|frontsegP(sk4,tl(sk1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_55]), c_0_14])]), c_0_112])).
% 12.62/12.82  cnf(c_0_117, negated_conjecture, (tl(sk1)=nil|sk1=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_75]), c_0_4])]), c_0_76])).
% 12.62/12.82  cnf(c_0_118, plain, (frontsegP(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_41]), c_0_4])])).
% 12.62/12.82  cnf(c_0_119, plain, (ssItem(X1)|~singletonP(X2)|~ssList(X2)), inference(spm,[status(thm)],[c_0_113, c_0_13])).
% 12.62/12.82  cnf(c_0_120, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_4])])).
% 12.62/12.82  cnf(c_0_121, negated_conjecture, (sk1=nil|frontsegP(sk4,nil)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
% 12.62/12.82  cnf(c_0_122, plain, (app(nil,skaf45(nil,nil))=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_118]), c_0_4])])).
% 12.62/12.82  cnf(c_0_123, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(er,[status(thm)],[c_0_42])).
% 12.62/12.82  cnf(c_0_124, plain, (ssList(app(X1,app(X2,X3)))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_43]), c_0_13])).
% 12.62/12.82  cnf(c_0_125, plain, (ssItem(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_74]), c_0_120])).
% 12.62/12.82  cnf(c_0_126, plain, (frontsegP(app(X1,app(X2,X3)),app(X1,X2))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_43]), c_0_13])).
% 12.62/12.82  cnf(c_0_127, plain, (app(X1,skaf45(X1,X1))=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_33, c_0_44])).
% 12.62/12.82  cnf(c_0_128, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_36]), c_0_4])])).
% 12.62/12.82  cnf(c_0_129, negated_conjecture, (app(nil,skaf45(sk4,nil))=sk4|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_121]), c_0_4]), c_0_55])])).
% 12.62/12.82  cnf(c_0_130, plain, (skaf45(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_122]), c_0_34]), c_0_4])])).
% 12.62/12.82  cnf(c_0_131, plain, (segmentP(app(X1,app(X2,X3)),X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_43]), c_0_124])).
% 12.62/12.82  cnf(c_0_132, negated_conjecture, (app(sk1,X1)=cons(sk5,X1)|sk4=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_62]), c_0_65])).
% 12.62/12.82  cnf(c_0_133, plain, (skaf83(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_22]), c_0_30]), c_0_23])])]), c_0_24]), c_0_47])).
% 12.62/12.82  cnf(c_0_134, plain, (skaf83(X1)=X2|nil=X1|~frontsegP(X1,cons(X2,X3))|~ssItem(X2)|~ssList(X3)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_22]), c_0_30]), c_0_23])])).
% 12.62/12.82  cnf(c_0_135, plain, (ssItem(X1)), inference(spm,[status(thm)],[c_0_125, c_0_29])).
% 12.62/12.82  cnf(c_0_136, plain, (frontsegP(app(X1,X2),app(X1,X2))|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_34])])).
% 12.62/12.82  cnf(c_0_137, negated_conjecture, (app(nil,cons(sk5,sk4))=cons(sk5,sk4)|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_107]), c_0_14]), c_0_55])])).
% 12.62/12.82  cnf(c_0_138, negated_conjecture, (sk1=nil|ssList(cons(X1,sk4))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_129]), c_0_34]), c_0_4])]), c_0_112])).
% 12.62/12.82  cnf(c_0_139, plain, (cons(X1,app(X2,X3))!=nil|~ssItem(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_37]), c_0_24]), c_0_47])).
% 12.62/12.82  cnf(c_0_140, negated_conjecture, (app(sk1,nil)=cons(sk5,nil)|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_127]), c_0_130]), c_0_34]), c_0_4])])).
% 12.62/12.82  cnf(c_0_141, negated_conjecture, (cons(sk5,app(nil,X1))=app(sk1,X1)|sk4=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_62]), c_0_4])]), c_0_65])).
% 12.62/12.82  cnf(c_0_142, negated_conjecture, (sk4=nil|segmentP(app(X1,cons(sk5,X2)),sk1)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_14])])).
% 12.62/12.82  cnf(c_0_143, plain, (skaf44(X1)=skaf83(X1)|~singletonP(X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_28]), c_0_29]), c_0_4])])).
% 12.62/12.82  cnf(c_0_144, plain, (skaf83(X1)=X2|nil=X1|~frontsegP(X1,cons(X2,X3))|~ssList(X3)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_134, c_0_135])])).
% 12.62/12.82  cnf(c_0_145, negated_conjecture, (sk1=nil|frontsegP(cons(sk5,sk4),cons(sk5,sk4))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_4])]), c_0_138])).
% 12.62/12.82  cnf(c_0_146, negated_conjecture, (sk1=nil|cons(X1,sk4)!=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_129]), c_0_34]), c_0_4])]), c_0_112])).
% 12.62/12.82  cnf(c_0_147, negated_conjecture, (sk1=nil|X1=nil|app(sk1,X1)!=cons(sk5,nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_140]), c_0_4]), c_0_14])])).
% 12.62/12.82  cnf(c_0_148, negated_conjecture, (app(sk1,sk4)=cons(sk5,sk4)|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_103]), c_0_55])])).
% 12.62/12.82  cnf(c_0_149, negated_conjecture, (~segmentP(sk4,sk1)|~singletonP(sk1)), inference(rw,[status(thm)],[c_0_51, c_0_1])).
% 12.62/12.82  cnf(c_0_150, negated_conjecture, (sk4=nil|segmentP(sk4,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_59]), c_0_10]), c_0_11])])).
% 12.62/12.82  cnf(c_0_151, negated_conjecture, (sk4=nil|singletonP(sk1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_62]), c_0_57])).
% 12.62/12.82  cnf(c_0_152, plain, (~neq(X1,X1)|~ssItem(X1)), inference(er,[status(thm)],[c_0_52])).
% 12.62/12.82  cnf(c_0_153, plain, (cons(skaf83(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), inference(spm,[status(thm)],[c_0_28, c_0_143])).
% 12.62/12.82  cnf(c_0_154, negated_conjecture, (skaf83(cons(sk5,sk4))=sk5|sk1=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_55])]), c_0_138]), c_0_146])).
% 12.62/12.82  cnf(c_0_155, negated_conjecture, (sk1=nil|cons(sk5,sk4)!=cons(sk5,nil)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_148]), c_0_55])]), c_0_67])).
% 12.62/12.82  cnf(c_0_156, negated_conjecture, (neq(sk4,nil)|~segmentP(sk4,sk1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_1]), c_0_1])).
% 12.62/12.82  cnf(c_0_157, negated_conjecture, (sk4=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_151])).
% 12.62/12.82  cnf(c_0_158, plain, (~neq(X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_152, c_0_135])])).
% 12.62/12.82  cnf(c_0_159, negated_conjecture, (sk1=nil|~singletonP(cons(sk5,sk4))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154]), c_0_138]), c_0_155])).
% 12.62/12.82  cnf(c_0_160, plain, (singletonP(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_135])])).
% 12.62/12.82  cnf(c_0_161, negated_conjecture, (~segmentP(nil,sk1)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_156, c_0_157]), c_0_157]), c_0_158])).
% 12.62/12.82  cnf(c_0_162, negated_conjecture, (sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_159, c_0_157]), c_0_160])])).
% 12.62/12.82  cnf(c_0_163, plain, (segmentP(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_54]), c_0_4])])).
% 12.62/12.82  cnf(c_0_164, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_161, c_0_162]), c_0_163])]), ['proof']).
% 12.62/12.82  # SZS output end CNFRefutation
% 12.62/12.82  # Training examples: 0 positive, 0 negative
% 12.62/12.85  
% 12.62/12.85  # -------------------------------------------------
% 12.62/12.85  # User time                : 12.302 s
% 12.62/12.85  # System time              : 0.324 s
% 12.62/12.85  # Total time               : 12.626 s
% 12.62/12.85  # Maximum resident set size: 1856 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
