<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC083+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC083+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC083+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC083%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC083+1.p">SWC083+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n153.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:35 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :  115 (10809 expanded)
%              Number of clauses        :   83 (3968 expanded)
%              Number of leaves         :   16 (2627 expanded)
%              Depth                    :   25
%              Number of atoms          :  370 (97365 expanded)
%              Number of equality atoms :   97 (23052 expanded)
%              Maximal formula depth    :   22 (   4 average)
%              Maximal clause size      :   18 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ~ ssList(X4)
                  | X2 != X4
                  | X1 != X3
                  | ~ neq(X2,nil)
                  | ? [X5] :
                      ( ssList(X5)
                      & neq(X5,nil)
                      & segmentP(X2,X5)
                      & segmentP(X1,X5) )
                  | ! [X6] :
                      ( ssList(X6)
                     => ! [X7] :
                          ( ssList(X7)
                         => ! [X8] :
                              ( ~ ssList(X8)
                              | app(app(X6,X7),X8) != X3
                              | app(X6,X8) != X4 ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( rearsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax6)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X3,X2) = app(X1,X2)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax79)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( segmentP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax7)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( frontsegP(X1,X2)
               => frontsegP(app(X1,X3),X2) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax43)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,nil) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax45)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( frontsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X2,X3) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax5)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X2,X3) = app(X2,X1)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax80)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,X1) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax42)).

<A NAME="c_0_16"></A>fof(c_0_16,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ~ ssList(X4)
                    | X2 != X4
                    | X1 != X3
                    | ~ neq(X2,nil)
                    | ? [X5] :
                        ( ssList(X5)
                        & neq(X5,nil)
                        & segmentP(X2,X5)
                        & segmentP(X1,X5) )
                    | ! [X6] :
                        ( ssList(X6)
                       => ! [X7] :
                            ( ssList(X7)
                           => ! [X8] :
                                ( ~ ssList(X8)
                                | app(app(X6,X7),X8) != X3
                                | app(X6,X8) != X4 ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_17"></A>fof(c_0_17,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk7_2(X4,X5))
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(esk7_2(X4,X5),X5) = X4
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X7,X5) != X4
        | rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

<A NAME="c_0_18"></A>fof(c_0_18,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_19"></A>fof(c_0_19,negated_conjecture,(
    ! [X13] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & neq(esk49_0,nil)
      & ( ~ ssList(X13)
        | ~ neq(X13,nil)
        | ~ segmentP(esk49_0,X13)
        | ~ segmentP(esk48_0,X13) )
      & ssList(esk52_0)
      & ssList(esk53_0)
      & ssList(esk54_0)
      & app(app(esk52_0,esk53_0),esk54_0) = esk50_0
      & app(esk52_0,esk54_0) = esk51_0 ) ),
    inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])).

<A NAME="c_0_20"></A>cnf(c_0_20,plain,
    ( rearsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X3,X2) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( app(app(esk52_0,esk53_0),esk54_0) = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_24"></A>fof(c_0_24,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X6,X5) != app(X4,X5)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

<A NAME="c_0_25"></A>cnf(c_0_25,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_20]),<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( app(app(esk52_0,esk53_0),esk54_0) = esk48_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( ssList(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,plain,
    ( X1 = X2
    | app(X1,X3) != app(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( rearsegP(esk48_0,esk54_0)
    | ~ ssList(app(esk52_0,esk53_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26]),c_0_27])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( ssList(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( ssList(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( X1 = app(esk52_0,esk53_0)
    | app(X1,esk54_0) != esk48_0
    | ~ ssList(app(esk52_0,esk53_0))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_26]),c_0_27])])).

<A NAME="c_0_33"></A>cnf(c_0_33,plain,
    ( app(esk7_2(X1,X2),X2) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( rearsegP(esk48_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_21]),c_0_30]),c_0_31])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( ssList(esk7_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( X1 = app(esk52_0,esk53_0)
    | app(X1,esk54_0) != esk48_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_21]),c_0_30]),c_0_31])])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( app(esk7_2(esk48_0,esk54_0),esk54_0) = esk48_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_34]),c_0_27]),c_0_35])])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( ssList(esk7_2(esk48_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_34]),c_0_27]),c_0_35])])).

<A NAME="c_0_40"></A>fof(c_0_40,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( esk7_2(esk48_0,esk54_0) = app(esk52_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_38]),c_0_39])])).

<A NAME="c_0_42"></A>cnf(c_0_42,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( ssList(app(esk52_0,esk53_0)) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_44"></A>fof(c_0_44,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk8_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ssList(esk9_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(app(X9,X6),X10) != X5
        | segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( app(app(esk52_0,esk53_0),app(esk54_0,X1)) = app(esk48_0,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_26]),c_0_27])]),c_0_43])])).

<A NAME="c_0_46"></A>cnf(c_0_46,plain,
    ( segmentP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(app(X3,X2),X4) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( ssList(app(esk48_0,X1))
    | ~ ssList(app(esk54_0,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_45]),c_0_43])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( app(esk52_0,esk54_0) = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_50"></A>cnf(c_0_50,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( ssList(app(esk48_0,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_21]),c_0_27])])).

<A NAME="c_0_52"></A>fof(c_0_52,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( app(esk52_0,esk54_0) = esk49_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( segmentP(app(esk48_0,X1),esk54_0)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_26]),c_0_27])]),c_0_43])]),<A HREF="#c_0_51">c_0_51</A>])).

<A NAME="c_0_55"></A>cnf(c_0_55,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( segmentP(app(esk49_0,X1),esk54_0)
    | ~ ssList(app(esk49_0,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_53]),c_0_31]),c_0_27])])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( segmentP(esk48_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_55]),c_0_56]),c_0_35])])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( segmentP(esk49_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,c_0_55]),c_0_58]),c_0_56])])).

<A NAME="c_0_62"></A>fof(c_0_62,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( ~ neq(esk54_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,c_0_60]),c_0_61]),c_0_27])])).

<A NAME="c_0_64"></A>cnf(c_0_64,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,negated_conjecture,
    ( esk54_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>,c_0_64]),c_0_56]),c_0_27])])).

<A NAME="c_0_66"></A>fof(c_0_66,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ frontsegP(X4,X5)
      | frontsegP(app(X4,X6),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_67"></A>fof(c_0_67,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( app(esk52_0,nil) = esk49_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,<A HREF="#c_0_65">c_0_65</A>])).

<A NAME="c_0_69"></A>fof(c_0_69,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk6_2(X4,X5))
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(X5,esk6_2(X4,X5)) = X4
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X5,X7) != X4
        | frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).

<A NAME="c_0_70"></A>fof(c_0_70,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).

<A NAME="c_0_71"></A>cnf(c_0_71,plain,
    ( frontsegP(app(X1,X2),X3)
    | ~ frontsegP(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,negated_conjecture,
    ( esk52_0 = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_68]),c_0_31])])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( app(esk48_0,nil) = esk48_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_55]),c_0_26]),c_0_56]),c_0_27])])).

<A NAME="c_0_75"></A>cnf(c_0_75,plain,
    ( frontsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X2,X3) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,plain,
    ( nil = app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != X1
    | nil != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>,c_0_72]),c_0_56])])).

<A NAME="c_0_78"></A>cnf(c_0_78,negated_conjecture,
    ( app(esk49_0,nil) = esk49_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,<A HREF="#c_0_73">c_0_73</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( rearsegP(esk48_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_74]),c_0_35]),c_0_56])])).

<A NAME="c_0_80"></A>fof(c_0_80,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X5,X6) != app(X5,X4)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_75]),<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_76])]),c_0_56])])).

<A NAME="c_0_83"></A>cnf(c_0_83,plain,
    ( app(X2,esk6_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( frontsegP(esk49_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_78]),c_0_56]),c_0_58])])).

<A NAME="c_0_85"></A>cnf(c_0_85,plain,
    ( ssList(esk6_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( app(esk49_0,app(nil,X1)) = app(esk49_0,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_78]),c_0_56]),c_0_58])])).

<A NAME="c_0_87"></A>fof(c_0_87,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( app(esk7_2(esk48_0,nil),nil) = esk48_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_79]),c_0_56]),c_0_35])])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( ssList(esk7_2(esk48_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_79]),c_0_56]),c_0_35])])).

<A NAME="c_0_90"></A>cnf(c_0_90,plain,
    ( X1 = X2
    | app(X3,X1) != app(X3,X2)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>])).

<A NAME="c_0_91"></A>cnf(c_0_91,negated_conjecture,
    ( frontsegP(esk49_0,esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_53]),c_0_27]),c_0_31])])).

<A NAME="c_0_92"></A>cnf(c_0_92,plain,
    ( app(nil,app(nil,X1)) = app(nil,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_82]),c_0_56])])).

<A NAME="c_0_93"></A>cnf(c_0_93,negated_conjecture,
    ( app(nil,esk6_2(esk49_0,nil)) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_84]),c_0_56]),c_0_58])])).

<A NAME="c_0_94"></A>cnf(c_0_94,negated_conjecture,
    ( ssList(esk6_2(esk49_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_84]),c_0_56]),c_0_58])])).

<A NAME="c_0_95"></A>cnf(c_0_95,negated_conjecture,
    ( ssList(app(esk49_0,X1))
    | ~ ssList(app(nil,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_86]),c_0_58])])).

<A NAME="c_0_96"></A>cnf(c_0_96,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,negated_conjecture,
    ( esk7_2(esk48_0,nil) = esk48_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_88]),c_0_89])])).

<A NAME="c_0_98"></A>fof(c_0_98,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).

<A NAME="c_0_99"></A>cnf(c_0_99,negated_conjecture,
    ( X1 = esk54_0
    | app(esk52_0,X1) != esk49_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,c_0_53]),c_0_31]),c_0_27])])).

<A NAME="c_0_100"></A>cnf(c_0_100,negated_conjecture,
    ( app(esk52_0,esk6_2(esk49_0,esk52_0)) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_91]),c_0_31]),c_0_58])])).

<A NAME="c_0_101"></A>cnf(c_0_101,negated_conjecture,
    ( ssList(esk6_2(esk49_0,esk52_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_91]),c_0_31]),c_0_58])])).

<A NAME="c_0_102"></A>cnf(c_0_102,negated_conjecture,
    ( app(nil,esk49_0) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_93]),c_0_94])])).

<A NAME="c_0_103"></A>cnf(c_0_103,negated_conjecture,
    ( ssList(app(esk49_0,X1))
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,<A HREF="#c_0_96">c_0_96</A>])).

<A NAME="c_0_104"></A>cnf(c_0_104,negated_conjecture,
    ( app(esk52_0,esk53_0) = esk48_0 ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_65]),<A HREF="#c_0_97">c_0_97</A>])).

<A NAME="c_0_105"></A>cnf(c_0_105,plain,
    ( frontsegP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_98">c_0_98</A>])).

<A NAME="c_0_106"></A>cnf(c_0_106,negated_conjecture,
    ( esk6_2(esk49_0,esk52_0) = esk54_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_100]),c_0_101])])).

<A NAME="c_0_107"></A>cnf(c_0_107,negated_conjecture,
    ( segmentP(app(esk49_0,X1),esk49_0)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_102]),c_0_56]),c_0_58])]),<A HREF="#c_0_103">c_0_103</A>])).

<A NAME="c_0_108"></A>cnf(c_0_108,negated_conjecture,
    ( app(esk49_0,esk53_0) = esk48_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>,<A HREF="#c_0_73">c_0_73</A>])).

<A NAME="c_0_109"></A>cnf(c_0_109,plain,
    ( app(X1,esk6_2(X1,X1)) = X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,<A HREF="#c_0_105">c_0_105</A>])).

<A NAME="c_0_110"></A>cnf(c_0_110,negated_conjecture,
    ( esk6_2(esk49_0,esk49_0) = nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>,c_0_65]),<A HREF="#c_0_73">c_0_73</A>])).

<A NAME="c_0_111"></A>cnf(c_0_111,negated_conjecture,
    ( segmentP(esk48_0,esk49_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>,c_0_108]),c_0_30])])).

<A NAME="c_0_112"></A>cnf(c_0_112,negated_conjecture,
    ( segmentP(esk49_0,esk49_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>,c_0_109]),c_0_110]),c_0_56]),c_0_58])])).

<A NAME="c_0_113"></A>cnf(c_0_113,negated_conjecture,
    ( neq(esk49_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_114"></A>cnf(c_0_114,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,c_0_111]),c_0_112]),c_0_58]),c_0_113])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC083+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n153.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:35:08 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.38  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.38  # and selection function SelectCQPrecWNTNp.
% 0.06/0.38  #
% 0.06/0.38  # Preprocessing time       : 0.014 s
% 0.06/0.38  # Presaturation interreduction done
% 0.06/0.38  
% 0.06/0.38  # Proof found!
% 0.06/0.38  # SZS status Theorem
% 0.06/0.38  # SZS output start CNFRefutation
% 0.06/0.38  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~(ssList(X4))|X2!=X4)|X1!=X3)|~(neq(X2,nil)))|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))|![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>![X8]:((~(ssList(X8))|app(app(X6,X7),X8)!=X3)|app(X6,X8)!=X4)))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.06/0.38  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax6)).
% 0.06/0.38  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.06/0.38  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax79)).
% 0.06/0.38  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.06/0.38  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax7)).
% 0.06/0.38  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax84)).
% 0.06/0.38  fof(c_0_7, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.38  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.06/0.38  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(frontsegP(X1,X2)=>frontsegP(app(X1,X3),X2)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax43)).
% 0.06/0.38  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax45)).
% 0.06/0.38  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax5)).
% 0.06/0.38  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.06/0.38  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X2,X3)=app(X2,X1)=>X3=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax80)).
% 0.06/0.38  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
% 0.06/0.38  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax42)).
% 0.06/0.38  fof(c_0_16, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~ssList(X4)|X2!=X4)|X1!=X3)|~neq(X2,nil))|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))|![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>![X8]:((~ssList(X8)|app(app(X6,X7),X8)!=X3)|app(X6,X8)!=X4))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.38  fof(c_0_17, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
% 0.06/0.38  fof(c_0_18, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.06/0.38  fof(c_0_19, negated_conjecture, (![X13]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&neq(esk49_0,nil))&(((~ssList(X13)|~neq(X13,nil))|~segmentP(esk49_0,X13))|~segmentP(esk48_0,X13)))&(ssList(esk52_0)&(ssList(esk53_0)&((ssList(esk54_0)&app(app(esk52_0,esk53_0),esk54_0)=esk50_0)&app(esk52_0,esk54_0)=esk51_0)))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])).
% 0.06/0.38  cnf(c_0_20, plain, (rearsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X3,X2)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
% 0.06/0.38  cnf(c_0_21, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
% 0.06/0.38  cnf(c_0_22, negated_conjecture, (app(app(esk52_0,esk53_0),esk54_0)=esk50_0), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.38  cnf(c_0_23, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.38  fof(c_0_24, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
% 0.06/0.38  cnf(c_0_25, plain, (rearsegP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_20]), c_0_21])).
% 0.06/0.38  cnf(c_0_26, negated_conjecture, (app(app(esk52_0,esk53_0),esk54_0)=esk48_0), inference(rw,[status(thm)],[c_0_22, c_0_23])).
% 0.06/0.38  cnf(c_0_27, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.38  cnf(c_0_28, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
% 0.06/0.38  cnf(c_0_29, negated_conjecture, (rearsegP(esk48_0,esk54_0)|~ssList(app(esk52_0,esk53_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
% 0.06/0.38  cnf(c_0_30, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.38  cnf(c_0_31, negated_conjecture, (ssList(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.38  cnf(c_0_32, negated_conjecture, (X1=app(esk52_0,esk53_0)|app(X1,esk54_0)!=esk48_0|~ssList(app(esk52_0,esk53_0))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_26]), c_0_27])])).
% 0.06/0.38  cnf(c_0_33, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
% 0.06/0.38  cnf(c_0_34, negated_conjecture, (rearsegP(esk48_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_21]), c_0_30]), c_0_31])])).
% 0.06/0.38  cnf(c_0_35, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.38  cnf(c_0_36, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
% 0.06/0.38  cnf(c_0_37, negated_conjecture, (X1=app(esk52_0,esk53_0)|app(X1,esk54_0)!=esk48_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_21]), c_0_30]), c_0_31])])).
% 0.06/0.38  cnf(c_0_38, negated_conjecture, (app(esk7_2(esk48_0,esk54_0),esk54_0)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_27]), c_0_35])])).
% 0.06/0.38  cnf(c_0_39, negated_conjecture, (ssList(esk7_2(esk48_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_34]), c_0_27]), c_0_35])])).
% 0.06/0.38  fof(c_0_40, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 0.06/0.38  cnf(c_0_41, negated_conjecture, (esk7_2(esk48_0,esk54_0)=app(esk52_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])])).
% 0.06/0.38  cnf(c_0_42, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
% 0.06/0.38  cnf(c_0_43, negated_conjecture, (ssList(app(esk52_0,esk53_0))), inference(rw,[status(thm)],[c_0_39, c_0_41])).
% 0.06/0.38  fof(c_0_44, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
% 0.06/0.38  cnf(c_0_45, negated_conjecture, (app(app(esk52_0,esk53_0),app(esk54_0,X1))=app(esk48_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_26]), c_0_27])]), c_0_43])])).
% 0.06/0.38  cnf(c_0_46, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
% 0.06/0.38  cnf(c_0_47, negated_conjecture, (ssList(app(esk48_0,X1))|~ssList(app(esk54_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_45]), c_0_43])])).
% 0.06/0.38  cnf(c_0_48, negated_conjecture, (app(esk52_0,esk54_0)=esk51_0), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.38  cnf(c_0_49, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.38  cnf(c_0_50, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_46])).
% 0.06/0.38  cnf(c_0_51, negated_conjecture, (ssList(app(esk48_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_21]), c_0_27])])).
% 0.06/0.38  fof(c_0_52, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
% 0.06/0.38  cnf(c_0_53, negated_conjecture, (app(esk52_0,esk54_0)=esk49_0), inference(rw,[status(thm)],[c_0_48, c_0_49])).
% 0.06/0.38  cnf(c_0_54, negated_conjecture, (segmentP(app(esk48_0,X1),esk54_0)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_26]), c_0_27])]), c_0_43])]), c_0_51])).
% 0.06/0.38  cnf(c_0_55, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
% 0.06/0.38  cnf(c_0_56, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.06/0.38  cnf(c_0_57, negated_conjecture, (segmentP(app(esk49_0,X1),esk54_0)|~ssList(app(esk49_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_53]), c_0_31]), c_0_27])])).
% 0.06/0.38  cnf(c_0_58, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.38  cnf(c_0_59, negated_conjecture, (~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.38  cnf(c_0_60, negated_conjecture, (segmentP(esk48_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), c_0_35])])).
% 0.06/0.38  cnf(c_0_61, negated_conjecture, (segmentP(esk49_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_55]), c_0_58]), c_0_56])])).
% 0.06/0.38  fof(c_0_62, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
% 0.06/0.38  cnf(c_0_63, negated_conjecture, (~neq(esk54_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61]), c_0_27])])).
% 0.06/0.38  cnf(c_0_64, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
% 0.06/0.38  cnf(c_0_65, negated_conjecture, (esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_56]), c_0_27])])).
% 0.06/0.38  fof(c_0_66, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~frontsegP(X4,X5)|frontsegP(app(X4,X6),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 0.06/0.38  fof(c_0_67, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
% 0.06/0.38  cnf(c_0_68, negated_conjecture, (app(esk52_0,nil)=esk49_0), inference(rw,[status(thm)],[c_0_53, c_0_65])).
% 0.06/0.38  fof(c_0_69, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
% 0.06/0.38  fof(c_0_70, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
% 0.06/0.38  cnf(c_0_71, plain, (frontsegP(app(X1,X2),X3)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
% 0.06/0.38  cnf(c_0_72, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
% 0.06/0.38  cnf(c_0_73, negated_conjecture, (esk52_0=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_68]), c_0_31])])).
% 0.06/0.38  cnf(c_0_74, negated_conjecture, (app(esk48_0,nil)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_55]), c_0_26]), c_0_56]), c_0_27])])).
% 0.06/0.38  cnf(c_0_75, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_69])).
% 0.06/0.38  cnf(c_0_76, plain, (nil=app(X1,X2)|~ssList(X1)|~ssList(X2)|nil!=X1|nil!=X2), inference(split_conjunct,[status(thm)],[c_0_70])).
% 0.06/0.38  cnf(c_0_77, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_56])])).
% 0.06/0.38  cnf(c_0_78, negated_conjecture, (app(esk49_0,nil)=esk49_0), inference(rw,[status(thm)],[c_0_68, c_0_73])).
% 0.06/0.38  cnf(c_0_79, negated_conjecture, (rearsegP(esk48_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_74]), c_0_35]), c_0_56])])).
% 0.06/0.38  fof(c_0_80, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X5,X6)!=app(X5,X4)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
% 0.06/0.38  cnf(c_0_81, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_75]), c_0_21])).
% 0.06/0.38  cnf(c_0_82, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_76])]), c_0_56])])).
% 0.06/0.38  cnf(c_0_83, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
% 0.06/0.38  cnf(c_0_84, negated_conjecture, (frontsegP(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_56]), c_0_58])])).
% 0.06/0.38  cnf(c_0_85, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
% 0.06/0.38  cnf(c_0_86, negated_conjecture, (app(esk49_0,app(nil,X1))=app(esk49_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_78]), c_0_56]), c_0_58])])).
% 0.06/0.38  fof(c_0_87, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).
% 0.06/0.38  cnf(c_0_88, negated_conjecture, (app(esk7_2(esk48_0,nil),nil)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_79]), c_0_56]), c_0_35])])).
% 0.06/0.38  cnf(c_0_89, negated_conjecture, (ssList(esk7_2(esk48_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_79]), c_0_56]), c_0_35])])).
% 0.06/0.38  cnf(c_0_90, plain, (X1=X2|app(X3,X1)!=app(X3,X2)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_80])).
% 0.06/0.38  cnf(c_0_91, negated_conjecture, (frontsegP(esk49_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_53]), c_0_27]), c_0_31])])).
% 0.06/0.38  cnf(c_0_92, plain, (app(nil,app(nil,X1))=app(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_82]), c_0_56])])).
% 0.06/0.38  cnf(c_0_93, negated_conjecture, (app(nil,esk6_2(esk49_0,nil))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_56]), c_0_58])])).
% 0.06/0.38  cnf(c_0_94, negated_conjecture, (ssList(esk6_2(esk49_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_84]), c_0_56]), c_0_58])])).
% 0.06/0.38  cnf(c_0_95, negated_conjecture, (ssList(app(esk49_0,X1))|~ssList(app(nil,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_86]), c_0_58])])).
% 0.06/0.38  cnf(c_0_96, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
% 0.06/0.38  cnf(c_0_97, negated_conjecture, (esk7_2(esk48_0,nil)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_88]), c_0_89])])).
% 0.06/0.38  fof(c_0_98, plain, (![X2]:(~ssList(X2)|frontsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
% 0.06/0.38  cnf(c_0_99, negated_conjecture, (X1=esk54_0|app(esk52_0,X1)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_53]), c_0_31]), c_0_27])])).
% 0.06/0.38  cnf(c_0_100, negated_conjecture, (app(esk52_0,esk6_2(esk49_0,esk52_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_91]), c_0_31]), c_0_58])])).
% 0.06/0.38  cnf(c_0_101, negated_conjecture, (ssList(esk6_2(esk49_0,esk52_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_91]), c_0_31]), c_0_58])])).
% 0.06/0.38  cnf(c_0_102, negated_conjecture, (app(nil,esk49_0)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94])])).
% 0.06/0.38  cnf(c_0_103, negated_conjecture, (ssList(app(esk49_0,X1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
% 0.06/0.38  cnf(c_0_104, negated_conjecture, (app(esk52_0,esk53_0)=esk48_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_65]), c_0_97])).
% 0.06/0.38  cnf(c_0_105, plain, (frontsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_98])).
% 0.06/0.38  cnf(c_0_106, negated_conjecture, (esk6_2(esk49_0,esk52_0)=esk54_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101])])).
% 0.06/0.38  cnf(c_0_107, negated_conjecture, (segmentP(app(esk49_0,X1),esk49_0)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_102]), c_0_56]), c_0_58])]), c_0_103])).
% 0.06/0.38  cnf(c_0_108, negated_conjecture, (app(esk49_0,esk53_0)=esk48_0), inference(rw,[status(thm)],[c_0_104, c_0_73])).
% 0.06/0.38  cnf(c_0_109, plain, (app(X1,esk6_2(X1,X1))=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_83, c_0_105])).
% 0.06/0.38  cnf(c_0_110, negated_conjecture, (esk6_2(esk49_0,esk49_0)=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_65]), c_0_73])).
% 0.06/0.38  cnf(c_0_111, negated_conjecture, (segmentP(esk48_0,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_30])])).
% 0.06/0.38  cnf(c_0_112, negated_conjecture, (segmentP(esk49_0,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_109]), c_0_110]), c_0_56]), c_0_58])])).
% 0.06/0.38  cnf(c_0_113, negated_conjecture, (neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.38  cnf(c_0_114, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_111]), c_0_112]), c_0_58]), c_0_113])]), ['proof']).
% 0.06/0.38  # SZS output end CNFRefutation
% 0.06/0.38  # Training examples: 0 positive, 0 negative
% 0.06/0.38  
% 0.06/0.38  # -------------------------------------------------
% 0.06/0.38  # User time                : 0.145 s
% 0.06/0.38  # System time              : 0.008 s
% 0.06/0.38  # Total time               : 0.153 s
% 0.06/0.38  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
