<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC388+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC388+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC388+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC388%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC388+1.p">SWC388+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n127.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:45 EDT 2016

% Result     : Theorem 52.74s
% Output     : CNFRefutation 52.74s
% Verified   : 
% Statistics : Number of formulae       :  118 ( 847 expanded)
%              Number of clauses        :   74 ( 332 expanded)
%              Number of leaves         :   22 ( 219 expanded)
%              Depth                    :   20
%              Number of atoms          :  479 (6156 expanded)
%              Number of equality atoms :  160 (2083 expanded)
%              Maximal formula depth    :   17 (   5 average)
%              Maximal clause size      :   18 (   3 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ~ segmentP(X4,X3)
                    | ? [X5] :
                        ( ssItem(X5)
                        & cons(X5,nil) = X1
                        & memberP(X2,X5) )
                    | ( ~ singletonP(X3)
                      & neq(X4,nil) )
                    | ( nil = X2
                      & nil = X1 ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( singletonP(X1)
      &lt;=> ? [X2] :
            ( ssItem(X2)
            & cons(X2,nil) = X1 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax4)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => cons(hd(X1),tl(X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax78)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => ! [X4] :
                  ( ssItem(X4)
                 => ( cons(X3,X1) = cons(X4,X2)
                   => ( X3 = X4
                      & X2 = X1 ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax19)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssItem(hd(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax22)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil = X1
        | ? [X2] :
            ( ssList(X2)
            & ? [X3] :
                ( ssItem(X3)
                & cons(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax20)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => hd(cons(X2,X1)) = X2 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax23)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(cons(X2,X3),X1)
              &lt;=> ( X1 = X2
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax37)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) != X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax18)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(app(X2,X3),X1)
              &lt;=> ( memberP(X2,X1)
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax36)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssList(tl(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax86)).

<A NAME="c_0_17"></A>fof(c_0_17,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_18"></A>fof(c_0_18,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_19"></A>fof(c_0_19,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( segmentP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax7)).

<A NAME="c_0_20"></A>fof(c_0_20,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_21"></A>fof(c_0_21,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( segmentP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax58)).

<A NAME="c_0_22"></A>fof(c_0_22,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ~ segmentP(X4,X3)
                      | ? [X5] :
                          ( ssItem(X5)
                          & cons(X5,nil) = X1
                          & memberP(X2,X5) )
                      | ( ~ singletonP(X3)
                        & neq(X4,nil) )
                      | ( nil = X2
                        & nil = X1 ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_23"></A>fof(c_0_23,negated_conjecture,(
    ! [X10] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & segmentP(esk51_0,esk50_0)
      & ( ~ ssItem(X10)
        | cons(X10,nil) != esk48_0
        | ~ memberP(esk49_0,X10) )
      & ( singletonP(esk50_0)
        | ~ neq(esk51_0,nil) )
      & ( nil != esk49_0
        | nil != esk48_0 ) ) ),
    inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).

<A NAME="c_0_24"></A>fof(c_0_24,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_25"></A>fof(c_0_25,plain,(
    ! [X3,X5] :
      ( ( ssItem(esk5_1(X3))
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( cons(esk5_1(X3),nil) = X3
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( ~ ssItem(X5)
        | cons(X5,nil) != X3
        | singletonP(X3)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( singletonP(esk50_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_28"></A>fof(c_0_28,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,plain,
    ( cons(esk5_1(X1),nil) = X1
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,plain,
    ( ssItem(esk5_1(X1))
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( singletonP(esk48_0)
    | ~ neq(esk51_0,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_38"></A>fof(c_0_38,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | cons(hd(X2),tl(X2)) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).

<A NAME="c_0_39"></A>cnf(c_0_39,plain,
    ( tl(X1) = nil
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_32]),c_0_33])]),<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( esk51_0 = nil
    | singletonP(esk48_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_36]),c_0_33]),c_0_37])])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_42"></A>fof(c_0_42,plain,(
    ! [X5,X6,X7,X8] :
      ( ( X7 = X8
        | cons(X7,X5) != cons(X8,X6)
        | ~ ssItem(X8)
        | ~ ssItem(X7)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( X6 = X5
        | cons(X7,X5) != cons(X8,X6)
        | ~ ssItem(X8)
        | ~ ssItem(X7)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).

<A NAME="c_0_43"></A>cnf(c_0_43,plain,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( tl(esk48_0) = nil
    | esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_40]),c_0_41])])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssItem(X3)
    | ~ ssItem(X4)
    | cons(X3,X1) != cons(X4,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( cons(hd(esk48_0),nil) = esk48_0
    | esk51_0 = nil
    | esk48_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_44]),c_0_41])])).

<A NAME="c_0_47"></A>fof(c_0_47,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssItem(hd(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( esk48_0 = nil
    | esk51_0 = nil
    | nil = X1
    | cons(X2,X1) != esk48_0
    | ~ ssList(X1)
    | ~ ssItem(hd(esk48_0))
    | ~ ssItem(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_46]),c_0_33])])).

<A NAME="c_0_49"></A>cnf(c_0_49,plain,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_50"></A>fof(c_0_50,plain,(
    ! [X4] :
      ( ( ssList(esk44_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( ssItem(esk45_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( cons(esk45_1(X4),esk44_1(X4)) = X4
        | nil = X4
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( esk51_0 = nil
    | esk48_0 = nil
    | nil = X1
    | cons(X2,X1) != esk48_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_49]),c_0_41])])).

<A NAME="c_0_52"></A>cnf(c_0_52,plain,
    ( nil = X1
    | cons(esk45_1(X1),esk44_1(X1)) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( esk44_1(esk48_0) = nil
    | esk48_0 = nil
    | esk51_0 = nil
    | ~ ssList(esk44_1(esk48_0))
    | ~ ssItem(esk45_1(esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,c_0_52])]),c_0_41])])).

<A NAME="c_0_54"></A>cnf(c_0_54,plain,
    ( nil = X1
    | ssList(esk44_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_55"></A>fof(c_0_55,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | hd(cons(X4,X3)) = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( esk44_1(esk48_0) = nil
    | esk51_0 = nil
    | esk48_0 = nil
    | ~ ssItem(esk45_1(esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_54]),c_0_41])])).

<A NAME="c_0_57"></A>cnf(c_0_57,plain,
    ( nil = X1
    | ssItem(esk45_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,plain,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( esk44_1(esk48_0) = nil
    | esk48_0 = nil
    | esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_57]),c_0_41])])).

<A NAME="c_0_60"></A>cnf(c_0_60,plain,
    ( esk5_1(X1) = hd(X1)
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_32]),c_0_33])]),<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_61"></A>fof(c_0_61,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(cons(X5,X6),X4)
        | X4 = X5
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( X4 != X5
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( ~ memberP(esk49_0,X1)
    | cons(X1,nil) != esk48_0
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( cons(esk45_1(esk48_0),nil) = esk48_0
    | esk51_0 = nil
    | esk48_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_59]),c_0_41])])).

<A NAME="c_0_64"></A>fof(c_0_64,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) != X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( cons(hd(X1),nil) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,<A HREF="#c_0_60">c_0_60</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,plain,
    ( ssItem(hd(X1))
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,<A HREF="#c_0_60">c_0_60</A>])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( nil != esk48_0
    | nil != esk49_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_68"></A>fof(c_0_68,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(app(X5,X6),X4)
        | memberP(X5,X4)
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X5,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( memberP(cons(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | X1 != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>])).

<A NAME="c_0_70"></A>fof(c_0_70,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).

<A NAME="c_0_71"></A>fof(c_0_71,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssItem(X6)
      | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).

<A NAME="c_0_72"></A>fof(c_0_72,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssList(tl(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).

<A NAME="c_0_73"></A>fof(c_0_73,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).

<A NAME="c_0_74"></A>fof(c_0_74,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).

<A NAME="c_0_75"></A>fof(c_0_75,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])).

<A NAME="c_0_76"></A>fof(c_0_76,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk8_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ssList(esk9_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(app(X9,X6),X10) != X5
        | segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])])).

<A NAME="c_0_77"></A>cnf(c_0_77,negated_conjecture,
    ( segmentP(esk51_0,esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_78"></A>cnf(c_0_78,negated_conjecture,
    ( cons(X1,nil) != esk48_0
    | ~ memberP(esk51_0,X1)
    | ~ ssItem(X1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( esk45_1(esk48_0) = hd(esk48_0)
    | esk48_0 = nil
    | esk51_0 = nil
    | ~ ssItem(esk45_1(esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_63]),c_0_33])])).

<A NAME="c_0_80"></A>cnf(c_0_80,plain,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>])).

<A NAME="c_0_81"></A>cnf(c_0_81,negated_conjecture,
    ( cons(hd(esk48_0),nil) = esk48_0
    | esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_40]),c_0_41])])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( esk51_0 = nil
    | ssItem(hd(esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>,c_0_40]),c_0_41])])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( esk48_0 != nil
    | esk51_0 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,plain,
    ( memberP(app(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(X3,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_85"></A>cnf(c_0_85,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>])).

<A NAME="c_0_87"></A>cnf(c_0_87,plain,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,plain,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,plain,
    ( tl(app(X1,X2)) = app(tl(X1),X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>])).

<A NAME="c_0_90"></A>cnf(c_0_90,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>])).

<A NAME="c_0_91"></A>cnf(c_0_91,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>])).

<A NAME="c_0_92"></A>fof(c_0_92,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).

<A NAME="c_0_93"></A>cnf(c_0_93,plain,
    ( app(app(esk8_2(X1,X2),X2),esk9_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ segmentP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_94"></A>cnf(c_0_94,negated_conjecture,
    ( segmentP(esk51_0,esk48_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_95"></A>cnf(c_0_95,plain,
    ( ssList(esk9_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ segmentP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_96"></A>cnf(c_0_96,plain,
    ( ssList(esk8_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ segmentP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,negated_conjecture,
    ( esk48_0 = nil
    | esk51_0 = nil
    | ~ memberP(esk51_0,esk45_1(esk48_0))
    | ~ ssItem(esk45_1(esk48_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_98"></A>cnf(c_0_98,negated_conjecture,
    ( esk45_1(esk48_0) = hd(esk48_0)
    | esk51_0 = nil
    | esk48_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_79">c_0_79</A>,c_0_57]),c_0_41])])).

<A NAME="c_0_99"></A>cnf(c_0_99,negated_conjecture,
    ( esk48_0 != nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,c_0_81]),c_0_33])]),c_0_82]),<A HREF="#c_0_83">c_0_83</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,plain,
    ( memberP(app(X1,cons(X2,X3)),X2)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,c_0_85]),<A HREF="#c_0_86">c_0_86</A>])).

<A NAME="c_0_101"></A>cnf(c_0_101,plain,
    ( cons(hd(X1),app(tl(X1),X2)) = app(X1,X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,c_0_43]),c_0_49]),<A HREF="#c_0_88">c_0_88</A>])).

<A NAME="c_0_102"></A>cnf(c_0_102,plain,
    ( nil = X1
    | ssList(app(tl(X1),X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>,c_0_89]),c_0_90]),<A HREF="#c_0_91">c_0_91</A>])).

<A NAME="c_0_103"></A>cnf(c_0_103,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>])).

<A NAME="c_0_104"></A>cnf(c_0_104,negated_conjecture,
    ( app(app(esk8_2(esk51_0,esk48_0),esk48_0),esk9_2(esk51_0,esk48_0)) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_94]),c_0_41]),c_0_37])])).

<A NAME="c_0_105"></A>cnf(c_0_105,negated_conjecture,
    ( ssList(esk9_2(esk51_0,esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_94]),c_0_41]),c_0_37])])).

<A NAME="c_0_106"></A>cnf(c_0_106,negated_conjecture,
    ( ssList(esk8_2(esk51_0,esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,c_0_94]),c_0_41]),c_0_37])])).

<A NAME="c_0_107"></A>cnf(c_0_107,negated_conjecture,
    ( esk51_0 = nil
    | esk48_0 = nil
    | ~ memberP(esk51_0,esk45_1(esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_57]),c_0_41])])).

<A NAME="c_0_108"></A>cnf(c_0_108,negated_conjecture,
    ( esk45_1(esk48_0) = hd(esk48_0)
    | esk51_0 = nil ),
    inference(sr,[status(thm)],[<A HREF="#c_0_98">c_0_98</A>,<A HREF="#c_0_99">c_0_99</A>])).

<A NAME="c_0_109"></A>cnf(c_0_109,plain,
    ( nil = X1
    | memberP(app(X2,app(X1,X3)),hd(X1))
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>,c_0_101]),c_0_49]),<A HREF="#c_0_102">c_0_102</A>])).

<A NAME="c_0_110"></A>cnf(c_0_110,negated_conjecture,
    ( app(esk8_2(esk51_0,esk48_0),app(esk48_0,esk9_2(esk51_0,esk48_0))) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_103">c_0_103</A>,c_0_104]),c_0_105]),c_0_41]),c_0_106])])).

<A NAME="c_0_111"></A>cnf(c_0_111,negated_conjecture,
    ( esk51_0 = nil
    | ~ memberP(esk51_0,hd(esk48_0)) ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>,c_0_108]),<A HREF="#c_0_99">c_0_99</A>])).

<A NAME="c_0_112"></A>cnf(c_0_112,negated_conjecture,
    ( memberP(esk51_0,hd(esk48_0)) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>,c_0_110]),c_0_106]),c_0_105]),c_0_41])]),<A HREF="#c_0_99">c_0_99</A>])).

<A NAME="c_0_113"></A>fof(c_0_113,plain,(
    ! [X2] :
      ( ( ~ segmentP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | segmentP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).

<A NAME="c_0_114"></A>cnf(c_0_114,negated_conjecture,
    ( esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_111">c_0_111</A>,c_0_112])])).

<A NAME="c_0_115"></A>cnf(c_0_115,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ segmentP(nil,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_113">c_0_113</A>])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( segmentP(nil,esk48_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>,<A HREF="#c_0_114">c_0_114</A>])).

<A NAME="c_0_117"></A>cnf(c_0_117,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_115">c_0_115</A>,c_0_116]),c_0_41])]),<A HREF="#c_0_99">c_0_99</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC388+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n127.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:43:23 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 52.74/52.94  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 52.74/52.94  # and selection function SelectCQPrecWNTNp.
% 52.74/52.94  #
% 52.74/52.94  # Preprocessing time       : 0.014 s
% 52.74/52.94  # Presaturation interreduction done
% 52.74/52.94  
% 52.74/52.94  # Proof found!
% 52.74/52.94  # SZS status Theorem
% 52.74/52.94  # SZS output start CNFRefutation
% 52.74/52.94  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~(segmentP(X4,X3)))|?[X5]:((ssItem(X5)&cons(X5,nil)=X1)&memberP(X2,X5)))|(~(singletonP(X3))&neq(X4,nil)))|(nil=X2&nil=X1))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 52.74/52.94  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
% 52.74/52.94  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)&lt;=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax4)).
% 52.74/52.94  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
% 52.74/52.94  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 52.74/52.94  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax78)).
% 52.74/52.94  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssItem(X4)=>(cons(X3,X1)=cons(X4,X2)=>(X3=X4&X2=X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax19)).
% 52.74/52.94  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax22)).
% 52.74/52.94  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax20)).
% 52.74/52.94  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax23)).
% 52.74/52.94  fof(c_0_10, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)&lt;=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax37)).
% 52.74/52.94  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax18)).
% 52.74/52.94  fof(c_0_12, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(memberP(app(X2,X3),X1)&lt;=>(memberP(X2,X1)|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax36)).
% 52.74/52.94  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
% 52.74/52.94  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
% 52.74/52.94  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
% 52.74/52.94  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax86)).
% 52.74/52.94  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 52.74/52.94  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 52.74/52.94  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax7)).
% 52.74/52.94  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 52.74/52.94  fof(c_0_21, axiom, (![X1]:(ssList(X1)=>(segmentP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax58)).
% 52.74/52.94  fof(c_0_22, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~segmentP(X4,X3))|?[X5]:((ssItem(X5)&cons(X5,nil)=X1)&memberP(X2,X5)))|(~singletonP(X3)&neq(X4,nil)))|(nil=X2&nil=X1)))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 52.74/52.94  fof(c_0_23, negated_conjecture, (![X10]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((((esk49_0=esk51_0&esk48_0=esk50_0)&segmentP(esk51_0,esk50_0))&((~ssItem(X10)|cons(X10,nil)!=esk48_0)|~memberP(esk49_0,X10)))&(singletonP(esk50_0)|~neq(esk51_0,nil)))&(nil!=esk49_0|nil!=esk48_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).
% 52.74/52.94  fof(c_0_24, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 52.74/52.94  fof(c_0_25, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
% 52.74/52.94  cnf(c_0_26, negated_conjecture, (singletonP(esk50_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 52.74/52.94  cnf(c_0_27, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_23])).
% 52.74/52.94  fof(c_0_28, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
% 52.74/52.94  cnf(c_0_29, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 52.74/52.94  cnf(c_0_30, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_23])).
% 52.74/52.94  cnf(c_0_31, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
% 52.74/52.94  cnf(c_0_32, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
% 52.74/52.94  cnf(c_0_33, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 52.74/52.94  cnf(c_0_34, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
% 52.74/52.94  cnf(c_0_35, negated_conjecture, (singletonP(esk48_0)|~neq(esk51_0,nil)), inference(rw,[status(thm)],[c_0_26, c_0_27])).
% 52.74/52.94  cnf(c_0_36, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
% 52.74/52.94  cnf(c_0_37, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_29, c_0_30])).
% 52.74/52.94  fof(c_0_38, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
% 52.74/52.94  cnf(c_0_39, plain, (tl(X1)=nil|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])]), c_0_34])).
% 52.74/52.94  cnf(c_0_40, negated_conjecture, (esk51_0=nil|singletonP(esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_33]), c_0_37])])).
% 52.74/52.94  cnf(c_0_41, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 52.74/52.94  fof(c_0_42, plain, (![X5]:![X6]:![X7]:![X8]:((((((X7=X8|cons(X7,X5)!=cons(X8,X6))|~ssItem(X8))|~ssItem(X7))|~ssList(X6))|~ssList(X5))&(((((X6=X5|cons(X7,X5)!=cons(X8,X6))|~ssItem(X8))|~ssItem(X7))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
% 52.74/52.94  cnf(c_0_43, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
% 52.74/52.94  cnf(c_0_44, negated_conjecture, (tl(esk48_0)=nil|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])])).
% 52.74/52.94  cnf(c_0_45, plain, (X2=X1|~ssList(X1)|~ssList(X2)|~ssItem(X3)|~ssItem(X4)|cons(X3,X1)!=cons(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
% 52.74/52.94  cnf(c_0_46, negated_conjecture, (cons(hd(esk48_0),nil)=esk48_0|esk51_0=nil|esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_41])])).
% 52.74/52.94  fof(c_0_47, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
% 52.74/52.94  cnf(c_0_48, negated_conjecture, (esk48_0=nil|esk51_0=nil|nil=X1|cons(X2,X1)!=esk48_0|~ssList(X1)|~ssItem(hd(esk48_0))|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_33])])).
% 52.74/52.94  cnf(c_0_49, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
% 52.74/52.94  fof(c_0_50, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
% 52.74/52.94  cnf(c_0_51, negated_conjecture, (esk51_0=nil|esk48_0=nil|nil=X1|cons(X2,X1)!=esk48_0|~ssList(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_41])])).
% 52.74/52.94  cnf(c_0_52, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
% 52.74/52.94  cnf(c_0_53, negated_conjecture, (esk44_1(esk48_0)=nil|esk48_0=nil|esk51_0=nil|~ssList(esk44_1(esk48_0))|~ssItem(esk45_1(esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52])]), c_0_41])])).
% 52.74/52.94  cnf(c_0_54, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
% 52.74/52.94  fof(c_0_55, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 52.74/52.94  cnf(c_0_56, negated_conjecture, (esk44_1(esk48_0)=nil|esk51_0=nil|esk48_0=nil|~ssItem(esk45_1(esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_41])])).
% 52.74/52.94  cnf(c_0_57, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
% 52.74/52.94  cnf(c_0_58, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
% 52.74/52.94  cnf(c_0_59, negated_conjecture, (esk44_1(esk48_0)=nil|esk48_0=nil|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_41])])).
% 52.74/52.94  cnf(c_0_60, plain, (esk5_1(X1)=hd(X1)|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_32]), c_0_33])]), c_0_34])).
% 52.74/52.94  fof(c_0_61, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).
% 52.74/52.94  cnf(c_0_62, negated_conjecture, (~memberP(esk49_0,X1)|cons(X1,nil)!=esk48_0|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 52.74/52.94  cnf(c_0_63, negated_conjecture, (cons(esk45_1(esk48_0),nil)=esk48_0|esk51_0=nil|esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_59]), c_0_41])])).
% 52.74/52.94  fof(c_0_64, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
% 52.74/52.94  cnf(c_0_65, plain, (cons(hd(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), inference(spm,[status(thm)],[c_0_32, c_0_60])).
% 52.74/52.94  cnf(c_0_66, plain, (ssItem(hd(X1))|~singletonP(X1)|~ssList(X1)), inference(spm,[status(thm)],[c_0_34, c_0_60])).
% 52.74/52.94  cnf(c_0_67, negated_conjecture, (nil!=esk48_0|nil!=esk49_0), inference(split_conjunct,[status(thm)],[c_0_23])).
% 52.74/52.94  fof(c_0_68, plain, (![X4]:![X5]:![X6]:(((((~memberP(app(X5,X6),X4)|(memberP(X5,X4)|memberP(X6,X4)))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&(((((~memberP(X5,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
% 52.74/52.94  cnf(c_0_69, plain, (memberP(cons(X2,X3),X1)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_61])).
% 52.74/52.94  fof(c_0_70, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
% 52.74/52.94  fof(c_0_71, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
% 52.74/52.94  fof(c_0_72, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
% 52.74/52.94  fof(c_0_73, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
% 52.74/52.94  fof(c_0_74, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
% 52.74/52.94  fof(c_0_75, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])).
% 52.74/52.94  fof(c_0_76, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])])).
% 52.74/52.94  cnf(c_0_77, negated_conjecture, (segmentP(esk51_0,esk50_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 52.74/52.94  cnf(c_0_78, negated_conjecture, (cons(X1,nil)!=esk48_0|~memberP(esk51_0,X1)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_62, c_0_30])).
% 52.74/52.94  cnf(c_0_79, negated_conjecture, (esk45_1(esk48_0)=hd(esk48_0)|esk48_0=nil|esk51_0=nil|~ssItem(esk45_1(esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_63]), c_0_33])])).
% 52.74/52.94  cnf(c_0_80, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
% 52.74/52.94  cnf(c_0_81, negated_conjecture, (cons(hd(esk48_0),nil)=esk48_0|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_40]), c_0_41])])).
% 52.74/52.94  cnf(c_0_82, negated_conjecture, (esk51_0=nil|ssItem(hd(esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_40]), c_0_41])])).
% 52.74/52.94  cnf(c_0_83, negated_conjecture, (esk48_0!=nil|esk51_0!=nil), inference(rw,[status(thm)],[c_0_67, c_0_30])).
% 52.74/52.94  cnf(c_0_84, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
% 52.74/52.94  cnf(c_0_85, plain, (memberP(cons(X1,X2),X1)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_69])).
% 52.74/52.94  cnf(c_0_86, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
% 52.74/52.94  cnf(c_0_87, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_71])).
% 52.74/52.94  cnf(c_0_88, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
% 52.74/52.94  cnf(c_0_89, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
% 52.74/52.94  cnf(c_0_90, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
% 52.74/52.94  cnf(c_0_91, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
% 52.74/52.94  fof(c_0_92, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
% 52.74/52.94  cnf(c_0_93, plain, (app(app(esk8_2(X1,X2),X2),esk9_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
% 52.74/52.94  cnf(c_0_94, negated_conjecture, (segmentP(esk51_0,esk48_0)), inference(rw,[status(thm)],[c_0_77, c_0_27])).
% 52.74/52.94  cnf(c_0_95, plain, (ssList(esk9_2(X1,X2))|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
% 52.74/52.94  cnf(c_0_96, plain, (ssList(esk8_2(X1,X2))|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
% 52.74/52.94  cnf(c_0_97, negated_conjecture, (esk48_0=nil|esk51_0=nil|~memberP(esk51_0,esk45_1(esk48_0))|~ssItem(esk45_1(esk48_0))), inference(spm,[status(thm)],[c_0_78, c_0_63])).
% 52.74/52.94  cnf(c_0_98, negated_conjecture, (esk45_1(esk48_0)=hd(esk48_0)|esk51_0=nil|esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_57]), c_0_41])])).
% 52.74/52.94  cnf(c_0_99, negated_conjecture, (esk48_0!=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_33])]), c_0_82]), c_0_83])).
% 52.74/52.94  cnf(c_0_100, plain, (memberP(app(X1,cons(X2,X3)),X2)|~ssList(X1)|~ssList(X3)|~ssItem(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86])).
% 52.74/52.94  cnf(c_0_101, plain, (cons(hd(X1),app(tl(X1),X2))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_43]), c_0_49]), c_0_88])).
% 52.74/52.94  cnf(c_0_102, plain, (nil=X1|ssList(app(tl(X1),X2))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90]), c_0_91])).
% 52.74/52.94  cnf(c_0_103, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
% 52.74/52.94  cnf(c_0_104, negated_conjecture, (app(app(esk8_2(esk51_0,esk48_0),esk48_0),esk9_2(esk51_0,esk48_0))=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_41]), c_0_37])])).
% 52.74/52.94  cnf(c_0_105, negated_conjecture, (ssList(esk9_2(esk51_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_94]), c_0_41]), c_0_37])])).
% 52.74/52.94  cnf(c_0_106, negated_conjecture, (ssList(esk8_2(esk51_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_94]), c_0_41]), c_0_37])])).
% 52.74/52.94  cnf(c_0_107, negated_conjecture, (esk51_0=nil|esk48_0=nil|~memberP(esk51_0,esk45_1(esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_57]), c_0_41])])).
% 52.74/52.94  cnf(c_0_108, negated_conjecture, (esk45_1(esk48_0)=hd(esk48_0)|esk51_0=nil), inference(sr,[status(thm)],[c_0_98, c_0_99])).
% 52.74/52.94  cnf(c_0_109, plain, (nil=X1|memberP(app(X2,app(X1,X3)),hd(X1))|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_49]), c_0_102])).
% 52.74/52.94  cnf(c_0_110, negated_conjecture, (app(esk8_2(esk51_0,esk48_0),app(esk48_0,esk9_2(esk51_0,esk48_0)))=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_105]), c_0_41]), c_0_106])])).
% 52.74/52.94  cnf(c_0_111, negated_conjecture, (esk51_0=nil|~memberP(esk51_0,hd(esk48_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_99])).
% 52.74/52.94  cnf(c_0_112, negated_conjecture, (memberP(esk51_0,hd(esk48_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_106]), c_0_105]), c_0_41])]), c_0_99])).
% 52.74/52.94  fof(c_0_113, plain, (![X2]:(((~segmentP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|segmentP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
% 52.74/52.94  cnf(c_0_114, negated_conjecture, (esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_112])])).
% 52.74/52.94  cnf(c_0_115, plain, (nil=X1|~ssList(X1)|~segmentP(nil,X1)), inference(split_conjunct,[status(thm)],[c_0_113])).
% 52.74/52.94  cnf(c_0_116, negated_conjecture, (segmentP(nil,esk48_0)), inference(rw,[status(thm)],[c_0_94, c_0_114])).
% 52.74/52.94  cnf(c_0_117, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_41])]), c_0_99]), ['proof']).
% 52.74/52.94  # SZS output end CNFRefutation
% 52.74/52.94  # Training examples: 0 positive, 0 negative
% 52.84/53.06  
% 52.84/53.06  # -------------------------------------------------
% 52.84/53.06  # User time                : 51.672 s
% 52.84/53.06  # System time              : 1.187 s
% 52.84/53.06  # Total time               : 52.858 s
% 52.84/53.06  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
