<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC297-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC297-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC297-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC297-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC297-1.p">SWC297-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n148.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:23 EDT 2016

% Result     : Unsatisfiable 0.07s
% Output     : CNFRefutation 0.07s
% Verified   : 
% Statistics : Number of clauses        :  119 (3082 expanded)
%              Number of leaves         :   44 (2039 expanded)
%              Depth                    :   20
%              Number of atoms          :  295 (5879 expanded)
%              Number of equality atoms :  104 (1301 expanded)
%              Maximal clause size      :    7 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( app(app(sk6,cons(sk5,nil)),sk7) = sk1 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( ssList(sk6) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssList(sk7) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ssItem(sk5) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( rearsegP(X3,X2)
    | app(X1,X2) != X3
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause143)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( X2 = X3
    | app(X1,X2) != app(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause150)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( X1 = X3
    | app(X1,X2) != app(X3,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause151)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( app(skaf46(X1,X2),X2) = X1
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause131)).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( cons(sk9,nil) = sk3
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_18)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( ssItem(sk9)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_17)).

<A NAME="c_0_23"></A>cnf(c_0_23,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_24"></A>cnf(c_0_24,axiom,
    ( ssList(skaf46(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause50)).

<A NAME="c_0_25"></A>cnf(c_0_25,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_26"></A>cnf(c_0_26,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_27"></A>cnf(c_0_27,axiom,
    ( memberP(X2,X3)
    | X3 = X1
    | ~ memberP(cons(X1,X2),X3)
    | ~ ssList(X2)
    | ~ ssItem(X1)
    | ~ ssItem(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause161)).

<A NAME="c_0_28"></A>cnf(c_0_28,axiom,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause71)).

<A NAME="c_0_29"></A>cnf(c_0_29,axiom,
    ( memberP(X4,X2)
    | app(X1,cons(X2,X3)) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause175)).

<A NAME="c_0_30"></A>cnf(c_0_30,axiom,
    ( frontsegP(X3,X1)
    | app(X1,X2) != X3
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause144)).

<A NAME="c_0_31"></A>cnf(c_0_31,axiom,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause104)).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( nil = sk3
    | nil != sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_16)).

<A NAME="c_0_33"></A>cnf(c_0_33,axiom,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause97)).

<A NAME="c_0_34"></A>cnf(c_0_34,axiom,
    ( X2 = X1
    | ~ segmentP(X1,X2)
    | ~ segmentP(X2,X1)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause127)).

<A NAME="c_0_35"></A>cnf(c_0_35,axiom,
    ( rearsegP(app(X3,X1),X2)
    | ~ rearsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause136)).

<A NAME="c_0_36"></A>cnf(c_0_36,axiom,
    ( rearsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause58)).

<A NAME="c_0_37"></A>cnf(c_0_37,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_38"></A>cnf(c_0_38,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_39"></A>cnf(c_0_39,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_40"></A>cnf(c_0_40,axiom,
    ( X4 = X2
    | cons(X1,X2) != cons(X3,X4)
    | ~ ssItem(X3)
    | ~ ssItem(X1)
    | ~ ssList(X4)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause171)).

<A NAME="c_0_41"></A>cnf(c_0_41,axiom,
    ( nil = X1
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause118)).

<A NAME="c_0_42"></A>cnf(c_0_42,axiom,
    ( nil = X2
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause119)).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( memberP(sk6,sk8)
    | memberP(sk7,sk8) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( app(app(sk6,cons(sk5,nil)),sk7) = sk3 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( app(sk6,app(cons(sk5,nil),sk7)) = sk3
    | ~ ssList(cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_2]),c_0_3]),c_0_4])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( app(sk6,app(cons(sk5,nil),sk7)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_5]),c_0_6]),c_0_7])])).

<A NAME="c_0_47"></A>cnf(c_0_47,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( app(sk6,cons(sk5,sk7)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_9]),c_0_6]),c_0_4])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( ~ ssList(app(sk3,cons(sk5,X1)))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(sk3,cons(sk5,X1))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_48]),c_0_6]),c_0_4]),c_0_3])])).

<A NAME="c_0_50"></A>cnf(c_0_50,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_10]),<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(app(sk3,cons(sk5,X2)))
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( rearsegP(sk3,sk7)
    | ~ ssList(app(sk6,cons(sk5,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_44]),c_0_4])])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(cons(sk5,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,c_0_11]),c_0_52])])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( rearsegP(sk3,sk7)
    | ~ ssList(cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_11]),c_0_3])])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_5]),c_0_6])])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( rearsegP(sk3,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_5]),c_0_6]),c_0_7])])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( X1 = app(cons(sk5,nil),sk7)
    | app(sk6,X1) != sk3
    | ~ ssList(app(cons(sk5,nil),sk7))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_46]),c_0_3])])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( X1 = app(sk6,cons(sk5,nil))
    | app(X1,sk7) != sk3
    | ~ ssList(app(sk6,cons(sk5,nil)))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_44]),c_0_4])])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( app(skaf46(sk3,sk7),sk7) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_57]),c_0_4]),c_0_52])])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( cons(sk9,nil) = sk3
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_21]),c_0_58]),c_0_7])])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( sk4 = nil
    | ssItem(sk9) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_21]),c_0_58]),c_0_7])])).

<A NAME="c_0_65"></A>cnf(c_0_65,negated_conjecture,
    ( app(cons(sk5,nil),sk7) = cons(sk5,sk7)
    | ~ ssList(app(cons(sk5,nil),sk7))
    | ~ ssList(cons(sk5,sk7)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,plain,
    ( ssList(cons(X1,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,c_0_60])])).

<A NAME="c_0_67"></A>cnf(c_0_67,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( skaf46(sk3,sk7) = app(sk6,cons(sk5,nil))
    | ~ ssList(app(sk6,cons(sk5,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_62]),c_0_24])])).

<A NAME="c_0_69"></A>cnf(c_0_69,negated_conjecture,
    ( tl(sk3) = nil
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_63]),c_0_7])]),<A HREF="#c_0_64">c_0_64</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( app(cons(sk5,nil),sk7) = cons(sk5,sk7)
    | ~ ssList(app(cons(sk5,nil),sk7)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_66]),c_0_4])])).

<A NAME="c_0_71"></A>cnf(c_0_71,plain,
    ( segmentP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,c_0_26]),c_0_7])]),<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_72"></A>cnf(c_0_72,negated_conjecture,
    ( skaf46(sk3,sk7) = app(sk6,cons(sk5,nil))
    | ~ ssList(cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_11]),c_0_3])])).

<A NAME="c_0_73"></A>cnf(c_0_73,plain,
    ( X1 = X2
    | memberP(X3,X2)
    | ~ memberP(cons(X1,X3),X2)
    | ~ ssList(X3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_60]),c_0_60])])).

<A NAME="c_0_74"></A>cnf(c_0_74,plain,
    ( ~ memberP(nil,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_60])])).

<A NAME="c_0_75"></A>cnf(c_0_75,plain,
    ( memberP(app(X1,cons(X2,X3)),X2)
    | ~ ssItem(X2)
    | ~ ssList(app(X1,cons(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_30]),<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,negated_conjecture,
    ( cons(hd(sk3),nil) = sk3
    | sk3 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_69]),c_0_52])]),<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_78"></A>cnf(c_0_78,negated_conjecture,
    ( hd(sk3) = sk9
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_63]),c_0_7])]),<A HREF="#c_0_64">c_0_64</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( app(cons(sk5,nil),sk7) = cons(sk5,sk7)
    | ~ ssList(cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>,c_0_11]),c_0_4])])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( segmentP(sk3,skaf46(sk3,sk7)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>,c_0_62]),c_0_4]),c_0_24])])).

<A NAME="c_0_81"></A>cnf(c_0_81,negated_conjecture,
    ( skaf46(sk3,sk7) = app(sk6,cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_66]),c_0_7])])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( sk4 = nil
    | sk9 = X1
    | ~ memberP(sk3,X1) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>,c_0_63]),c_0_7])]),<A HREF="#c_0_74">c_0_74</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( memberP(sk3,sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_48]),c_0_6]),c_0_52]),c_0_4]),c_0_3])])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( frontsegP(sk3,sk6)
    | ~ ssList(cons(sk5,sk7)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>,c_0_48]),c_0_3])])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( cons(sk9,nil) = sk3
    | sk3 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_78]),<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( app(cons(sk5,nil),sk7) = cons(sk5,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_79">c_0_79</A>,c_0_66]),c_0_7])])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( segmentP(sk3,app(sk6,cons(sk5,nil))) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,<A HREF="#c_0_81">c_0_81</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( ssList(app(sk6,cons(sk5,nil))) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,<A HREF="#c_0_81">c_0_81</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( sk9 = sk5
    | sk4 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,<A HREF="#c_0_83">c_0_83</A>])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( frontsegP(sk3,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,c_0_5]),c_0_6]),c_0_4])])).

<A NAME="c_0_91"></A>cnf(c_0_91,negated_conjecture,
    ( sk3 = nil
    | sk9 = X1
    | ~ memberP(sk3,X1) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>,c_0_85]),c_0_7])]),<A HREF="#c_0_74">c_0_74</A>])).

<A NAME="c_0_92"></A>cnf(c_0_92,negated_conjecture,
    ( ssList(cons(sk5,sk7))
    | ~ ssList(cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_86]),c_0_4])])).

<A NAME="c_0_93"></A>cnf(c_0_93,negated_conjecture,
    ( app(sk6,cons(sk5,nil)) = sk3
    | ~ segmentP(app(sk6,cons(sk5,nil)),sk3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_87]),c_0_52]),c_0_88])])).

<A NAME="c_0_94"></A>cnf(c_0_94,negated_conjecture,
    ( cons(sk5,nil) = sk3
    | sk4 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>,<A HREF="#c_0_89">c_0_89</A>])).

<A NAME="c_0_95"></A>cnf(c_0_95,plain,
    ( rearsegP(app(X1,X2),nil)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_36]),c_0_7])])).

<A NAME="c_0_96"></A>cnf(c_0_96,negated_conjecture,
    ( app(sk6,skaf45(sk3,sk6)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_90]),c_0_3]),c_0_52])])).

<A NAME="c_0_97"></A>cnf(c_0_97,negated_conjecture,
    ( sk9 = sk5
    | sk3 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,<A HREF="#c_0_83">c_0_83</A>])).

<A NAME="c_0_98"></A>cnf(c_0_98,negated_conjecture,
    ( ssList(cons(sk5,sk7)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_66]),c_0_7])])).

<A NAME="c_0_99"></A>cnf(c_0_99,negated_conjecture,
    ( app(sk6,sk3) = sk3
    | sk4 = nil
    | ~ segmentP(app(sk6,sk3),sk3) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,<A HREF="#c_0_94">c_0_94</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,plain,
    ( segmentP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,c_0_38]),c_0_7])]),<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_101"></A>cnf(c_0_101,negated_conjecture,
    ( rearsegP(sk3,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_96]),c_0_3]),c_0_39])])).

<A NAME="c_0_102"></A>cnf(c_0_102,negated_conjecture,
    ( cons(sk5,nil) = sk3
    | sk3 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,<A HREF="#c_0_97">c_0_97</A>])).

<A NAME="c_0_103"></A>cnf(c_0_103,plain,
    ( X1 = X2
    | cons(X3,X1) != cons(X4,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_60]),c_0_60])])).

<A NAME="c_0_104"></A>cnf(c_0_104,negated_conjecture,
    ( X1 = cons(sk5,sk7)
    | app(sk6,X1) != sk3
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,c_0_86]),c_0_86]),c_0_98])])).

<A NAME="c_0_105"></A>cnf(c_0_105,negated_conjecture,
    ( app(sk6,sk3) = sk3
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_100]),c_0_52]),c_0_3])])).

<A NAME="c_0_106"></A>cnf(c_0_106,negated_conjecture,
    ( app(skaf46(sk3,nil),nil) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_101]),c_0_7]),c_0_52])])).

<A NAME="c_0_107"></A>cnf(c_0_107,plain,
    ( X1 = nil
    | app(X1,X2) != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_26]),c_0_7])])).

<A NAME="c_0_108"></A>cnf(c_0_108,negated_conjecture,
    ( app(sk6,app(sk3,sk7)) = sk3
    | sk3 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,<A HREF="#c_0_102">c_0_102</A>])).

<A NAME="c_0_109"></A>cnf(c_0_109,negated_conjecture,
    ( sk6 = nil
    | sk3 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_96]),c_0_39]),c_0_3])])).

<A NAME="c_0_110"></A>cnf(c_0_110,negated_conjecture,
    ( sk3 = nil
    | X1 = nil
    | cons(X2,X1) != sk3
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_103">c_0_103</A>,c_0_85]),c_0_7])])).

<A NAME="c_0_111"></A>cnf(c_0_111,negated_conjecture,
    ( cons(sk5,sk7) = sk3
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>,c_0_105]),c_0_52])])).

<A NAME="c_0_112"></A>cnf(c_0_112,negated_conjecture,
    ( sk7 = nil
    | sk3 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_62]),c_0_4]),c_0_24])])).

<A NAME="c_0_113"></A>cnf(c_0_113,negated_conjecture,
    ( skaf46(sk3,nil) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_106]),c_0_24])])).

<A NAME="c_0_114"></A>cnf(c_0_114,negated_conjecture,
    ( sk6 = nil
    | app(sk3,sk7) != sk3
    | ~ ssList(app(sk3,sk7)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>,c_0_108]),c_0_3])]),<A HREF="#c_0_109">c_0_109</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,negated_conjecture,
    ( sk7 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_110">c_0_110</A>,c_0_111]),c_0_4])]),c_0_32]),<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( app(sk3,nil) = sk3 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>,<A HREF="#c_0_113">c_0_113</A>])).

<A NAME="c_0_117"></A>cnf(c_0_117,negated_conjecture,
    ( sk6 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_114">c_0_114</A>,c_0_115]),c_0_116]),c_0_115]),c_0_116]),c_0_52])])).

<A NAME="c_0_118"></A>cnf(c_0_118,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_115]),c_0_74]),c_0_117]),<A HREF="#c_0_74">c_0_74</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC297-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n148.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:39:53 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.07/0.49  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.07/0.49  # and selection function SelectCQArNTEqFirst.
% 0.07/0.49  #
% 0.07/0.49  # Preprocessing time       : 0.009 s
% 0.07/0.49  # Presaturation interreduction done
% 0.07/0.49  
% 0.07/0.49  # Proof found!
% 0.07/0.49  # SZS status Unsatisfiable
% 0.07/0.49  # SZS output start CNFRefutation
% 0.07/0.49  cnf(c_0_0, negated_conjecture, (app(app(sk6,cons(sk5,nil)),sk7)=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
% 0.07/0.49  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 0.07/0.49  cnf(c_0_2, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause149)).
% 0.07/0.49  cnf(c_0_3, negated_conjecture, (ssList(sk6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
% 0.07/0.49  cnf(c_0_4, negated_conjecture, (ssList(sk7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
% 0.07/0.49  cnf(c_0_5, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.07/0.49  cnf(c_0_6, negated_conjecture, (ssItem(sk5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 0.07/0.49  cnf(c_0_7, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.07/0.49  cnf(c_0_8, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause179)).
% 0.07/0.49  cnf(c_0_9, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause120)).
% 0.07/0.49  cnf(c_0_10, axiom, (rearsegP(X3,X2)|app(X1,X2)!=X3|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause143)).
% 0.07/0.49  cnf(c_0_11, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.07/0.49  cnf(c_0_12, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause72)).
% 0.07/0.49  cnf(c_0_13, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 0.07/0.49  cnf(c_0_14, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
% 0.07/0.49  cnf(c_0_15, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 0.07/0.49  cnf(c_0_16, axiom, (X2=X3|app(X1,X2)!=app(X1,X3)|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause150)).
% 0.07/0.49  cnf(c_0_17, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause53)).
% 0.07/0.49  cnf(c_0_18, axiom, (X1=X3|app(X1,X2)!=app(X3,X2)|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause151)).
% 0.07/0.49  cnf(c_0_19, axiom, (app(skaf46(X1,X2),X2)=X1|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause131)).
% 0.07/0.49  cnf(c_0_20, negated_conjecture, (cons(sk9,nil)=sk3|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_18)).
% 0.07/0.49  cnf(c_0_21, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause100)).
% 0.07/0.49  cnf(c_0_22, negated_conjecture, (ssItem(sk9)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_17)).
% 0.07/0.49  cnf(c_0_23, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause173)).
% 0.07/0.49  cnf(c_0_24, axiom, (ssList(skaf46(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause50)).
% 0.07/0.49  cnf(c_0_25, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause96)).
% 0.07/0.49  cnf(c_0_26, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause74)).
% 0.07/0.49  cnf(c_0_27, axiom, (memberP(X2,X3)|X3=X1|~memberP(cons(X1,X2),X3)|~ssList(X2)|~ssItem(X1)|~ssItem(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause161)).
% 0.07/0.49  cnf(c_0_28, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause71)).
% 0.07/0.49  cnf(c_0_29, axiom, (memberP(X4,X2)|app(X1,cons(X2,X3))!=X4|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause175)).
% 0.07/0.49  cnf(c_0_30, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause144)).
% 0.07/0.49  cnf(c_0_31, axiom, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause104)).
% 0.07/0.49  cnf(c_0_32, negated_conjecture, (nil=sk3|nil!=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_16)).
% 0.07/0.49  cnf(c_0_33, axiom, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause97)).
% 0.07/0.49  cnf(c_0_34, axiom, (X2=X1|~segmentP(X1,X2)|~segmentP(X2,X1)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause127)).
% 0.07/0.49  cnf(c_0_35, axiom, (rearsegP(app(X3,X1),X2)|~rearsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause136)).
% 0.07/0.49  cnf(c_0_36, axiom, (rearsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause58)).
% 0.07/0.49  cnf(c_0_37, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause132)).
% 0.07/0.49  cnf(c_0_38, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause73)).
% 0.07/0.49  cnf(c_0_39, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause51)).
% 0.07/0.49  cnf(c_0_40, axiom, (X4=X2|cons(X1,X2)!=cons(X3,X4)|~ssItem(X3)|~ssItem(X1)|~ssList(X4)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause171)).
% 0.07/0.49  cnf(c_0_41, axiom, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause118)).
% 0.07/0.49  cnf(c_0_42, axiom, (nil=X2|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause119)).
% 0.07/0.49  cnf(c_0_43, negated_conjecture, (memberP(sk6,sk8)|memberP(sk7,sk8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_12)).
% 0.07/0.49  cnf(c_0_44, negated_conjecture, (app(app(sk6,cons(sk5,nil)),sk7)=sk3), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 0.07/0.49  cnf(c_0_45, negated_conjecture, (app(sk6,app(cons(sk5,nil),sk7))=sk3|~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_2]), c_0_3]), c_0_4])])).
% 0.07/0.49  cnf(c_0_46, negated_conjecture, (app(sk6,app(cons(sk5,nil),sk7))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_5]), c_0_6]), c_0_7])])).
% 0.07/0.49  cnf(c_0_47, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_8])).
% 0.07/0.49  cnf(c_0_48, negated_conjecture, (app(sk6,cons(sk5,sk7))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_9]), c_0_6]), c_0_4])])).
% 0.07/0.49  cnf(c_0_49, negated_conjecture, (~ssList(app(sk3,cons(sk5,X1)))|~ssList(X1)|~duplicatefreeP(app(sk3,cons(sk5,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_6]), c_0_4]), c_0_3])])).
% 0.07/0.49  cnf(c_0_50, plain, (rearsegP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_10]), c_0_11])).
% 0.07/0.49  cnf(c_0_51, negated_conjecture, (ssItem(X1)|~ssList(app(sk3,cons(sk5,X2)))|~ssList(X2)), inference(spm,[status(thm)],[c_0_49, c_0_12])).
% 0.07/0.49  cnf(c_0_52, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_13, c_0_1])).
% 0.07/0.49  cnf(c_0_53, negated_conjecture, (rearsegP(sk3,sk7)|~ssList(app(sk6,cons(sk5,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_44]), c_0_4])])).
% 0.07/0.49  cnf(c_0_54, negated_conjecture, (ssItem(X1)|~ssList(cons(sk5,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_11]), c_0_52])])).
% 0.07/0.49  cnf(c_0_55, negated_conjecture, (rearsegP(sk3,sk7)|~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_11]), c_0_3])])).
% 0.07/0.49  cnf(c_0_56, negated_conjecture, (ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_5]), c_0_6])])).
% 0.07/0.49  cnf(c_0_57, negated_conjecture, (rearsegP(sk3,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_5]), c_0_6]), c_0_7])])).
% 0.07/0.49  cnf(c_0_58, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_14, c_0_15])).
% 0.07/0.49  cnf(c_0_59, negated_conjecture, (X1=app(cons(sk5,nil),sk7)|app(sk6,X1)!=sk3|~ssList(app(cons(sk5,nil),sk7))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_46]), c_0_3])])).
% 0.07/0.49  cnf(c_0_60, negated_conjecture, (ssItem(X1)), inference(spm,[status(thm)],[c_0_56, c_0_17])).
% 0.07/0.49  cnf(c_0_61, negated_conjecture, (X1=app(sk6,cons(sk5,nil))|app(X1,sk7)!=sk3|~ssList(app(sk6,cons(sk5,nil)))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_44]), c_0_4])])).
% 0.07/0.49  cnf(c_0_62, negated_conjecture, (app(skaf46(sk3,sk7),sk7)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_57]), c_0_4]), c_0_52])])).
% 0.07/0.49  cnf(c_0_63, negated_conjecture, (cons(sk9,nil)=sk3|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_58]), c_0_7])])).
% 0.07/0.49  cnf(c_0_64, negated_conjecture, (sk4=nil|ssItem(sk9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_21]), c_0_58]), c_0_7])])).
% 0.07/0.49  cnf(c_0_65, negated_conjecture, (app(cons(sk5,nil),sk7)=cons(sk5,sk7)|~ssList(app(cons(sk5,nil),sk7))|~ssList(cons(sk5,sk7))), inference(spm,[status(thm)],[c_0_59, c_0_48])).
% 0.07/0.49  cnf(c_0_66, plain, (ssList(cons(X1,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_60])])).
% 0.07/0.49  cnf(c_0_67, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(er,[status(thm)],[c_0_23])).
% 0.07/0.49  cnf(c_0_68, negated_conjecture, (skaf46(sk3,sk7)=app(sk6,cons(sk5,nil))|~ssList(app(sk6,cons(sk5,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_24])])).
% 0.07/0.49  cnf(c_0_69, negated_conjecture, (tl(sk3)=nil|sk4=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_63]), c_0_7])]), c_0_64])).
% 0.07/0.49  cnf(c_0_70, negated_conjecture, (app(cons(sk5,nil),sk7)=cons(sk5,sk7)|~ssList(app(cons(sk5,nil),sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_4])])).
% 0.07/0.49  cnf(c_0_71, plain, (segmentP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_26]), c_0_7])]), c_0_11])).
% 0.07/0.49  cnf(c_0_72, negated_conjecture, (skaf46(sk3,sk7)=app(sk6,cons(sk5,nil))|~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_11]), c_0_3])])).
% 0.07/0.49  cnf(c_0_73, plain, (X1=X2|memberP(X3,X2)|~memberP(cons(X1,X3),X2)|~ssList(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_60]), c_0_60])])).
% 0.07/0.49  cnf(c_0_74, plain, (~memberP(nil,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_60])])).
% 0.07/0.49  cnf(c_0_75, plain, (memberP(app(X1,cons(X2,X3)),X2)|~ssItem(X2)|~ssList(app(X1,cons(X2,X3)))|~ssList(X3)|~ssList(X1)), inference(er,[status(thm)],[c_0_29])).
% 0.07/0.49  cnf(c_0_76, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_30]), c_0_11])).
% 0.07/0.49  cnf(c_0_77, negated_conjecture, (cons(hd(sk3),nil)=sk3|sk3=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_69]), c_0_52])]), c_0_32])).
% 0.07/0.49  cnf(c_0_78, negated_conjecture, (hd(sk3)=sk9|sk4=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_63]), c_0_7])]), c_0_64])).
% 0.07/0.49  cnf(c_0_79, negated_conjecture, (app(cons(sk5,nil),sk7)=cons(sk5,sk7)|~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_11]), c_0_4])])).
% 0.07/0.49  cnf(c_0_80, negated_conjecture, (segmentP(sk3,skaf46(sk3,sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_62]), c_0_4]), c_0_24])])).
% 0.07/0.49  cnf(c_0_81, negated_conjecture, (skaf46(sk3,sk7)=app(sk6,cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_66]), c_0_7])])).
% 0.07/0.49  cnf(c_0_82, negated_conjecture, (sk4=nil|sk9=X1|~memberP(sk3,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_63]), c_0_7])]), c_0_74])).
% 0.07/0.49  cnf(c_0_83, negated_conjecture, (memberP(sk3,sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_48]), c_0_6]), c_0_52]), c_0_4]), c_0_3])])).
% 0.07/0.49  cnf(c_0_84, negated_conjecture, (frontsegP(sk3,sk6)|~ssList(cons(sk5,sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_48]), c_0_3])])).
% 0.07/0.49  cnf(c_0_85, negated_conjecture, (cons(sk9,nil)=sk3|sk3=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_32])).
% 0.07/0.49  cnf(c_0_86, negated_conjecture, (app(cons(sk5,nil),sk7)=cons(sk5,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_66]), c_0_7])])).
% 0.07/0.49  cnf(c_0_87, negated_conjecture, (segmentP(sk3,app(sk6,cons(sk5,nil)))), inference(rw,[status(thm)],[c_0_80, c_0_81])).
% 0.07/0.49  cnf(c_0_88, negated_conjecture, (ssList(app(sk6,cons(sk5,nil)))), inference(spm,[status(thm)],[c_0_24, c_0_81])).
% 0.07/0.49  cnf(c_0_89, negated_conjecture, (sk9=sk5|sk4=nil), inference(spm,[status(thm)],[c_0_82, c_0_83])).
% 0.07/0.49  cnf(c_0_90, negated_conjecture, (frontsegP(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_5]), c_0_6]), c_0_4])])).
% 0.07/0.49  cnf(c_0_91, negated_conjecture, (sk3=nil|sk9=X1|~memberP(sk3,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_85]), c_0_7])]), c_0_74])).
% 0.07/0.49  cnf(c_0_92, negated_conjecture, (ssList(cons(sk5,sk7))|~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_86]), c_0_4])])).
% 0.07/0.49  cnf(c_0_93, negated_conjecture, (app(sk6,cons(sk5,nil))=sk3|~segmentP(app(sk6,cons(sk5,nil)),sk3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_87]), c_0_52]), c_0_88])])).
% 0.07/0.49  cnf(c_0_94, negated_conjecture, (cons(sk5,nil)=sk3|sk4=nil), inference(spm,[status(thm)],[c_0_63, c_0_89])).
% 0.07/0.49  cnf(c_0_95, plain, (rearsegP(app(X1,X2),nil)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_7])])).
% 0.07/0.49  cnf(c_0_96, negated_conjecture, (app(sk6,skaf45(sk3,sk6))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_90]), c_0_3]), c_0_52])])).
% 0.07/0.49  cnf(c_0_97, negated_conjecture, (sk9=sk5|sk3=nil), inference(spm,[status(thm)],[c_0_91, c_0_83])).
% 0.07/0.49  cnf(c_0_98, negated_conjecture, (ssList(cons(sk5,sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_66]), c_0_7])])).
% 0.07/0.49  cnf(c_0_99, negated_conjecture, (app(sk6,sk3)=sk3|sk4=nil|~segmentP(app(sk6,sk3),sk3)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
% 0.07/0.49  cnf(c_0_100, plain, (segmentP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_38]), c_0_7])]), c_0_11])).
% 0.07/0.49  cnf(c_0_101, negated_conjecture, (rearsegP(sk3,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_3]), c_0_39])])).
% 0.07/0.49  cnf(c_0_102, negated_conjecture, (cons(sk5,nil)=sk3|sk3=nil), inference(spm,[status(thm)],[c_0_85, c_0_97])).
% 0.07/0.49  cnf(c_0_103, plain, (X1=X2|cons(X3,X1)!=cons(X4,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_60]), c_0_60])])).
% 0.07/0.49  cnf(c_0_104, negated_conjecture, (X1=cons(sk5,sk7)|app(sk6,X1)!=sk3|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_86]), c_0_86]), c_0_98])])).
% 0.07/0.49  cnf(c_0_105, negated_conjecture, (app(sk6,sk3)=sk3|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_52]), c_0_3])])).
% 0.07/0.49  cnf(c_0_106, negated_conjecture, (app(skaf46(sk3,nil),nil)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_101]), c_0_7]), c_0_52])])).
% 0.07/0.49  cnf(c_0_107, plain, (X1=nil|app(X1,X2)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_26]), c_0_7])])).
% 0.07/0.49  cnf(c_0_108, negated_conjecture, (app(sk6,app(sk3,sk7))=sk3|sk3=nil), inference(spm,[status(thm)],[c_0_46, c_0_102])).
% 0.07/0.49  cnf(c_0_109, negated_conjecture, (sk6=nil|sk3!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_96]), c_0_39]), c_0_3])])).
% 0.07/0.49  cnf(c_0_110, negated_conjecture, (sk3=nil|X1=nil|cons(X2,X1)!=sk3|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_85]), c_0_7])])).
% 0.07/0.49  cnf(c_0_111, negated_conjecture, (cons(sk5,sk7)=sk3|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_52])])).
% 0.07/0.49  cnf(c_0_112, negated_conjecture, (sk7=nil|sk3!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_62]), c_0_4]), c_0_24])])).
% 0.07/0.49  cnf(c_0_113, negated_conjecture, (skaf46(sk3,nil)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_106]), c_0_24])])).
% 0.07/0.49  cnf(c_0_114, negated_conjecture, (sk6=nil|app(sk3,sk7)!=sk3|~ssList(app(sk3,sk7))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_3])]), c_0_109])).
% 0.07/0.49  cnf(c_0_115, negated_conjecture, (sk7=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_4])]), c_0_32]), c_0_112])).
% 0.07/0.49  cnf(c_0_116, negated_conjecture, (app(sk3,nil)=sk3), inference(rw,[status(thm)],[c_0_106, c_0_113])).
% 0.07/0.49  cnf(c_0_117, negated_conjecture, (sk6=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_115]), c_0_116]), c_0_115]), c_0_116]), c_0_52])])).
% 0.07/0.49  cnf(c_0_118, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_115]), c_0_74]), c_0_117]), c_0_74]), ['proof']).
% 0.07/0.49  # SZS output end CNFRefutation
% 0.07/0.49  # Training examples: 0 positive, 0 negative
% 0.07/0.49  
% 0.07/0.49  # -------------------------------------------------
% 0.07/0.49  # User time                : 0.251 s
% 0.07/0.49  # System time              : 0.006 s
% 0.07/0.49  # Total time               : 0.257 s
% 0.07/0.49  # Maximum resident set size: 1864 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
