<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC053+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC053+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC053+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC053%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC053+1.p">SWC053+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n158.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:27 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :  126 (2049 expanded)
%              Number of clauses        :   84 ( 754 expanded)
%              Number of leaves         :   21 ( 503 expanded)
%              Depth                    :   21
%              Number of atoms          :  488 (26641 expanded)
%              Number of equality atoms :  162 (9154 expanded)
%              Maximal formula depth    :   24 (   5 average)
%              Maximal clause size      :   36 (   3 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ( ! [X5] :
                          ( ssList(X5)
                         => ( ~ neq(X5,nil)
                            | ~ frontsegP(X4,X5)
                            | ~ frontsegP(X3,X5) ) )
                      & ( nil != X4
                        | nil != X3 ) )
                    | ( ( nil != X2
                        | nil = X1 )
                      & ( ~ neq(X2,nil)
                        | ? [X6] :
                            ( ssList(X6)
                            & neq(X6,nil)
                            & segmentP(X2,X6)
                            & segmentP(X1,X6) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( frontsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X2,X3) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax5)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => hd(app(X1,X2)) = hd(X1) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax85)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => cons(hd(X1),tl(X1)) = X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax78)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => nil != cons(X2,X1) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax21)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssItem(hd(X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax22)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( frontsegP(cons(X1,X3),cons(X2,X4))
                  &lt;=> ( X1 = X2
                      & frontsegP(X3,X4) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax44)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,X1) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax42)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssList(tl(X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax24)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( segmentP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax7)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( frontsegP(X1,X2)
               => frontsegP(app(X1,X3),X2) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax43)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,nil) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax45)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_17"></A>fof(c_0_17,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( ( frontsegP(X1,X2)
              & frontsegP(X2,X1) )
           => X1 = X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax41)).

<A NAME="c_0_18"></A>fof(c_0_18,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_19"></A>fof(c_0_19,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,nil) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax57)).

<A NAME="c_0_20"></A>fof(c_0_20,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_21"></A>fof(c_0_21,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ( ! [X5] :
                            ( ssList(X5)
                           => ( ~ neq(X5,nil)
                              | ~ frontsegP(X4,X5)
                              | ~ frontsegP(X3,X5) ) )
                        & ( nil != X4
                          | nil != X3 ) )
                      | ( ( nil != X2
                          | nil = X1 )
                        & ( ~ neq(X2,nil)
                          | ? [X6] :
                              ( ssList(X6)
                              & neq(X6,nil)
                              & segmentP(X2,X6)
                              & segmentP(X1,X6) ) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_22"></A>fof(c_0_22,negated_conjecture,(
    ! [X12] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ( nil = esk51_0
        | ssList(esk52_0) )
      & ( nil = esk50_0
        | ssList(esk52_0) )
      & ( nil = esk51_0
        | neq(esk52_0,nil) )
      & ( nil = esk50_0
        | neq(esk52_0,nil) )
      & ( nil = esk51_0
        | frontsegP(esk51_0,esk52_0) )
      & ( nil = esk50_0
        | frontsegP(esk51_0,esk52_0) )
      & ( nil = esk51_0
        | frontsegP(esk50_0,esk52_0) )
      & ( nil = esk50_0
        | frontsegP(esk50_0,esk52_0) )
      & ( neq(esk49_0,nil)
        | nil = esk49_0 )
      & ( ~ ssList(X12)
        | ~ neq(X12,nil)
        | ~ segmentP(esk49_0,X12)
        | ~ segmentP(esk48_0,X12)
        | nil = esk49_0 )
      & ( neq(esk49_0,nil)
        | nil != esk48_0 )
      & ( ~ ssList(X12)
        | ~ neq(X12,nil)
        | ~ segmentP(esk49_0,X12)
        | ~ segmentP(esk48_0,X12)
        | nil != esk48_0 ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).

<A NAME="c_0_23"></A>fof(c_0_23,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk6_2(X4,X5))
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(X5,esk6_2(X4,X5)) = X4
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X5,X7) != X4
        | frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_26"></A>fof(c_0_26,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | hd(app(X3,X4)) = hd(X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_27"></A>cnf(c_0_27,plain,
    ( app(X2,esk6_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( frontsegP(esk51_0,esk52_0)
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( ssList(esk52_0)
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( ssList(esk6_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_32"></A>fof(c_0_32,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | cons(hd(X2),tl(X2)) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

<A NAME="c_0_33"></A>fof(c_0_33,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_34"></A>fof(c_0_34,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_35"></A>fof(c_0_35,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | nil != cons(X4,X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( hd(app(X1,X2)) = hd(X1)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( app(esk52_0,esk6_2(esk51_0,esk52_0)) = esk51_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_28]),c_0_29])]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( esk50_0 = nil
    | ssList(esk6_2(esk51_0,esk52_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_28]),c_0_29])]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_39"></A>fof(c_0_39,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssItem(hd(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

<A NAME="c_0_40"></A>fof(c_0_40,plain,(
    ! [X5,X6,X7,X8] :
      ( ( X5 = X6
        | ~ frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) )
      & ( frontsegP(X7,X8)
        | ~ frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) )
      & ( X5 != X6
        | ~ frontsegP(X7,X8)
        | frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,plain,
    ( nil != cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( hd(esk52_0) = hd(esk51_0)
    | esk50_0 = nil
    | esk52_0 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_37]),c_0_30]),<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,plain,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_47"></A>cnf(c_0_47,plain,
    ( frontsegP(cons(X1,X3),cons(X2,X4))
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ ssList(X4)
    | ~ frontsegP(X3,X4)
    | X1 != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_48"></A>fof(c_0_48,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).

<A NAME="c_0_49"></A>cnf(c_0_49,plain,
    ( cons(hd(cons(X1,X2)),X2) = cons(X1,X2)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_42]),c_0_43]),<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( cons(hd(esk51_0),tl(esk52_0)) = esk52_0
    | esk50_0 = nil
    | esk52_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_45]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( esk50_0 = nil
    | esk52_0 = nil
    | ssItem(hd(esk51_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_45]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_52"></A>fof(c_0_52,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssList(tl(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( frontsegP(cons(X1,X2),cons(X1,X3))
    | ~ frontsegP(X2,X3)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_54"></A>cnf(c_0_54,plain,
    ( frontsegP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( cons(hd(esk52_0),tl(esk52_0)) = esk52_0
    | esk52_0 = nil
    | esk50_0 = nil
    | ~ ssList(tl(esk52_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_50]),<A HREF="#c_0_51">c_0_51</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,plain,
    ( frontsegP(cons(X1,X2),cons(X1,X2))
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( cons(hd(esk52_0),tl(esk52_0)) = esk52_0
    | esk50_0 = nil
    | esk52_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_56]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_59"></A>fof(c_0_59,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk8_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ssList(esk9_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(app(X9,X6),X10) != X5
        | segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( esk52_0 = nil
    | esk50_0 = nil
    | frontsegP(esk52_0,esk52_0)
    | ~ ssList(tl(esk52_0))
    | ~ ssItem(hd(esk52_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,<A HREF="#c_0_58">c_0_58</A>])).

<A NAME="c_0_61"></A>cnf(c_0_61,plain,
    ( segmentP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(app(X3,X2),X4) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>])).

<A NAME="c_0_62"></A>fof(c_0_62,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).

<A NAME="c_0_63"></A>fof(c_0_63,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_66"></A>fof(c_0_66,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ frontsegP(X4,X5)
      | frontsegP(app(X4,X6),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).

<A NAME="c_0_67"></A>fof(c_0_67,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( esk50_0 = nil
    | esk52_0 = nil
    | frontsegP(esk52_0,esk52_0)
    | ~ ssItem(hd(esk52_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>,c_0_56]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>])).

<A NAME="c_0_71"></A>cnf(c_0_71,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,negated_conjecture,
    ( frontsegP(esk50_0,esk52_0)
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,<A HREF="#c_0_65">c_0_65</A>])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( frontsegP(esk51_0,esk52_0)
    | nil = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,negated_conjecture,
    ( ssList(esk52_0)
    | nil = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,plain,
    ( frontsegP(app(X1,X2),X3)
    | ~ frontsegP(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>])).

<A NAME="c_0_78"></A>cnf(c_0_78,plain,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( esk52_0 = nil
    | esk50_0 = nil
    | frontsegP(esk52_0,esk52_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_46]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( nil = esk49_0
    | ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( segmentP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>,c_0_70]),c_0_71])]),<A HREF="#c_0_72">c_0_72</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( app(esk52_0,esk6_2(esk50_0,esk52_0)) = esk50_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_73]),c_0_74])]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( esk50_0 = nil
    | ssList(esk6_2(esk50_0,esk52_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_73]),c_0_74])]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( app(esk52_0,esk6_2(esk51_0,esk52_0)) = esk51_0
    | esk51_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_75]),c_0_29])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( esk51_0 = nil
    | ssList(esk6_2(esk51_0,esk52_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_75]),c_0_29])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_86"></A>fof(c_0_86,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ~ frontsegP(X3,X4)
      | ~ frontsegP(X4,X3)
      | X3 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).

<A NAME="c_0_87"></A>cnf(c_0_87,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_78]),c_0_71])])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( app(esk52_0,esk6_2(esk52_0,esk52_0)) = esk52_0
    | esk50_0 = nil
    | esk52_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_79]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( esk50_0 = nil
    | esk52_0 = nil
    | ssList(esk6_2(esk52_0,esk52_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_79]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( esk51_0 = nil
    | ~ segmentP(esk50_0,X1)
    | ~ segmentP(esk51_0,X1)
    | ~ ssList(X1)
    | ~ neq(X1,nil) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,c_0_25]),c_0_65]),<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_91"></A>cnf(c_0_91,negated_conjecture,
    ( esk50_0 = nil
    | segmentP(esk50_0,esk52_0) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_82]),c_0_30]),<A HREF="#c_0_83">c_0_83</A>])).

<A NAME="c_0_92"></A>cnf(c_0_92,negated_conjecture,
    ( neq(esk52_0,nil)
    | nil = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_93"></A>cnf(c_0_93,negated_conjecture,
    ( esk51_0 = nil
    | segmentP(esk51_0,esk52_0) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_84]),c_0_76]),<A HREF="#c_0_85">c_0_85</A>])).

<A NAME="c_0_94"></A>fof(c_0_94,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])).

<A NAME="c_0_95"></A>cnf(c_0_95,plain,
    ( X1 = X2
    | ~ frontsegP(X2,X1)
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>])).

<A NAME="c_0_96"></A>cnf(c_0_96,negated_conjecture,
    ( esk52_0 = nil
    | esk50_0 = nil
    | frontsegP(esk52_0,nil) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,c_0_88]),c_0_30]),<A HREF="#c_0_89">c_0_89</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,negated_conjecture,
    ( frontsegP(esk50_0,esk52_0)
    | nil = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_98"></A>cnf(c_0_98,negated_conjecture,
    ( esk50_0 = nil
    | esk51_0 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,c_0_91]),c_0_92]),c_0_76]),<A HREF="#c_0_93">c_0_93</A>])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,negated_conjecture,
    ( nil != esk48_0
    | ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_101"></A>fof(c_0_101,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).

<A NAME="c_0_102"></A>fof(c_0_102,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])).

<A NAME="c_0_103"></A>cnf(c_0_103,negated_conjecture,
    ( esk50_0 = nil
    | esk52_0 = nil
    | ~ frontsegP(nil,esk52_0) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_96]),c_0_71])]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_104"></A>cnf(c_0_104,negated_conjecture,
    ( esk51_0 = nil
    | frontsegP(nil,esk52_0) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,<A HREF="#c_0_98">c_0_98</A>])).

<A NAME="c_0_105"></A>cnf(c_0_105,negated_conjecture,
    ( esk50_0 = nil
    | esk52_0 = nil
    | esk51_0 != nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_37]),c_0_30]),<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_106"></A>cnf(c_0_106,negated_conjecture,
    ( esk50_0 != nil
    | ~ segmentP(esk50_0,X1)
    | ~ segmentP(esk51_0,X1)
    | ~ ssList(X1)
    | ~ neq(X1,nil) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>,c_0_65]),c_0_65]),<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_107"></A>cnf(c_0_107,plain,
    ( segmentP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_101">c_0_101</A>])).

<A NAME="c_0_108"></A>cnf(c_0_108,plain,
    ( ~ ssList(X1)
    | ~ ssList(X2)
    | X1 != X2
    | ~ neq(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>])).

<A NAME="c_0_109"></A>cnf(c_0_109,negated_conjecture,
    ( neq(esk52_0,nil)
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_110"></A>cnf(c_0_110,negated_conjecture,
    ( esk52_0 = nil
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_103">c_0_103</A>,c_0_104]),<A HREF="#c_0_105">c_0_105</A>])).

<A NAME="c_0_111"></A>cnf(c_0_111,negated_conjecture,
    ( app(esk52_0,esk6_2(esk50_0,esk52_0)) = esk50_0
    | esk51_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_97]),c_0_74])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_112"></A>cnf(c_0_112,negated_conjecture,
    ( esk51_0 = nil
    | ssList(esk6_2(esk50_0,esk52_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_97]),c_0_74])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_113"></A>cnf(c_0_113,negated_conjecture,
    ( esk50_0 != nil
    | ~ segmentP(esk51_0,nil)
    | ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>,c_0_107]),c_0_71]),c_0_74])])).

<A NAME="c_0_114"></A>cnf(c_0_114,plain,
    ( ~ ssList(X1)
    | ~ neq(X1,X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,negated_conjecture,
    ( esk50_0 = nil
    | neq(nil,nil) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>,<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( neq(esk49_0,nil)
    | nil != esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_117"></A>cnf(c_0_117,negated_conjecture,
    ( esk51_0 = nil
    | esk52_0 = nil
    | esk50_0 != nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_111]),c_0_76]),<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_118"></A>cnf(c_0_118,negated_conjecture,
    ( esk50_0 != nil
    | ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_113">c_0_113</A>,c_0_107]),c_0_29])])).

<A NAME="c_0_119"></A>cnf(c_0_119,negated_conjecture,
    ( esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_114">c_0_114</A>,c_0_115]),c_0_71])])).

<A NAME="c_0_120"></A>cnf(c_0_120,negated_conjecture,
    ( neq(esk51_0,nil)
    | esk50_0 != nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_116">c_0_116</A>,c_0_25]),<A HREF="#c_0_65">c_0_65</A>])).

<A NAME="c_0_121"></A>cnf(c_0_121,negated_conjecture,
    ( esk52_0 = nil
    | esk51_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_117">c_0_117</A>,<A HREF="#c_0_98">c_0_98</A>])).

<A NAME="c_0_122"></A>cnf(c_0_122,negated_conjecture,
    ( ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_118">c_0_118</A>,c_0_119])])).

<A NAME="c_0_123"></A>cnf(c_0_123,negated_conjecture,
    ( neq(esk51_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_120">c_0_120</A>,c_0_119])])).

<A NAME="c_0_124"></A>cnf(c_0_124,negated_conjecture,
    ( esk51_0 = nil ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_121]),<A HREF="#c_0_122">c_0_122</A>])).

<A NAME="c_0_125"></A>cnf(c_0_125,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>,c_0_124]),<A HREF="#c_0_122">c_0_122</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC053+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.22  % Computer   : n158.star.cs.uiowa.edu
% 0.03/0.22  % Model      : x86_64 x86_64
% 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.22  % Memory     : 32218.75MB
% 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.22  % CPULimit   : 300
% 0.03/0.22  % DateTime   : Tue May 10 20:34:53 CDT 2016
% 0.03/0.22  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.47  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.47  # and selection function SelectCQPrecWNTNp.
% 0.06/0.47  #
% 0.06/0.47  # Preprocessing time       : 0.014 s
% 0.06/0.47  # Presaturation interreduction done
% 0.06/0.47  
% 0.06/0.47  # Proof found!
% 0.06/0.47  # SZS status Theorem
% 0.06/0.47  # SZS output start CNFRefutation
% 0.06/0.47  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|(![X5]:(ssList(X5)=>((~(neq(X5,nil))|~(frontsegP(X4,X5)))|~(frontsegP(X3,X5))))&(nil!=X4|nil!=X3)))|((nil!=X2|nil=X1)&(~(neq(X2,nil))|?[X6]:(((ssList(X6)&neq(X6,nil))&segmentP(X2,X6))&segmentP(X1,X6)))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.06/0.47  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax5)).
% 0.06/0.47  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>hd(app(X1,X2))=hd(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax85)).
% 0.06/0.47  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax78)).
% 0.06/0.47  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax25)).
% 0.06/0.47  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.06/0.47  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>nil!=cons(X2,X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax21)).
% 0.06/0.47  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax22)).
% 0.06/0.47  fof(c_0_8, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(frontsegP(cons(X1,X3),cons(X2,X4))&lt;=>(X1=X2&frontsegP(X3,X4)))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax44)).
% 0.06/0.47  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax42)).
% 0.06/0.47  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax24)).
% 0.06/0.47  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax7)).
% 0.06/0.47  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
% 0.06/0.47  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.06/0.47  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(frontsegP(X1,X2)=>frontsegP(app(X1,X3),X2)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax43)).
% 0.06/0.47  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax45)).
% 0.06/0.47  fof(c_0_16, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.47  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>((frontsegP(X1,X2)&frontsegP(X2,X1))=>X1=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax41)).
% 0.06/0.47  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.06/0.47  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>segmentP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax57)).
% 0.06/0.47  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.06/0.47  fof(c_0_21, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|(![X5]:(ssList(X5)=>((~neq(X5,nil)|~frontsegP(X4,X5))|~frontsegP(X3,X5)))&(nil!=X4|nil!=X3)))|((nil!=X2|nil=X1)&(~neq(X2,nil)|?[X6]:(((ssList(X6)&neq(X6,nil))&segmentP(X2,X6))&segmentP(X1,X6))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.47  fof(c_0_22, negated_conjecture, (![X12]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((esk49_0=esk51_0&esk48_0=esk50_0)&(((nil=esk51_0|ssList(esk52_0))&(nil=esk50_0|ssList(esk52_0)))&((((nil=esk51_0|neq(esk52_0,nil))&(nil=esk50_0|neq(esk52_0,nil)))&((nil=esk51_0|frontsegP(esk51_0,esk52_0))&(nil=esk50_0|frontsegP(esk51_0,esk52_0))))&((nil=esk51_0|frontsegP(esk50_0,esk52_0))&(nil=esk50_0|frontsegP(esk50_0,esk52_0))))))&(((neq(esk49_0,nil)|nil=esk49_0)&((((~ssList(X12)|~neq(X12,nil))|~segmentP(esk49_0,X12))|~segmentP(esk48_0,X12))|nil=esk49_0))&((neq(esk49_0,nil)|nil!=esk48_0)&((((~ssList(X12)|~neq(X12,nil))|~segmentP(esk49_0,X12))|~segmentP(esk48_0,X12))|nil!=esk48_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).
% 0.06/0.47  fof(c_0_23, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
% 0.06/0.47  cnf(c_0_24, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_25, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  fof(c_0_26, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|hd(app(X3,X4))=hd(X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.06/0.47  cnf(c_0_27, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 0.06/0.47  cnf(c_0_28, negated_conjecture, (frontsegP(esk51_0,esk52_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_29, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_24, c_0_25])).
% 0.06/0.47  cnf(c_0_30, negated_conjecture, (ssList(esk52_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_31, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 0.06/0.47  fof(c_0_32, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
% 0.06/0.47  fof(c_0_33, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 0.06/0.47  fof(c_0_34, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 0.06/0.47  fof(c_0_35, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|nil!=cons(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
% 0.06/0.47  cnf(c_0_36, plain, (hd(app(X1,X2))=hd(X1)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
% 0.06/0.47  cnf(c_0_37, negated_conjecture, (app(esk52_0,esk6_2(esk51_0,esk52_0))=esk51_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])]), c_0_30])).
% 0.06/0.47  cnf(c_0_38, negated_conjecture, (esk50_0=nil|ssList(esk6_2(esk51_0,esk52_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_28]), c_0_29])]), c_0_30])).
% 0.06/0.47  fof(c_0_39, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
% 0.06/0.47  fof(c_0_40, plain, (![X5]:![X6]:![X7]:![X8]:(((((((X5=X6|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5))&(((((frontsegP(X7,X8)|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))&((((((X5!=X6|~frontsegP(X7,X8))|frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
% 0.06/0.47  cnf(c_0_41, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.47  cnf(c_0_42, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
% 0.06/0.47  cnf(c_0_43, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
% 0.06/0.47  cnf(c_0_44, plain, (nil!=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
% 0.06/0.47  cnf(c_0_45, negated_conjecture, (hd(esk52_0)=hd(esk51_0)|esk50_0=nil|esk52_0=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_30]), c_0_38])).
% 0.06/0.47  cnf(c_0_46, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.47  cnf(c_0_47, plain, (frontsegP(cons(X1,X3),cons(X2,X4))|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~ssList(X4)|~frontsegP(X3,X4)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_40])).
% 0.06/0.47  fof(c_0_48, plain, (![X2]:(~ssList(X2)|frontsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).
% 0.06/0.47  cnf(c_0_49, plain, (cons(hd(cons(X1,X2)),X2)=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44])).
% 0.06/0.47  cnf(c_0_50, negated_conjecture, (cons(hd(esk51_0),tl(esk52_0))=esk52_0|esk50_0=nil|esk52_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_45]), c_0_30])).
% 0.06/0.47  cnf(c_0_51, negated_conjecture, (esk50_0=nil|esk52_0=nil|ssItem(hd(esk51_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_45]), c_0_30])).
% 0.06/0.47  fof(c_0_52, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
% 0.06/0.47  cnf(c_0_53, plain, (frontsegP(cons(X1,X2),cons(X1,X3))|~frontsegP(X2,X3)|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_47])).
% 0.06/0.47  cnf(c_0_54, plain, (frontsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
% 0.06/0.47  cnf(c_0_55, negated_conjecture, (cons(hd(esk52_0),tl(esk52_0))=esk52_0|esk52_0=nil|esk50_0=nil|~ssList(tl(esk52_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])).
% 0.06/0.47  cnf(c_0_56, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
% 0.06/0.47  cnf(c_0_57, plain, (frontsegP(cons(X1,X2),cons(X1,X2))|~ssList(X2)|~ssItem(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
% 0.06/0.47  cnf(c_0_58, negated_conjecture, (cons(hd(esk52_0),tl(esk52_0))=esk52_0|esk50_0=nil|esk52_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_30])).
% 0.06/0.47  fof(c_0_59, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
% 0.06/0.47  cnf(c_0_60, negated_conjecture, (esk52_0=nil|esk50_0=nil|frontsegP(esk52_0,esk52_0)|~ssList(tl(esk52_0))|~ssItem(hd(esk52_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
% 0.06/0.47  cnf(c_0_61, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_59])).
% 0.06/0.47  fof(c_0_62, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).
% 0.06/0.47  fof(c_0_63, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
% 0.06/0.47  cnf(c_0_64, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_65, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  fof(c_0_66, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~frontsegP(X4,X5)|frontsegP(app(X4,X6),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
% 0.06/0.47  fof(c_0_67, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
% 0.06/0.47  cnf(c_0_68, negated_conjecture, (esk50_0=nil|esk52_0=nil|frontsegP(esk52_0,esk52_0)|~ssItem(hd(esk52_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_56]), c_0_30])).
% 0.06/0.47  cnf(c_0_69, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_61])).
% 0.06/0.47  cnf(c_0_70, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
% 0.06/0.47  cnf(c_0_71, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.47  cnf(c_0_72, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
% 0.06/0.47  cnf(c_0_73, negated_conjecture, (frontsegP(esk50_0,esk52_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_74, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_64, c_0_65])).
% 0.06/0.47  cnf(c_0_75, negated_conjecture, (frontsegP(esk51_0,esk52_0)|nil=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_76, negated_conjecture, (ssList(esk52_0)|nil=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_77, plain, (frontsegP(app(X1,X2),X3)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
% 0.06/0.47  cnf(c_0_78, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
% 0.06/0.47  cnf(c_0_79, negated_conjecture, (esk52_0=nil|esk50_0=nil|frontsegP(esk52_0,esk52_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_46]), c_0_30])).
% 0.06/0.47  cnf(c_0_80, negated_conjecture, (nil=esk49_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_81, plain, (segmentP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])]), c_0_72])).
% 0.06/0.47  cnf(c_0_82, negated_conjecture, (app(esk52_0,esk6_2(esk50_0,esk52_0))=esk50_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_73]), c_0_74])]), c_0_30])).
% 0.06/0.47  cnf(c_0_83, negated_conjecture, (esk50_0=nil|ssList(esk6_2(esk50_0,esk52_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_73]), c_0_74])]), c_0_30])).
% 0.06/0.47  cnf(c_0_84, negated_conjecture, (app(esk52_0,esk6_2(esk51_0,esk52_0))=esk51_0|esk51_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_75]), c_0_29])]), c_0_76])).
% 0.06/0.47  cnf(c_0_85, negated_conjecture, (esk51_0=nil|ssList(esk6_2(esk51_0,esk52_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_75]), c_0_29])]), c_0_76])).
% 0.06/0.47  fof(c_0_86, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|((~frontsegP(X3,X4)|~frontsegP(X4,X3))|X3=X4)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
% 0.06/0.47  cnf(c_0_87, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_71])])).
% 0.06/0.47  cnf(c_0_88, negated_conjecture, (app(esk52_0,esk6_2(esk52_0,esk52_0))=esk52_0|esk50_0=nil|esk52_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_79]), c_0_30])).
% 0.06/0.47  cnf(c_0_89, negated_conjecture, (esk50_0=nil|esk52_0=nil|ssList(esk6_2(esk52_0,esk52_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_79]), c_0_30])).
% 0.06/0.47  cnf(c_0_90, negated_conjecture, (esk51_0=nil|~segmentP(esk50_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_25]), c_0_65]), c_0_25])).
% 0.06/0.47  cnf(c_0_91, negated_conjecture, (esk50_0=nil|segmentP(esk50_0,esk52_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_30]), c_0_83])).
% 0.06/0.47  cnf(c_0_92, negated_conjecture, (neq(esk52_0,nil)|nil=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_93, negated_conjecture, (esk51_0=nil|segmentP(esk51_0,esk52_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_84]), c_0_76]), c_0_85])).
% 0.06/0.47  fof(c_0_94, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])).
% 0.06/0.47  cnf(c_0_95, plain, (X1=X2|~frontsegP(X2,X1)|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
% 0.06/0.47  cnf(c_0_96, negated_conjecture, (esk52_0=nil|esk50_0=nil|frontsegP(esk52_0,nil)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_30]), c_0_89])).
% 0.06/0.47  cnf(c_0_97, negated_conjecture, (frontsegP(esk50_0,esk52_0)|nil=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_98, negated_conjecture, (esk50_0=nil|esk51_0=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_92]), c_0_76]), c_0_93])).
% 0.06/0.47  cnf(c_0_99, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_94])).
% 0.06/0.47  cnf(c_0_100, negated_conjecture, (nil!=esk48_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  fof(c_0_101, plain, (![X2]:(~ssList(X2)|segmentP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).
% 0.06/0.47  fof(c_0_102, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])).
% 0.06/0.47  cnf(c_0_103, negated_conjecture, (esk50_0=nil|esk52_0=nil|~frontsegP(nil,esk52_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_71])]), c_0_30])).
% 0.06/0.47  cnf(c_0_104, negated_conjecture, (esk51_0=nil|frontsegP(nil,esk52_0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
% 0.06/0.47  cnf(c_0_105, negated_conjecture, (esk50_0=nil|esk52_0=nil|esk51_0!=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_37]), c_0_30]), c_0_38])).
% 0.06/0.47  cnf(c_0_106, negated_conjecture, (esk50_0!=nil|~segmentP(esk50_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_65]), c_0_65]), c_0_25])).
% 0.06/0.47  cnf(c_0_107, plain, (segmentP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
% 0.06/0.47  cnf(c_0_108, plain, (~ssList(X1)|~ssList(X2)|X1!=X2|~neq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
% 0.06/0.47  cnf(c_0_109, negated_conjecture, (neq(esk52_0,nil)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_110, negated_conjecture, (esk52_0=nil|esk50_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_105])).
% 0.06/0.47  cnf(c_0_111, negated_conjecture, (app(esk52_0,esk6_2(esk50_0,esk52_0))=esk50_0|esk51_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_97]), c_0_74])]), c_0_76])).
% 0.06/0.47  cnf(c_0_112, negated_conjecture, (esk51_0=nil|ssList(esk6_2(esk50_0,esk52_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_97]), c_0_74])]), c_0_76])).
% 0.06/0.47  cnf(c_0_113, negated_conjecture, (esk50_0!=nil|~segmentP(esk51_0,nil)|~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_71]), c_0_74])])).
% 0.06/0.47  cnf(c_0_114, plain, (~ssList(X1)|~neq(X1,X1)), inference(er,[status(thm)],[c_0_108])).
% 0.06/0.47  cnf(c_0_115, negated_conjecture, (esk50_0=nil|neq(nil,nil)), inference(spm,[status(thm)],[c_0_109, c_0_110])).
% 0.06/0.47  cnf(c_0_116, negated_conjecture, (neq(esk49_0,nil)|nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.47  cnf(c_0_117, negated_conjecture, (esk51_0=nil|esk52_0=nil|esk50_0!=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_111]), c_0_76]), c_0_112])).
% 0.06/0.47  cnf(c_0_118, negated_conjecture, (esk50_0!=nil|~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_107]), c_0_29])])).
% 0.06/0.47  cnf(c_0_119, negated_conjecture, (esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_71])])).
% 0.06/0.47  cnf(c_0_120, negated_conjecture, (neq(esk51_0,nil)|esk50_0!=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_25]), c_0_65])).
% 0.06/0.47  cnf(c_0_121, negated_conjecture, (esk52_0=nil|esk51_0=nil), inference(spm,[status(thm)],[c_0_117, c_0_98])).
% 0.06/0.47  cnf(c_0_122, negated_conjecture, (~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_118, c_0_119])])).
% 0.06/0.47  cnf(c_0_123, negated_conjecture, (neq(esk51_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_119])])).
% 0.06/0.47  cnf(c_0_124, negated_conjecture, (esk51_0=nil), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_121]), c_0_122])).
% 0.06/0.47  cnf(c_0_125, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_123, c_0_124]), c_0_122]), ['proof']).
% 0.06/0.47  # SZS output end CNFRefutation
% 0.06/0.47  # Training examples: 0 positive, 0 negative
% 0.06/0.47  
% 0.06/0.47  # -------------------------------------------------
% 0.06/0.47  # User time                : 0.234 s
% 0.06/0.47  # System time              : 0.010 s
% 0.06/0.47  # Total time               : 0.245 s
% 0.06/0.47  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
