<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC417-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC417-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC417-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC417-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC417-1.p">SWC417-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n158.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:52 EDT 2016

% Result     : Unsatisfiable 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of clauses        :   23 (  42 expanded)
%              Number of leaves         :   13 (  26 expanded)
%              Depth                    :    6
%              Number of atoms          :   43 (  68 expanded)
%              Number of equality atoms :   15 (  21 expanded)
%              Maximal clause size      :    6 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( neq(sk2,nil)
    | neq(sk2,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | app(app(X2,cons(X1,nil)),X3) != sk2
    | app(app(X3,cons(X1,nil)),X2) != sk1
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_14)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( app(sk6,cons(sk5,nil)) = sk3
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_18)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( ssItem(sk5)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_15)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ssList(sk6)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_16)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( app(cons(sk5,nil),sk6) = sk4
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_17)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( neq(sk2,nil) ),
    inference(cn,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( app(app(X1,cons(X2,nil)),X3) != sk1
    | app(app(X3,cons(X2,nil)),X1) != sk2
    | ~ ssItem(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_1">c_0_1</A>,c_0_2]),c_0_13])])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( app(sk6,cons(sk5,nil)) = sk1 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>,c_0_4]),c_0_2]),c_0_13])])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ssItem(sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,c_0_2]),c_0_13])])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( ssList(sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>,c_0_2]),c_0_13])])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( app(app(X1,cons(sk5,nil)),sk6) != sk2
    | app(sk1,X1) != sk1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_15]),c_0_16]),c_0_17])])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( app(app(nil,cons(sk5,nil)),sk6) != sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_7]),c_0_8]),c_0_9])])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( app(cons(sk5,nil),sk6) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_2]),c_0_2]),c_0_13])])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( ~ ssList(cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_11]),c_0_20])])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_12]),c_0_16]),c_0_8])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC417-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n158.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:43:38 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.26  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.03/0.26  # and selection function SelectCQArNTEqFirst.
% 0.03/0.26  #
% 0.03/0.26  # Preprocessing time       : 0.009 s
% 0.03/0.26  # Presaturation interreduction done
% 0.03/0.26  
% 0.03/0.26  # Proof found!
% 0.03/0.26  # SZS status Unsatisfiable
% 0.03/0.26  # SZS output start CNFRefutation
% 0.03/0.26  cnf(c_0_0, negated_conjecture, (neq(sk2,nil)|neq(sk2,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 0.03/0.26  cnf(c_0_1, negated_conjecture, (~ssItem(X1)|~ssList(X2)|~ssList(X3)|app(app(X2,cons(X1,nil)),X3)!=sk2|app(app(X3,cons(X1,nil)),X2)!=sk1|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_14)).
% 0.03/0.26  cnf(c_0_2, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 0.03/0.26  cnf(c_0_3, negated_conjecture, (app(sk6,cons(sk5,nil))=sk3|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_18)).
% 0.03/0.26  cnf(c_0_4, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 0.03/0.26  cnf(c_0_5, negated_conjecture, (ssItem(sk5)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_15)).
% 0.03/0.26  cnf(c_0_6, negated_conjecture, (ssList(sk6)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_16)).
% 0.03/0.26  cnf(c_0_7, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause73)).
% 0.03/0.26  cnf(c_0_8, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.03/0.26  cnf(c_0_9, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 0.03/0.26  cnf(c_0_10, negated_conjecture, (app(cons(sk5,nil),sk6)=sk4|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_17)).
% 0.03/0.26  cnf(c_0_11, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause74)).
% 0.03/0.26  cnf(c_0_12, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.03/0.26  cnf(c_0_13, negated_conjecture, (neq(sk2,nil)), inference(cn,[status(thm)],[c_0_0])).
% 0.03/0.26  cnf(c_0_14, negated_conjecture, (app(app(X1,cons(X2,nil)),X3)!=sk1|app(app(X3,cons(X2,nil)),X1)!=sk2|~ssItem(X2)|~ssList(X1)|~ssList(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_1, c_0_2]), c_0_13])])).
% 0.03/0.26  cnf(c_0_15, negated_conjecture, (app(sk6,cons(sk5,nil))=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_4]), c_0_2]), c_0_13])])).
% 0.03/0.26  cnf(c_0_16, negated_conjecture, (ssItem(sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_2]), c_0_13])])).
% 0.03/0.26  cnf(c_0_17, negated_conjecture, (ssList(sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_2]), c_0_13])])).
% 0.03/0.26  cnf(c_0_18, negated_conjecture, (app(app(X1,cons(sk5,nil)),sk6)!=sk2|app(sk1,X1)!=sk1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17])])).
% 0.03/0.26  cnf(c_0_19, negated_conjecture, (app(app(nil,cons(sk5,nil)),sk6)!=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_7]), c_0_8]), c_0_9])])).
% 0.03/0.26  cnf(c_0_20, negated_conjecture, (app(cons(sk5,nil),sk6)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_2]), c_0_2]), c_0_13])])).
% 0.03/0.26  cnf(c_0_21, negated_conjecture, (~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_11]), c_0_20])])).
% 0.03/0.26  cnf(c_0_22, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_12]), c_0_16]), c_0_8])]), ['proof']).
% 0.03/0.26  # SZS output end CNFRefutation
% 0.03/0.26  # Training examples: 0 positive, 0 negative
% 0.03/0.26  
% 0.03/0.26  # -------------------------------------------------
% 0.03/0.26  # User time                : 0.026 s
% 0.03/0.26  # System time              : 0.003 s
% 0.03/0.26  # Total time               : 0.029 s
% 0.03/0.26  # Maximum resident set size: 1864 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
