<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC023-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC023-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC023-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC023-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC023-1.p">SWC023-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n110.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:20 EDT 2016

% Result     : Unsatisfiable 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of clauses        :   27 (  50 expanded)
%              Number of leaves         :   16 (  33 expanded)
%              Depth                    :    6
%              Number of atoms          :   46 (  85 expanded)
%              Number of equality atoms :   15 (  24 expanded)
%              Maximal clause size      :    4 (   2 average)
%              Maximal term depth       :    1 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( ~ ssList(X1)
    | ~ neq(X1,nil)
    | ~ frontsegP(sk2,X1)
    | ~ frontsegP(sk1,X1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( frontsegP(sk3,sk5)
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_16)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssList(sk5)
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( frontsegP(sk4,sk5)
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_15)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( neq(sk5,nil)
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_14)).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( frontsegP(sk3,sk5)
    | nil = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause60)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( nil = X1
    | ~ frontsegP(nil,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause84)).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( ssList(sk5)
    | nil = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( neq(sk2,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( neq(sk5,nil)
    | nil = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ~ neq(X1,nil)
    | ~ frontsegP(sk3,X1)
    | ~ frontsegP(sk4,X1)
    | ~ ssList(X1) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,c_0_1]),<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( sk3 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_3]),c_0_4]),c_0_5]),<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( sk4 = nil
    | frontsegP(nil,sk5) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( ~ neq(nil,nil)
    | ~ frontsegP(sk4,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_9]),c_0_10]),c_0_18])])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( sk4 = nil
    | sk5 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_19]),<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_9]),c_0_21])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( neq(sk4,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( sk4 = nil ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_22]),<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_25]),<A HREF="#c_0_23">c_0_23</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC023-1 : TPTP v6.4.0. Released v2.4.0.
% 0.03/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n110.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:34:23 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.03/0.25  # and selection function SelectCQArNTEqFirst.
% 0.03/0.25  #
% 0.03/0.25  # Preprocessing time       : 0.009 s
% 0.03/0.25  # Presaturation interreduction done
% 0.03/0.25  
% 0.03/0.25  # Proof found!
% 0.03/0.25  # SZS status Unsatisfiable
% 0.03/0.25  # SZS output start CNFRefutation
% 0.03/0.25  cnf(c_0_0, negated_conjecture, (~ssList(X1)|~neq(X1,nil)|~frontsegP(sk2,X1)|~frontsegP(sk1,X1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
% 0.03/0.25  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.03/0.25  cnf(c_0_2, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 0.03/0.25  cnf(c_0_3, negated_conjecture, (frontsegP(sk3,sk5)|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_16)).
% 0.03/0.25  cnf(c_0_4, negated_conjecture, (ssList(sk5)|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_10)).
% 0.03/0.25  cnf(c_0_5, negated_conjecture, (frontsegP(sk4,sk5)|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_15)).
% 0.03/0.25  cnf(c_0_6, negated_conjecture, (neq(sk5,nil)|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_14)).
% 0.03/0.25  cnf(c_0_7, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.03/0.25  cnf(c_0_8, negated_conjecture, (frontsegP(sk3,sk5)|nil=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
% 0.03/0.25  cnf(c_0_9, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause60)).
% 0.03/0.25  cnf(c_0_10, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.03/0.25  cnf(c_0_11, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 0.03/0.25  cnf(c_0_12, axiom, (nil=X1|~frontsegP(nil,X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause84)).
% 0.03/0.25  cnf(c_0_13, negated_conjecture, (ssList(sk5)|nil=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
% 0.03/0.25  cnf(c_0_14, negated_conjecture, (neq(sk2,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 0.03/0.25  cnf(c_0_15, negated_conjecture, (neq(sk5,nil)|nil=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_11)).
% 0.03/0.25  cnf(c_0_16, negated_conjecture, (~neq(X1,nil)|~frontsegP(sk3,X1)|~frontsegP(sk4,X1)|~ssList(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
% 0.03/0.25  cnf(c_0_17, negated_conjecture, (sk3=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_3]), c_0_4]), c_0_5]), c_0_6])).
% 0.03/0.25  cnf(c_0_18, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_7, c_0_1])).
% 0.03/0.25  cnf(c_0_19, negated_conjecture, (sk4=nil|frontsegP(nil,sk5)), inference(rw,[status(thm)],[c_0_8, c_0_17])).
% 0.03/0.25  cnf(c_0_20, negated_conjecture, (~neq(nil,nil)|~frontsegP(sk4,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_9]), c_0_10]), c_0_18])])).
% 0.03/0.25  cnf(c_0_21, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_11, c_0_2])).
% 0.03/0.25  cnf(c_0_22, negated_conjecture, (sk4=nil|sk5=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_19]), c_0_13])).
% 0.03/0.25  cnf(c_0_23, negated_conjecture, (~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_9]), c_0_21])])).
% 0.03/0.25  cnf(c_0_24, negated_conjecture, (neq(sk4,nil)), inference(rw,[status(thm)],[c_0_14, c_0_2])).
% 0.03/0.25  cnf(c_0_25, negated_conjecture, (sk4=nil), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_22]), c_0_23])).
% 0.03/0.25  cnf(c_0_26, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_23]), ['proof']).
% 0.03/0.25  # SZS output end CNFRefutation
% 0.03/0.25  # Training examples: 0 positive, 0 negative
% 0.03/0.25  
% 0.03/0.25  # -------------------------------------------------
% 0.03/0.25  # User time                : 0.020 s
% 0.03/0.25  # System time              : 0.003 s
% 0.03/0.25  # Total time               : 0.023 s
% 0.03/0.25  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
