<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC392-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC392-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC392-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC392-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC392-1.p">SWC392-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n178.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:45 EDT 2016

% Result     : Unsatisfiable 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of clauses        :   39 (  87 expanded)
%              Number of leaves         :   21 (  55 expanded)
%              Depth                    :   13
%              Number of atoms          :   90 ( 172 expanded)
%              Number of equality atoms :   24 (  51 expanded)
%              Maximal clause size      :    7 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_3"></A>cnf(c_0_3,axiom,
    ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
    | ~ memberP(X1,X2)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause169)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( memberP(sk4,sk6)
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_16)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( ssItem(sk6)
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( ssList(skaf43(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause52)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( cons(sk6,nil) = sk3
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_15)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( memberP(sk1,sk5) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( memberP(app(X1,X3),X2)
    | ~ memberP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause140)).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ssItem(sk5) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( ~ memberP(sk2,sk5) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( memberP(X2,X3)
    | X3 = X1
    | ~ memberP(cons(X1,X2),X3)
    | ~ ssList(X2)
    | ~ ssItem(X1)
    | ~ ssItem(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause161)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause71)).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( app(skaf42(sk4,sk6),cons(sk6,skaf43(sk6,sk4))) = sk4
    | sk3 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>,c_0_4]),c_0_21])]),<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( sk3 = nil
    | ~ ssList(app(sk4,cons(sk6,X1)))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(sk4,cons(sk6,X1))) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_23]),c_0_6]),c_0_7])]),<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( sk3 = nil
    | ~ ssList(app(sk4,sk3))
    | ~ duplicatefreeP(app(sk4,sk3)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_8]),c_0_9])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( sk3 = nil
    | ssItem(X1)
    | ~ ssList(app(sk4,sk3)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( memberP(sk3,sk5) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( sk3 = nil
    | ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_14]),c_0_21]),c_0_27])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( memberP(app(sk3,X1),sk5)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_28]),c_0_16]),c_0_27])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( app(sk3,X1) = cons(sk6,X1)
    | sk3 = nil
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_8]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( sk3 = nil
    | memberP(cons(sk6,X1),sk5)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( ~ memberP(sk4,sk5) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( sk3 = nil
    | sk5 = sk6
    | memberP(X1,sk5)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_32]),c_0_16])]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( sk5 = sk6
    | sk3 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_34]),c_0_21])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( sk3 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_35]),<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( memberP(nil,sk5) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_37]),c_0_16])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC392-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.22  % Computer   : n178.star.cs.uiowa.edu
% 0.03/0.22  % Model      : x86_64 x86_64
% 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.22  % Memory     : 32218.75MB
% 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.22  % CPULimit   : 300
% 0.03/0.22  % DateTime   : Tue May 10 20:43:22 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.28  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.03/0.28  # and selection function SelectCQArNTEqFirst.
% 0.03/0.28  #
% 0.03/0.28  # Preprocessing time       : 0.010 s
% 0.03/0.28  # Presaturation interreduction done
% 0.03/0.28  
% 0.03/0.28  # Proof found!
% 0.03/0.28  # SZS status Unsatisfiable
% 0.03/0.28  # SZS output start CNFRefutation
% 0.03/0.28  cnf(c_0_0, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 0.03/0.28  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 0.03/0.28  cnf(c_0_2, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
% 0.03/0.28  cnf(c_0_3, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause169)).
% 0.03/0.28  cnf(c_0_4, negated_conjecture, (memberP(sk4,sk6)|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_16)).
% 0.03/0.28  cnf(c_0_5, negated_conjecture, (ssItem(sk6)|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_11)).
% 0.03/0.28  cnf(c_0_6, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause52)).
% 0.03/0.28  cnf(c_0_7, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause53)).
% 0.03/0.28  cnf(c_0_8, negated_conjecture, (cons(sk6,nil)=sk3|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_15)).
% 0.03/0.28  cnf(c_0_9, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.03/0.28  cnf(c_0_10, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
% 0.03/0.28  cnf(c_0_11, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.03/0.28  cnf(c_0_12, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.03/0.28  cnf(c_0_13, negated_conjecture, (memberP(sk1,sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
% 0.03/0.28  cnf(c_0_14, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.03/0.28  cnf(c_0_15, axiom, (memberP(app(X1,X3),X2)|~memberP(X1,X2)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause140)).
% 0.03/0.28  cnf(c_0_16, negated_conjecture, (ssItem(sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 0.03/0.28  cnf(c_0_17, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
% 0.03/0.28  cnf(c_0_18, negated_conjecture, (~memberP(sk2,sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
% 0.03/0.28  cnf(c_0_19, axiom, (memberP(X2,X3)|X3=X1|~memberP(cons(X1,X2),X3)|~ssList(X2)|~ssItem(X1)|~ssItem(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause161)).
% 0.03/0.28  cnf(c_0_20, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause71)).
% 0.03/0.28  cnf(c_0_21, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 0.03/0.28  cnf(c_0_22, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_2])).
% 0.03/0.28  cnf(c_0_23, negated_conjecture, (app(skaf42(sk4,sk6),cons(sk6,skaf43(sk6,sk4)))=sk4|sk3=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_3, c_0_4]), c_0_21])]), c_0_5])).
% 0.03/0.28  cnf(c_0_24, negated_conjecture, (sk3=nil|~ssList(app(sk4,cons(sk6,X1)))|~ssList(X1)|~duplicatefreeP(app(sk4,cons(sk6,X1)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_6]), c_0_7])]), c_0_5])).
% 0.03/0.28  cnf(c_0_25, negated_conjecture, (sk3=nil|~ssList(app(sk4,sk3))|~duplicatefreeP(app(sk4,sk3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_8]), c_0_9])])).
% 0.03/0.28  cnf(c_0_26, negated_conjecture, (sk3=nil|ssItem(X1)|~ssList(app(sk4,sk3))), inference(spm,[status(thm)],[c_0_25, c_0_10])).
% 0.03/0.28  cnf(c_0_27, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_11, c_0_12])).
% 0.03/0.28  cnf(c_0_28, negated_conjecture, (memberP(sk3,sk5)), inference(rw,[status(thm)],[c_0_13, c_0_12])).
% 0.03/0.28  cnf(c_0_29, negated_conjecture, (sk3=nil|ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_14]), c_0_21]), c_0_27])])).
% 0.03/0.28  cnf(c_0_30, negated_conjecture, (memberP(app(sk3,X1),sk5)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_28]), c_0_16]), c_0_27])])).
% 0.03/0.28  cnf(c_0_31, negated_conjecture, (app(sk3,X1)=cons(sk6,X1)|sk3=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_8]), c_0_29])).
% 0.03/0.28  cnf(c_0_32, negated_conjecture, (sk3=nil|memberP(cons(sk6,X1),sk5)|~ssList(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
% 0.03/0.28  cnf(c_0_33, negated_conjecture, (~memberP(sk4,sk5)), inference(rw,[status(thm)],[c_0_18, c_0_1])).
% 0.03/0.28  cnf(c_0_34, negated_conjecture, (sk3=nil|sk5=sk6|memberP(X1,sk5)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_32]), c_0_16])]), c_0_29])).
% 0.03/0.28  cnf(c_0_35, negated_conjecture, (sk5=sk6|sk3=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_21])])).
% 0.03/0.28  cnf(c_0_36, negated_conjecture, (sk3=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_35]), c_0_4])).
% 0.03/0.28  cnf(c_0_37, negated_conjecture, (memberP(nil,sk5)), inference(rw,[status(thm)],[c_0_28, c_0_36])).
% 0.03/0.28  cnf(c_0_38, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_37]), c_0_16])]), ['proof']).
% 0.03/0.28  # SZS output end CNFRefutation
% 0.03/0.28  # Training examples: 0 positive, 0 negative
% 0.03/0.28  
% 0.03/0.28  # -------------------------------------------------
% 0.03/0.28  # User time                : 0.054 s
% 0.03/0.28  # System time              : 0.003 s
% 0.03/0.28  # Total time               : 0.057 s
% 0.03/0.28  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
