<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC058+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC058+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC058+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC058%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC058+1.p">SWC058+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n104.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:28 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :   54 ( 607 expanded)
%              Number of clauses        :   38 ( 223 expanded)
%              Number of leaves         :    8 ( 140 expanded)
%              Depth                    :   11
%              Number of atoms          :  228 (7496 expanded)
%              Number of equality atoms :   70 (2606 expanded)
%              Maximal formula depth    :   19 (   5 average)
%              Maximal clause size      :   28 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ~ ssList(X4)
                  | X2 != X4
                  | X1 != X3
                  | ( ( nil != X2
                      | nil = X1 )
                    & ( ~ neq(X2,nil)
                      | ? [X5] :
                          ( ssList(X5)
                          & neq(X5,nil)
                          & segmentP(X2,X5)
                          & segmentP(X1,X5) ) ) )
                  | ( ( nil != X4
                      | nil != X3 )
                    & ( ~ neq(X3,nil)
                      | ~ rearsegP(X4,X3) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( segmentP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax7)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( rearsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax6)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,X1) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax55)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,nil) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax57)).

<A NAME="c_0_8"></A>fof(c_0_8,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ~ ssList(X4)
                    | X2 != X4
                    | X1 != X3
                    | ( ( nil != X2
                        | nil = X1 )
                      & ( ~ neq(X2,nil)
                        | ? [X5] :
                            ( ssList(X5)
                            & neq(X5,nil)
                            & segmentP(X2,X5)
                            & segmentP(X1,X5) ) ) )
                    | ( ( nil != X4
                        | nil != X3 )
                      & ( ~ neq(X3,nil)
                        | ~ rearsegP(X4,X3) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_9"></A>fof(c_0_9,negated_conjecture,(
    ! [X10] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ( neq(esk49_0,nil)
        | nil = esk49_0 )
      & ( ~ ssList(X10)
        | ~ neq(X10,nil)
        | ~ segmentP(esk49_0,X10)
        | ~ segmentP(esk48_0,X10)
        | nil = esk49_0 )
      & ( neq(esk49_0,nil)
        | nil != esk48_0 )
      & ( ~ ssList(X10)
        | ~ neq(X10,nil)
        | ~ segmentP(esk49_0,X10)
        | ~ segmentP(esk48_0,X10)
        | nil != esk48_0 )
      & ( neq(esk50_0,nil)
        | nil = esk51_0 )
      & ( rearsegP(esk51_0,esk50_0)
        | nil = esk51_0 )
      & ( neq(esk50_0,nil)
        | nil = esk50_0 )
      & ( rearsegP(esk51_0,esk50_0)
        | nil = esk50_0 ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).

<A NAME="c_0_10"></A>fof(c_0_10,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk8_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ssList(esk9_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(app(X9,X6),X10) != X5
        | segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

<A NAME="c_0_11"></A>fof(c_0_11,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk7_2(X4,X5))
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(esk7_2(X4,X5),X5) = X4
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X7,X5) != X4
        | rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( nil = esk50_0
    | rearsegP(esk51_0,esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,plain,
    ( segmentP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(app(X3,X2),X4) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,plain,
    ( app(esk7_2(X1,X2),X2) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( esk50_0 = nil
    | rearsegP(esk49_0,esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,plain,
    ( ssList(esk7_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( nil = esk49_0
    | ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_23"></A>fof(c_0_23,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( nil = esk51_0
    | neq(esk50_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( app(esk7_2(esk49_0,esk50_0),esk50_0) = esk49_0
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_18]),c_0_19]),c_0_20])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( esk50_0 = nil
    | ssList(esk7_2(esk49_0,esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_18]),c_0_19]),c_0_20])])).

<A NAME="c_0_28"></A>fof(c_0_28,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).

<A NAME="c_0_29"></A>fof(c_0_29,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( esk49_0 = nil
    | ~ segmentP(esk50_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ ssList(X1)
    | ~ neq(X1,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( segmentP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( esk49_0 = nil
    | neq(esk50_0,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( esk50_0 = nil
    | segmentP(app(esk49_0,X1),esk50_0)
    | ~ ssList(app(esk49_0,X1))
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26]),c_0_19])]),<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( nil = X2
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( nil != esk48_0
    | ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_38"></A>fof(c_0_38,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( esk49_0 = nil
    | ~ segmentP(esk49_0,esk50_0) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31]),c_0_19])]),<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( esk50_0 = nil
    | segmentP(esk49_0,esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_34]),c_0_20]),c_0_35])])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( esk50_0 = nil
    | esk49_0 != nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_26]),c_0_19])]),<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( esk50_0 != nil
    | ~ segmentP(esk50_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ ssList(X1)
    | ~ neq(X1,nil) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_15]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,plain,
    ( segmentP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( neq(esk49_0,nil)
    | nil != esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_40]),<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( esk50_0 != nil
    | ~ segmentP(esk49_0,nil)
    | ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_43]),c_0_35]),c_0_19])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( neq(esk49_0,nil)
    | esk50_0 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( esk49_0 = nil
    | ~ segmentP(esk49_0,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,<A HREF="#c_0_45">c_0_45</A>])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( esk50_0 != nil
    | ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_43]),c_0_20])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( neq(esk49_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_45])])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( esk49_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_43]),c_0_20])])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_45])])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_51]),<A HREF="#c_0_52">c_0_52</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.02/0.03  % Problem    : SWC058+1 : TPTP v6.4.0. Released v2.4.0.
% 0.02/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n104.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:34:39 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.42  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.42  # and selection function SelectCQPrecWNTNp.
% 0.06/0.42  #
% 0.06/0.42  # Preprocessing time       : 0.014 s
% 0.06/0.42  # Presaturation interreduction done
% 0.06/0.42  
% 0.06/0.42  # Proof found!
% 0.06/0.42  # SZS status Theorem
% 0.06/0.42  # SZS output start CNFRefutation
% 0.06/0.42  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~(ssList(X4))|X2!=X4)|X1!=X3)|((nil!=X2|nil=X1)&(~(neq(X2,nil))|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))))|((nil!=X4|nil!=X3)&(~(neq(X3,nil))|~(rearsegP(X4,X3))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.06/0.42  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax7)).
% 0.06/0.42  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax6)).
% 0.06/0.42  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>segmentP(X1,X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax55)).
% 0.06/0.42  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax84)).
% 0.06/0.42  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.06/0.42  fof(c_0_6, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.42  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>segmentP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax57)).
% 0.06/0.42  fof(c_0_8, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~ssList(X4)|X2!=X4)|X1!=X3)|((nil!=X2|nil=X1)&(~neq(X2,nil)|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))))|((nil!=X4|nil!=X3)&(~neq(X3,nil)|~rearsegP(X4,X3))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.42  fof(c_0_9, negated_conjecture, (![X10]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(((neq(esk49_0,nil)|nil=esk49_0)&((((~ssList(X10)|~neq(X10,nil))|~segmentP(esk49_0,X10))|~segmentP(esk48_0,X10))|nil=esk49_0))&((neq(esk49_0,nil)|nil!=esk48_0)&((((~ssList(X10)|~neq(X10,nil))|~segmentP(esk49_0,X10))|~segmentP(esk48_0,X10))|nil!=esk48_0))))&(((neq(esk50_0,nil)|nil=esk51_0)&(rearsegP(esk51_0,esk50_0)|nil=esk51_0))&((neq(esk50_0,nil)|nil=esk50_0)&(rearsegP(esk51_0,esk50_0)|nil=esk50_0)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
% 0.06/0.42  fof(c_0_10, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
% 0.06/0.42  fof(c_0_11, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
% 0.06/0.42  cnf(c_0_12, negated_conjecture, (nil=esk50_0|rearsegP(esk51_0,esk50_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.06/0.42  cnf(c_0_13, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.06/0.42  cnf(c_0_14, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.06/0.42  cnf(c_0_15, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.06/0.42  cnf(c_0_16, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.06/0.42  cnf(c_0_17, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.06/0.42  cnf(c_0_18, negated_conjecture, (esk50_0=nil|rearsegP(esk49_0,esk50_0)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
% 0.06/0.42  cnf(c_0_19, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_14, c_0_15])).
% 0.06/0.42  cnf(c_0_20, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.06/0.42  cnf(c_0_21, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.06/0.42  cnf(c_0_22, negated_conjecture, (nil=esk49_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.06/0.42  fof(c_0_23, plain, (![X2]:(~ssList(X2)|segmentP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
% 0.06/0.42  cnf(c_0_24, negated_conjecture, (nil=esk51_0|neq(esk50_0,nil)), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.06/0.42  cnf(c_0_25, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_16])).
% 0.06/0.42  cnf(c_0_26, negated_conjecture, (app(esk7_2(esk49_0,esk50_0),esk50_0)=esk49_0|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])])).
% 0.06/0.42  cnf(c_0_27, negated_conjecture, (esk50_0=nil|ssList(esk7_2(esk49_0,esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_18]), c_0_19]), c_0_20])])).
% 0.06/0.42  fof(c_0_28, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
% 0.06/0.42  fof(c_0_29, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
% 0.06/0.42  cnf(c_0_30, negated_conjecture, (esk49_0=nil|~segmentP(esk50_0,X1)|~segmentP(esk49_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[c_0_22, c_0_15])).
% 0.06/0.42  cnf(c_0_31, plain, (segmentP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 0.06/0.42  cnf(c_0_32, negated_conjecture, (esk49_0=nil|neq(esk50_0,nil)), inference(rw,[status(thm)],[c_0_24, c_0_13])).
% 0.06/0.42  cnf(c_0_33, negated_conjecture, (esk50_0=nil|segmentP(app(esk49_0,X1),esk50_0)|~ssList(app(esk49_0,X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_19])]), c_0_27])).
% 0.06/0.42  cnf(c_0_34, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
% 0.06/0.42  cnf(c_0_35, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
% 0.06/0.42  cnf(c_0_36, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
% 0.06/0.42  cnf(c_0_37, negated_conjecture, (nil!=esk48_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.06/0.42  fof(c_0_38, plain, (![X2]:(~ssList(X2)|segmentP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
% 0.06/0.42  cnf(c_0_39, negated_conjecture, (esk49_0=nil|~segmentP(esk49_0,esk50_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_19])]), c_0_32])).
% 0.06/0.42  cnf(c_0_40, negated_conjecture, (esk50_0=nil|segmentP(esk49_0,esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_20]), c_0_35])])).
% 0.06/0.42  cnf(c_0_41, negated_conjecture, (esk50_0=nil|esk49_0!=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_26]), c_0_19])]), c_0_27])).
% 0.06/0.42  cnf(c_0_42, negated_conjecture, (esk50_0!=nil|~segmentP(esk50_0,X1)|~segmentP(esk49_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_15]), c_0_15])).
% 0.06/0.42  cnf(c_0_43, plain, (segmentP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
% 0.06/0.42  cnf(c_0_44, negated_conjecture, (neq(esk49_0,nil)|nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.06/0.42  cnf(c_0_45, negated_conjecture, (esk50_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
% 0.06/0.42  cnf(c_0_46, negated_conjecture, (esk50_0!=nil|~segmentP(esk49_0,nil)|~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_35]), c_0_19])])).
% 0.06/0.42  cnf(c_0_47, negated_conjecture, (neq(esk49_0,nil)|esk50_0!=nil), inference(rw,[status(thm)],[c_0_44, c_0_15])).
% 0.06/0.42  cnf(c_0_48, negated_conjecture, (esk49_0=nil|~segmentP(esk49_0,nil)), inference(rw,[status(thm)],[c_0_39, c_0_45])).
% 0.06/0.42  cnf(c_0_49, negated_conjecture, (esk50_0!=nil|~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_43]), c_0_20])])).
% 0.06/0.42  cnf(c_0_50, negated_conjecture, (neq(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_45])])).
% 0.06/0.42  cnf(c_0_51, negated_conjecture, (esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_43]), c_0_20])])).
% 0.06/0.42  cnf(c_0_52, negated_conjecture, (~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_45])])).
% 0.06/0.42  cnf(c_0_53, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_51]), c_0_52]), ['proof']).
% 0.06/0.42  # SZS output end CNFRefutation
% 0.06/0.42  # Training examples: 0 positive, 0 negative
% 0.06/0.42  
% 0.06/0.42  # -------------------------------------------------
% 0.06/0.42  # User time                : 0.183 s
% 0.06/0.42  # System time              : 0.007 s
% 0.06/0.42  # Total time               : 0.190 s
% 0.06/0.42  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
