<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC176+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC176+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC176+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC176%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC176+1.p">SWC176+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n111.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:56 EDT 2016

% Result     : Theorem 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of formulae       :   48 ( 115 expanded)
%              Number of clauses        :   29 (  41 expanded)
%              Number of leaves         :    9 (  29 expanded)
%              Depth                    :   12
%              Number of atoms          :  241 (1895 expanded)
%              Number of equality atoms :   54 ( 565 expanded)
%              Maximal formula depth    :   30 (   6 average)
%              Maximal clause size      :   32 (   2 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ~ ssList(X4)
                  | X2 != X4
                  | X1 != X3
                  | ! [X5] :
                      ( ssList(X5)
                     => ! [X6] :
                          ( ~ ssList(X6)
                          | app(app(X5,X3),X6) != X4
                          | ~ strictorderedP(X3)
                          | ? [X7] :
                              ( ssItem(X7)
                              & ? [X8] :
                                  ( ssList(X8)
                                  & app(X8,cons(X7,nil)) = X5
                                  & ? [X9] :
                                      ( ssItem(X9)
                                      & ? [X10] :
                                          ( ssList(X10)
                                          & app(cons(X9,nil),X10) = X3
                                          & lt(X7,X9) ) ) ) )
                          | ? [X11] :
                              ( ssItem(X11)
                              & ? [X12] :
                                  ( ssList(X12)
                                  & app(cons(X11,nil),X12) = X6
                                  & ? [X13] :
                                      ( ssItem(X13)
                                      & ? [X14] :
                                          ( ssList(X14)
                                          & app(X14,cons(X13,nil)) = X3
                                          & lt(X13,X11) ) ) ) ) ) )
                  | ! [X15] :
                      ( ssItem(X15)
                     => ! [X16] :
                          ( ~ ssItem(X16)
                          | ! [X17] :
                              ( ssList(X17)
                             => ! [X18] :
                                  ( ssList(X18)
                                 => app(app(app(X17,cons(X15,nil)),cons(X16,nil)),X18) != X1 ) )
                          | neq(X15,X16) ) )
                  | ( nil != X4
                    & nil = X3 ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax1)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( strictorderedP(X1)
      &lt;=> ! [X2] :
            ( ssItem(X2)
           => ! [X3] :
                ( ssItem(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ! [X5] :
                        ( ssList(X5)
                       => ! [X6] :
                            ( ssList(X6)
                           => ( app(app(X4,cons(X2,X5)),cons(X3,X6)) = X1
                             => lt(X2,X3) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax12)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ~ lt(X1,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax90)).

<A NAME="c_0_9"></A>fof(c_0_9,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ~ ssList(X4)
                    | X2 != X4
                    | X1 != X3
                    | ! [X5] :
                        ( ssList(X5)
                       => ! [X6] :
                            ( ~ ssList(X6)
                            | app(app(X5,X3),X6) != X4
                            | ~ strictorderedP(X3)
                            | ? [X7] :
                                ( ssItem(X7)
                                & ? [X8] :
                                    ( ssList(X8)
                                    & app(X8,cons(X7,nil)) = X5
                                    & ? [X9] :
                                        ( ssItem(X9)
                                        & ? [X10] :
                                            ( ssList(X10)
                                            & app(cons(X9,nil),X10) = X3
                                            & lt(X7,X9) ) ) ) )
                            | ? [X11] :
                                ( ssItem(X11)
                                & ? [X12] :
                                    ( ssList(X12)
                                    & app(cons(X11,nil),X12) = X6
                                    & ? [X13] :
                                        ( ssItem(X13)
                                        & ? [X14] :
                                            ( ssList(X14)
                                            & app(X14,cons(X13,nil)) = X3
                                            & lt(X13,X11) ) ) ) ) ) )
                    | ! [X15] :
                        ( ssItem(X15)
                       => ! [X16] :
                            ( ~ ssItem(X16)
                            | ! [X17] :
                                ( ssList(X17)
                               => ! [X18] :
                                    ( ssList(X18)
                                   => app(app(app(X17,cons(X15,nil)),cons(X16,nil)),X18) != X1 ) )
                            | neq(X15,X16) ) )
                    | ( nil != X4
                      & nil = X3 ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_10"></A>fof(c_0_10,negated_conjecture,(
    ! [X25,X26,X27,X28,X29,X30,X31,X32] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ssList(esk52_0)
      & ssList(esk53_0)
      & app(app(esk52_0,esk50_0),esk53_0) = esk51_0
      & strictorderedP(esk50_0)
      & ( ~ ssItem(X25)
        | ~ ssList(X26)
        | app(X26,cons(X25,nil)) != esk52_0
        | ~ ssItem(X27)
        | ~ ssList(X28)
        | app(cons(X27,nil),X28) != esk50_0
        | ~ lt(X25,X27) )
      & ( ~ ssItem(X29)
        | ~ ssList(X30)
        | app(cons(X29,nil),X30) != esk53_0
        | ~ ssItem(X31)
        | ~ ssList(X32)
        | app(X32,cons(X31,nil)) != esk50_0
        | ~ lt(X31,X29) )
      & ssItem(esk54_0)
      & ssItem(esk55_0)
      & ssList(esk56_0)
      & ssList(esk57_0)
      & app(app(app(esk56_0,cons(esk54_0,nil)),cons(esk55_0,nil)),esk57_0) = esk48_0
      & ~ neq(esk54_0,esk55_0)
      & ( nil = esk51_0
        | nil != esk50_0 ) ) ),
    inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).

<A NAME="c_0_11"></A>fof(c_0_11,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssItem(X4)
        | ~ ssItem(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssItem(X4)
        | ~ ssItem(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( app(app(app(esk56_0,cons(esk54_0,nil)),cons(esk55_0,nil)),esk57_0) = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( ~ neq(esk54_0,esk55_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssItem(X1)
    | ~ ssItem(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ssItem(esk55_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( ssItem(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( app(app(app(esk56_0,cons(esk54_0,nil)),cons(esk55_0,nil)),esk57_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( esk55_0 = esk54_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_15]),c_0_16]),c_0_17])])).

<A NAME="c_0_20"></A>fof(c_0_20,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( app(app(app(esk56_0,cons(esk54_0,nil)),cons(esk54_0,nil)),esk57_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( ssList(esk57_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_24"></A>fof(c_0_24,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( app(app(esk56_0,cons(esk54_0,nil)),app(cons(esk54_0,nil),esk57_0)) = esk50_0
    | ~ ssList(app(esk56_0,cons(esk54_0,nil)))
    | ~ ssList(cons(esk54_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_22]),c_0_23])])).

<A NAME="c_0_26"></A>cnf(c_0_26,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_28"></A>fof(c_0_28,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( app(app(esk56_0,cons(esk54_0,nil)),app(cons(esk54_0,nil),esk57_0)) = esk50_0
    | ~ ssList(app(esk56_0,cons(esk54_0,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26]),c_0_27]),c_0_17])])).

<A NAME="c_0_30"></A>cnf(c_0_30,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( ssList(esk56_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_32"></A>fof(c_0_32,plain,(
    ! [X7,X8,X9,X10,X11,X12] :
      ( ( ~ strictorderedP(X7)
        | ~ ssItem(X8)
        | ~ ssItem(X9)
        | ~ ssList(X10)
        | ~ ssList(X11)
        | ~ ssList(X12)
        | app(app(X10,cons(X8,X11)),cons(X9,X12)) != X7
        | lt(X8,X9)
        | ~ ssList(X7) )
      & ( ssItem(esk30_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( ssItem(esk31_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk32_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk33_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk34_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( app(app(esk32_1(X7),cons(esk30_1(X7),esk33_1(X7))),cons(esk31_1(X7),esk34_1(X7))) = X7
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( ~ lt(esk30_1(X7),esk31_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( app(app(esk56_0,cons(esk54_0,nil)),app(cons(esk54_0,nil),esk57_0)) = esk50_0
    | ~ ssList(cons(esk54_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_30]),c_0_31])])).

<A NAME="c_0_34"></A>fof(c_0_34,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).

<A NAME="c_0_35"></A>fof(c_0_35,plain,(
    ! [X1] :
      ( ssItem(X1)
     => ~ lt(X1,X1) ) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( lt(X2,X3)
    | ~ ssList(X1)
    | app(app(X4,cons(X2,X5)),cons(X3,X6)) != X1
    | ~ ssList(X6)
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssItem(X3)
    | ~ ssItem(X2)
    | ~ strictorderedP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( app(app(esk56_0,cons(esk54_0,nil)),app(cons(esk54_0,nil),esk57_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_26]),c_0_27]),c_0_17])])).

<A NAME="c_0_38"></A>cnf(c_0_38,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_40"></A>fof(c_0_40,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | ~ lt(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( lt(X1,X2)
    | ~ strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( app(app(esk56_0,cons(esk54_0,nil)),cons(esk54_0,esk57_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_38]),c_0_23]),c_0_17])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( strictorderedP(esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( ~ lt(X1,X1)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( lt(esk54_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_42]),c_0_43]),c_0_44]),c_0_23]),c_0_27]),c_0_31]),c_0_17])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_46]),c_0_17])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC176+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n111.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:36:53 CDT 2016
% 0.02/0.22  % CPUTime    : 
% 0.02/0.22  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.29  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.02/0.29  # and selection function SelectCQPrecWNTNp.
% 0.02/0.29  #
% 0.02/0.29  # Preprocessing time       : 0.014 s
% 0.02/0.29  # Presaturation interreduction done
% 0.02/0.29  
% 0.02/0.29  # Proof found!
% 0.02/0.29  # SZS status Theorem
% 0.02/0.29  # SZS output start CNFRefutation
% 0.02/0.29  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~(ssList(X4))|X2!=X4)|X1!=X3)|![X5]:(ssList(X5)=>![X6]:((((~(ssList(X6))|app(app(X5,X3),X6)!=X4)|~(strictorderedP(X3)))|?[X7]:(ssItem(X7)&?[X8]:((ssList(X8)&app(X8,cons(X7,nil))=X5)&?[X9]:(ssItem(X9)&?[X10]:((ssList(X10)&app(cons(X9,nil),X10)=X3)&lt(X7,X9))))))|?[X11]:(ssItem(X11)&?[X12]:((ssList(X12)&app(cons(X11,nil),X12)=X6)&?[X13]:(ssItem(X13)&?[X14]:((ssList(X14)&app(X14,cons(X13,nil))=X3)&lt(X13,X11))))))))|![X15]:(ssItem(X15)=>![X16]:((~(ssItem(X16))|![X17]:(ssList(X17)=>![X18]:(ssList(X18)=>app(app(app(X17,cons(X15,nil)),cons(X16,nil)),X18)!=X1)))|neq(X15,X16))))|(nil!=X4&nil=X3)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.02/0.29  fof(c_0_1, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax1)).
% 0.02/0.29  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.02/0.29  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.02/0.29  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.02/0.29  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.02/0.29  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>(strictorderedP(X1)&lt;=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(app(X4,cons(X2,X5)),cons(X3,X6))=X1=>lt(X2,X3)))))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax12)).
% 0.02/0.29  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.02/0.29  fof(c_0_8, axiom, (![X1]:(ssItem(X1)=>~(lt(X1,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax90)).
% 0.02/0.29  fof(c_0_9, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~ssList(X4)|X2!=X4)|X1!=X3)|![X5]:(ssList(X5)=>![X6]:((((~ssList(X6)|app(app(X5,X3),X6)!=X4)|~strictorderedP(X3))|?[X7]:(ssItem(X7)&?[X8]:((ssList(X8)&app(X8,cons(X7,nil))=X5)&?[X9]:(ssItem(X9)&?[X10]:((ssList(X10)&app(cons(X9,nil),X10)=X3)&lt(X7,X9))))))|?[X11]:(ssItem(X11)&?[X12]:((ssList(X12)&app(cons(X11,nil),X12)=X6)&?[X13]:(ssItem(X13)&?[X14]:((ssList(X14)&app(X14,cons(X13,nil))=X3)&lt(X13,X11))))))))|![X15]:(ssItem(X15)=>![X16]:((~ssItem(X16)|![X17]:(ssList(X17)=>![X18]:(ssList(X18)=>app(app(app(X17,cons(X15,nil)),cons(X16,nil)),X18)!=X1)))|neq(X15,X16))))|(nil!=X4&nil=X3))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.02/0.29  fof(c_0_10, negated_conjecture, (![X25]:![X26]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(ssList(esk52_0)&((((ssList(esk53_0)&app(app(esk52_0,esk50_0),esk53_0)=esk51_0)&strictorderedP(esk50_0))&(~ssItem(X25)|((~ssList(X26)|app(X26,cons(X25,nil))!=esk52_0)|(~ssItem(X27)|((~ssList(X28)|app(cons(X27,nil),X28)!=esk50_0)|~lt(X25,X27))))))&(~ssItem(X29)|((~ssList(X30)|app(cons(X29,nil),X30)!=esk53_0)|(~ssItem(X31)|((~ssList(X32)|app(X32,cons(X31,nil))!=esk50_0)|~lt(X31,X29))))))))&(ssItem(esk54_0)&((ssItem(esk55_0)&(ssList(esk56_0)&(ssList(esk57_0)&app(app(app(esk56_0,cons(esk54_0,nil)),cons(esk55_0,nil)),esk57_0)=esk48_0)))&~neq(esk54_0,esk55_0))))&(nil=esk51_0|nil!=esk50_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).
% 0.02/0.29  fof(c_0_11, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssItem(X4))|~ssItem(X3))&(((X3=X4|neq(X3,X4))|~ssItem(X4))|~ssItem(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
% 0.02/0.29  cnf(c_0_12, negated_conjecture, (app(app(app(esk56_0,cons(esk54_0,nil)),cons(esk55_0,nil)),esk57_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  cnf(c_0_13, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  cnf(c_0_14, negated_conjecture, (~neq(esk54_0,esk55_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  cnf(c_0_15, plain, (neq(X1,X2)|X1=X2|~ssItem(X1)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.02/0.29  cnf(c_0_16, negated_conjecture, (ssItem(esk55_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  cnf(c_0_17, negated_conjecture, (ssItem(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  cnf(c_0_18, negated_conjecture, (app(app(app(esk56_0,cons(esk54_0,nil)),cons(esk55_0,nil)),esk57_0)=esk50_0), inference(rw,[status(thm)],[c_0_12, c_0_13])).
% 0.02/0.29  cnf(c_0_19, negated_conjecture, (esk55_0=esk54_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17])])).
% 0.02/0.29  fof(c_0_20, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.02/0.29  cnf(c_0_21, negated_conjecture, (app(app(app(esk56_0,cons(esk54_0,nil)),cons(esk54_0,nil)),esk57_0)=esk50_0), inference(rw,[status(thm)],[c_0_18, c_0_19])).
% 0.02/0.29  cnf(c_0_22, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
% 0.02/0.29  cnf(c_0_23, negated_conjecture, (ssList(esk57_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  fof(c_0_24, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
% 0.02/0.29  cnf(c_0_25, negated_conjecture, (app(app(esk56_0,cons(esk54_0,nil)),app(cons(esk54_0,nil),esk57_0))=esk50_0|~ssList(app(esk56_0,cons(esk54_0,nil)))|~ssList(cons(esk54_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])])).
% 0.02/0.29  cnf(c_0_26, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
% 0.02/0.29  cnf(c_0_27, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.29  fof(c_0_28, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 0.02/0.29  cnf(c_0_29, negated_conjecture, (app(app(esk56_0,cons(esk54_0,nil)),app(cons(esk54_0,nil),esk57_0))=esk50_0|~ssList(app(esk56_0,cons(esk54_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_17])])).
% 0.02/0.29  cnf(c_0_30, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
% 0.02/0.29  cnf(c_0_31, negated_conjecture, (ssList(esk56_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  fof(c_0_32, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((~strictorderedP(X7)|(~ssItem(X8)|(~ssItem(X9)|(~ssList(X10)|(~ssList(X11)|(~ssList(X12)|(app(app(X10,cons(X8,X11)),cons(X9,X12))!=X7|lt(X8,X9))))))))|~ssList(X7))&(((ssItem(esk30_1(X7))|strictorderedP(X7))|~ssList(X7))&(((ssItem(esk31_1(X7))|strictorderedP(X7))|~ssList(X7))&(((ssList(esk32_1(X7))|strictorderedP(X7))|~ssList(X7))&(((ssList(esk33_1(X7))|strictorderedP(X7))|~ssList(X7))&(((ssList(esk34_1(X7))|strictorderedP(X7))|~ssList(X7))&(((app(app(esk32_1(X7),cons(esk30_1(X7),esk33_1(X7))),cons(esk31_1(X7),esk34_1(X7)))=X7|strictorderedP(X7))|~ssList(X7))&((~lt(esk30_1(X7),esk31_1(X7))|strictorderedP(X7))|~ssList(X7)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
% 0.02/0.29  cnf(c_0_33, negated_conjecture, (app(app(esk56_0,cons(esk54_0,nil)),app(cons(esk54_0,nil),esk57_0))=esk50_0|~ssList(cons(esk54_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
% 0.02/0.29  fof(c_0_34, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
% 0.02/0.29  fof(c_0_35, plain, (![X1]:(ssItem(X1)=>~lt(X1,X1))), inference(fof_simplification,[status(thm)],[c_0_8])).
% 0.02/0.29  cnf(c_0_36, plain, (lt(X2,X3)|~ssList(X1)|app(app(X4,cons(X2,X5)),cons(X3,X6))!=X1|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssItem(X3)|~ssItem(X2)|~strictorderedP(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.02/0.29  cnf(c_0_37, negated_conjecture, (app(app(esk56_0,cons(esk54_0,nil)),app(cons(esk54_0,nil),esk57_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_26]), c_0_27]), c_0_17])])).
% 0.02/0.29  cnf(c_0_38, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
% 0.02/0.29  cnf(c_0_39, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  fof(c_0_40, plain, (![X2]:(~ssItem(X2)|~lt(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])).
% 0.02/0.29  cnf(c_0_41, plain, (lt(X1,X2)|~strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_36])).
% 0.02/0.29  cnf(c_0_42, negated_conjecture, (app(app(esk56_0,cons(esk54_0,nil)),cons(esk54_0,esk57_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_23]), c_0_17])])).
% 0.02/0.29  cnf(c_0_43, negated_conjecture, (strictorderedP(esk50_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  cnf(c_0_44, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_39, c_0_13])).
% 0.02/0.29  cnf(c_0_45, plain, (~lt(X1,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
% 0.02/0.29  cnf(c_0_46, negated_conjecture, (lt(esk54_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44]), c_0_23]), c_0_27]), c_0_31]), c_0_17])])).
% 0.02/0.29  cnf(c_0_47, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_17])]), ['proof']).
% 0.02/0.29  # SZS output end CNFRefutation
% 0.02/0.29  # Training examples: 0 positive, 0 negative
% 0.02/0.29  
% 0.02/0.29  # -------------------------------------------------
% 0.02/0.29  # User time                : 0.065 s
% 0.02/0.29  # System time              : 0.004 s
% 0.02/0.29  # Total time               : 0.068 s
% 0.02/0.29  # Maximum resident set size: 1892 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
