<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC075-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC075-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC075-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC075-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC075-1.p">SWC075-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n156.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:33 EDT 2016

% Result     : Unsatisfiable 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of clauses        :   59 ( 770 expanded)
%              Number of leaves         :   24 ( 519 expanded)
%              Depth                    :    9
%              Number of atoms          :  126 (1398 expanded)
%              Number of equality atoms :   42 ( 427 expanded)
%              Maximal clause size      :    6 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( app(sk6,sk5) = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( ssList(sk6) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssList(sk5) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( app(sk5,sk6) = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( frontsegP(X3,X1)
    | app(X1,X2) != X3
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause144)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( frontsegP(app(X1,X3),X2)
    | ~ frontsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause137)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause60)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( X1 = X3
    | app(X1,X2) != app(X3,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause151)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( X2 = X3
    | app(X1,X2) != app(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause150)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( ~ ssList(X1)
    | ~ neq(X1,nil)
    | ~ segmentP(sk2,X1)
    | ~ segmentP(sk1,X1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( nil != sk2
    | nil != sk1 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( app(sk6,sk5) = sk1 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( app(sk6,app(sk5,X1)) = app(sk1,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_24]),c_0_3]),c_0_4])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( app(sk5,sk6) = sk2 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_7]),<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( app(sk6,sk2) = app(sk1,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26]),c_0_3])])).

<A NAME="c_0_29"></A>cnf(c_0_29,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_11]),c_0_12])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( frontsegP(app(sk1,sk6),sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_28]),c_0_13]),c_0_3])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( ssList(app(sk1,sk6)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_28]),c_0_3]),c_0_13])])).

<A NAME="c_0_33"></A>cnf(c_0_33,plain,
    ( segmentP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_14]),c_0_12])]),<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( segmentP(app(sk2,X1),sk6)
    | ~ ssList(app(sk2,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_26]),c_0_3]),c_0_4])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( frontsegP(sk1,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_24]),c_0_4]),c_0_3])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( X1 = sk5
    | app(X1,sk6) != sk2
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_26]),c_0_4]),c_0_3])])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( X1 = sk2
    | app(sk6,X1) != app(sk1,sk6)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_28]),c_0_13]),c_0_3])])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( app(sk6,skaf45(app(sk1,sk6),sk6)) = app(sk1,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_31]),c_0_3]),c_0_32])])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( nil = X1
    | ~ segmentP(sk1,X1)
    | ~ segmentP(sk2,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_19]),c_0_12])])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( segmentP(sk1,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_24]),c_0_4]),c_0_3])])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( segmentP(sk2,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_20]),c_0_13]),c_0_12])])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( app(nil,skaf45(sk1,nil)) = sk1 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_35]),c_0_12]),c_0_21])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( segmentP(app(sk1,X1),sk5)
    | ~ ssList(app(sk1,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_24]),c_0_4]),c_0_3])])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( app(X1,X2) = sk5
    | app(X1,app(X2,sk6)) != sk2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_2]),c_0_3])]),<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( app(sk5,app(sk6,X1)) = app(sk2,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_26]),c_0_4]),c_0_3])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( skaf45(app(sk1,sk6),sk6) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_38]),c_0_22])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( sk6 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_40]),c_0_41]),c_0_3])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( app(sk1,nil) = sk1 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_20]),c_0_24]),c_0_12]),c_0_4])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( skaf45(sk1,nil) = sk1 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_42]),c_0_22])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( segmentP(sk1,sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_20]),c_0_21]),c_0_12])])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( sk5 = sk2
    | app(sk2,sk6) != sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_45]),c_0_26]),c_0_4]),c_0_3])])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( app(sk2,nil) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_20]),c_0_26]),c_0_12]),c_0_3])])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( sk2 = sk1 ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_47]),c_0_48]),c_0_47]),<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( sk5 = nil
    | ~ segmentP(sk2,sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_50]),c_0_4])])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( sk5 = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,c_0_47]),c_0_52])])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( segmentP(sk1,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_48]),c_0_12]),c_0_21])])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( sk1 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_53])])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_55]),c_0_55]),c_0_53]),c_0_53]),c_0_53]),c_0_56])]),<A HREF="#c_0_57">c_0_57</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC075-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n156.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:34:53 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.44  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.06/0.44  # and selection function SelectCQArNTEqFirst.
% 0.06/0.44  #
% 0.06/0.44  # Preprocessing time       : 0.009 s
% 0.06/0.44  # Presaturation interreduction done
% 0.06/0.44  
% 0.06/0.44  # Proof found!
% 0.06/0.44  # SZS status Unsatisfiable
% 0.06/0.44  # SZS output start CNFRefutation
% 0.06/0.44  cnf(c_0_0, negated_conjecture, (app(sk6,sk5)=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_11)).
% 0.06/0.44  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.06/0.44  cnf(c_0_2, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
% 0.06/0.44  cnf(c_0_3, negated_conjecture, (ssList(sk6)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
% 0.06/0.44  cnf(c_0_4, negated_conjecture, (ssList(sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
% 0.06/0.44  cnf(c_0_5, negated_conjecture, (app(sk5,sk6)=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_10)).
% 0.06/0.44  cnf(c_0_6, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 0.06/0.44  cnf(c_0_7, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause144)).
% 0.06/0.44  cnf(c_0_8, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.06/0.44  cnf(c_0_9, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause173)).
% 0.06/0.44  cnf(c_0_10, axiom, (frontsegP(app(X1,X3),X2)|~frontsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause137)).
% 0.06/0.44  cnf(c_0_11, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause60)).
% 0.06/0.44  cnf(c_0_12, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.06/0.44  cnf(c_0_13, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 0.06/0.44  cnf(c_0_14, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause74)).
% 0.06/0.44  cnf(c_0_15, axiom, (X1=X3|app(X1,X2)!=app(X3,X2)|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause151)).
% 0.06/0.44  cnf(c_0_16, axiom, (X2=X3|app(X1,X2)!=app(X1,X3)|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause150)).
% 0.06/0.44  cnf(c_0_17, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause132)).
% 0.06/0.44  cnf(c_0_18, negated_conjecture, (~ssList(X1)|~neq(X1,nil)|~segmentP(sk2,X1)|~segmentP(sk1,X1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 0.06/0.44  cnf(c_0_19, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause100)).
% 0.06/0.44  cnf(c_0_20, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause73)).
% 0.06/0.44  cnf(c_0_21, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.06/0.44  cnf(c_0_22, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause51)).
% 0.06/0.44  cnf(c_0_23, negated_conjecture, (nil!=sk2|nil!=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_12)).
% 0.06/0.44  cnf(c_0_24, negated_conjecture, (app(sk6,sk5)=sk1), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 0.06/0.44  cnf(c_0_25, negated_conjecture, (app(sk6,app(sk5,X1))=app(sk1,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_24]), c_0_3]), c_0_4])])).
% 0.06/0.44  cnf(c_0_26, negated_conjecture, (app(sk5,sk6)=sk2), inference(rw,[status(thm)],[c_0_5, c_0_6])).
% 0.06/0.44  cnf(c_0_27, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_7]), c_0_8])).
% 0.06/0.44  cnf(c_0_28, negated_conjecture, (app(sk6,sk2)=app(sk1,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_3])])).
% 0.06/0.44  cnf(c_0_29, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(er,[status(thm)],[c_0_9])).
% 0.06/0.44  cnf(c_0_30, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
% 0.06/0.44  cnf(c_0_31, negated_conjecture, (frontsegP(app(sk1,sk6),sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_13]), c_0_3])])).
% 0.06/0.44  cnf(c_0_32, negated_conjecture, (ssList(app(sk1,sk6))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_28]), c_0_3]), c_0_13])])).
% 0.06/0.44  cnf(c_0_33, plain, (segmentP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_14]), c_0_12])]), c_0_8])).
% 0.06/0.44  cnf(c_0_34, negated_conjecture, (segmentP(app(sk2,X1),sk6)|~ssList(app(sk2,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_26]), c_0_3]), c_0_4])])).
% 0.06/0.44  cnf(c_0_35, negated_conjecture, (frontsegP(sk1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_24]), c_0_4]), c_0_3])])).
% 0.06/0.44  cnf(c_0_36, negated_conjecture, (X1=sk5|app(X1,sk6)!=sk2|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_26]), c_0_4]), c_0_3])])).
% 0.06/0.44  cnf(c_0_37, negated_conjecture, (X1=sk2|app(sk6,X1)!=app(sk1,sk6)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_28]), c_0_13]), c_0_3])])).
% 0.06/0.44  cnf(c_0_38, negated_conjecture, (app(sk6,skaf45(app(sk1,sk6),sk6))=app(sk1,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_31]), c_0_3]), c_0_32])])).
% 0.06/0.44  cnf(c_0_39, negated_conjecture, (nil=X1|~segmentP(sk1,X1)|~segmentP(sk2,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_12])])).
% 0.06/0.44  cnf(c_0_40, negated_conjecture, (segmentP(sk1,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_24]), c_0_4]), c_0_3])])).
% 0.06/0.44  cnf(c_0_41, negated_conjecture, (segmentP(sk2,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_20]), c_0_13]), c_0_12])])).
% 0.06/0.44  cnf(c_0_42, negated_conjecture, (app(nil,skaf45(sk1,nil))=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_35]), c_0_12]), c_0_21])])).
% 0.06/0.44  cnf(c_0_43, negated_conjecture, (segmentP(app(sk1,X1),sk5)|~ssList(app(sk1,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_24]), c_0_4]), c_0_3])])).
% 0.06/0.44  cnf(c_0_44, negated_conjecture, (app(X1,X2)=sk5|app(X1,app(X2,sk6))!=sk2|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_2]), c_0_3])]), c_0_8])).
% 0.06/0.44  cnf(c_0_45, negated_conjecture, (app(sk5,app(sk6,X1))=app(sk2,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_26]), c_0_4]), c_0_3])])).
% 0.06/0.44  cnf(c_0_46, negated_conjecture, (skaf45(app(sk1,sk6),sk6)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_22])])).
% 0.06/0.44  cnf(c_0_47, negated_conjecture, (sk6=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_3])])).
% 0.06/0.44  cnf(c_0_48, negated_conjecture, (app(sk1,nil)=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_20]), c_0_24]), c_0_12]), c_0_4])])).
% 0.06/0.44  cnf(c_0_49, negated_conjecture, (skaf45(sk1,nil)=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_42]), c_0_22])])).
% 0.06/0.44  cnf(c_0_50, negated_conjecture, (segmentP(sk1,sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_20]), c_0_21]), c_0_12])])).
% 0.06/0.44  cnf(c_0_51, negated_conjecture, (sk5=sk2|app(sk2,sk6)!=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_26]), c_0_4]), c_0_3])])).
% 0.06/0.44  cnf(c_0_52, negated_conjecture, (app(sk2,nil)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_20]), c_0_26]), c_0_12]), c_0_3])])).
% 0.06/0.44  cnf(c_0_53, negated_conjecture, (sk2=sk1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_47]), c_0_49])).
% 0.06/0.44  cnf(c_0_54, negated_conjecture, (sk5=nil|~segmentP(sk2,sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_50]), c_0_4])])).
% 0.06/0.44  cnf(c_0_55, negated_conjecture, (sk5=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_47]), c_0_52])])).
% 0.06/0.44  cnf(c_0_56, negated_conjecture, (segmentP(sk1,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_48]), c_0_12]), c_0_21])])).
% 0.06/0.44  cnf(c_0_57, negated_conjecture, (sk1!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_53])])).
% 0.06/0.44  cnf(c_0_58, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55]), c_0_55]), c_0_53]), c_0_53]), c_0_53]), c_0_56])]), c_0_57]), ['proof']).
% 0.06/0.44  # SZS output end CNFRefutation
% 0.06/0.44  # Training examples: 0 positive, 0 negative
% 0.06/0.44  
% 0.06/0.44  # -------------------------------------------------
% 0.06/0.44  # User time                : 0.208 s
% 0.06/0.44  # System time              : 0.006 s
% 0.06/0.44  # Total time               : 0.214 s
% 0.06/0.44  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
