<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC126-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC126-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC126-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC126-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC126-1.p">SWC126-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n146.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:45 EDT 2016

% Result     : Unsatisfiable 13.32s
% Output     : CNFRefutation 13.32s
% Verified   : 
% Statistics : Number of clauses        :   91 (1611 expanded)
%              Number of leaves         :   39 (1026 expanded)
%              Depth                    :   20
%              Number of atoms          :  209 (2676 expanded)
%              Number of equality atoms :   41 ( 738 expanded)
%              Maximal clause size      :    7 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( neq(sk2,nil)
    | neq(sk2,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( cons(sk5,nil) = sk3
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_14)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssItem(sk5)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( frontsegP(X3,X1)
    | app(X1,X2) != X3
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause144)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause99)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( nil = X2
    | tl(app(X2,X1)) = app(tl(X2),X1)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause133)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( nil = X2
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause119)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( rearsegP(X3,X2)
    | app(X1,X2) != X3
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause143)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( frontsegP(app(X1,X3),X2)
    | ~ frontsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause137)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause60)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( rearsegP(app(X3,X1),X2)
    | ~ rearsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause136)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( memberP(cons(X2,X3),X1)
    | X1 != X2
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause138)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause157)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( app(skaf46(X1,X2),X2) = X1
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause131)).

<A NAME="c_0_23"></A>cnf(c_0_23,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_24"></A>cnf(c_0_24,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_25"></A>cnf(c_0_25,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_26"></A>cnf(c_0_26,axiom,
    ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
    | ~ memberP(X1,X2)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause169)).

<A NAME="c_0_27"></A>cnf(c_0_27,axiom,
    ( ssList(skaf46(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause50)).

<A NAME="c_0_28"></A>cnf(c_0_28,axiom,
    ( segmentP(app(app(X4,X1),X3),X2)
    | ~ segmentP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X4)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause172)).

<A NAME="c_0_29"></A>cnf(c_0_29,axiom,
    ( segmentP(X1,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause57)).

<A NAME="c_0_30"></A>cnf(c_0_30,axiom,
    ( ssList(skaf43(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause52)).

<A NAME="c_0_31"></A>cnf(c_0_31,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_32"></A>cnf(c_0_32,axiom,
    ( frontsegP(X1,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause61)).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( memberP(sk4,sk5)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_15)).

<A NAME="c_0_34"></A>cnf(c_0_34,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( ~ neq(sk1,nil)
    | ~ segmentP(sk2,sk1)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_16)).

<A NAME="c_0_37"></A>cnf(c_0_37,axiom,
    ( segmentP(X1,X3)
    | ~ segmentP(X1,X2)
    | ~ segmentP(X2,X3)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause152)).

<A NAME="c_0_38"></A>cnf(c_0_38,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( neq(sk2,nil) ),
    inference(cn,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( cons(sk5,nil) = sk1 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_1">c_0_1</A>,c_0_2]),c_0_3]),c_0_39])])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( ssItem(sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,c_0_3]),c_0_39])])).

<A NAME="c_0_42"></A>cnf(c_0_42,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_5]),<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( app(sk1,X1) = cons(sk5,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_40]),c_0_41])])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( frontsegP(cons(sk5,X1),sk1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_43]),c_0_8])])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( frontsegP(sk1,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_40]),c_0_9])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( app(sk1,skaf45(sk1,sk1)) = sk1 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_45]),c_0_8])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( tl(sk1) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_40]),c_0_41]),c_0_9])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( sk1 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_40]),c_0_41]),c_0_9])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( app(nil,skaf45(sk1,sk1)) = nil ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_46]),c_0_47]),c_0_47]),c_0_8]),c_0_14])]),<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( skaf45(sk1,sk1) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_49]),c_0_14]),c_0_9])])).

<A NAME="c_0_51"></A>cnf(c_0_51,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_16]),<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( app(sk1,nil) = sk1 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( rearsegP(sk1,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,c_0_52]),c_0_9]),c_0_8])])).

<A NAME="c_0_54"></A>cnf(c_0_54,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_18]),c_0_9])])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( rearsegP(app(X1,sk1),nil)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_53]),c_0_9]),c_0_8])])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( frontsegP(sk1,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_52]),c_0_9]),c_0_8])])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( rearsegP(cons(sk5,sk1),nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_43]),c_0_8])])).

<A NAME="c_0_58"></A>cnf(c_0_58,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( cons(sk5,app(nil,X1)) = app(sk1,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_40]),c_0_41]),c_0_9])])).

<A NAME="c_0_60"></A>cnf(c_0_60,plain,
    ( app(nil,skaf45(X1,nil)) = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_18]),c_0_9])])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( app(nil,skaf45(sk1,nil)) = sk1 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_56]),c_0_9]),c_0_8])])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( app(skaf46(cons(sk5,sk1),nil),nil) = cons(sk5,sk1)
    | ~ ssList(cons(sk5,sk1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_57]),c_0_9])])).

<A NAME="c_0_63"></A>cnf(c_0_63,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( memberP(sk1,sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_40]),c_0_41]),c_0_9])])).

<A NAME="c_0_65"></A>cnf(c_0_65,negated_conjecture,
    ( app(sk1,skaf45(X1,nil)) = cons(sk5,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,c_0_60]),c_0_14])])).

<A NAME="c_0_66"></A>cnf(c_0_66,negated_conjecture,
    ( skaf45(sk1,nil) = sk1 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_61]),c_0_14])])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( app(skaf46(cons(sk5,sk1),nil),nil) = cons(sk5,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_25]),c_0_41]),c_0_8])])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( ~ ssList(app(app(X1,cons(sk5,X2)),sk1))
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(app(X1,cons(sk5,X2)),sk1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>,c_0_40]),c_0_41]),c_0_9])])).

<A NAME="c_0_69"></A>cnf(c_0_69,negated_conjecture,
    ( app(skaf42(sk1,sk5),cons(sk5,skaf43(sk5,sk1))) = sk1 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_64]),c_0_41]),c_0_8])])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( app(sk1,sk1) = cons(sk5,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_66]),c_0_8])])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( ssList(cons(sk5,sk1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>,c_0_67]),c_0_27]),c_0_9])])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,negated_conjecture,
    ( app(nil,sk1) = sk1 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,<A HREF="#c_0_66">c_0_66</A>])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( ssList(cons(sk5,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>,c_0_43]),c_0_8])])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( ~ duplicatefreeP(cons(sk5,sk1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_69]),c_0_70]),c_0_71]),c_0_30]),c_0_31]),c_0_70])])).

<A NAME="c_0_76"></A>cnf(c_0_76,plain,
    ( app(X1,skaf45(X1,X1)) = X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,negated_conjecture,
    ( memberP(sk2,sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_3]),c_0_3]),c_0_39])])).

<A NAME="c_0_78"></A>cnf(c_0_78,negated_conjecture,
    ( segmentP(app(sk1,X1),sk1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_73]),c_0_9]),c_0_8])])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( app(sk1,skaf45(cons(sk5,X1),sk1)) = cons(sk5,X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_44]),c_0_8])]),<A HREF="#c_0_74">c_0_74</A>])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_34]),c_0_71])])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( segmentP(app(X1,X2),X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_76]),c_0_14])]),<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( app(skaf42(sk2,sk5),cons(sk5,skaf43(sk5,sk2))) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_77]),c_0_41]),c_0_35])])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( segmentP(cons(sk5,X1),sk1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_79]),c_0_14])])).

<A NAME="c_0_84"></A>cnf(c_0_84,plain,
    ( ssList(cons(X1,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_80])])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( segmentP(sk2,cons(sk5,skaf43(sk5,sk2)))
    | ~ ssList(cons(sk5,skaf43(sk5,sk2))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_82]),c_0_31])])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( ~ neq(sk1,nil)
    | ~ segmentP(sk2,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_3]),c_0_39])])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( segmentP(X1,sk1)
    | ~ segmentP(X1,cons(sk5,X2))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_83]),c_0_8])]),<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( segmentP(sk2,cons(sk5,skaf43(sk5,sk2))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_84]),c_0_30])])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( ~ segmentP(sk2,sk1) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>,c_0_38]),c_0_8]),c_0_9])]),<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,c_0_88]),c_0_35]),c_0_30])]),<A HREF="#c_0_89">c_0_89</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC126-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n146.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:35:23 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 13.32/13.54  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 13.32/13.54  # and selection function SelectCQArNTEqFirst.
% 13.32/13.54  #
% 13.32/13.54  # Preprocessing time       : 0.009 s
% 13.32/13.54  # Presaturation interreduction done
% 13.32/13.54  
% 13.32/13.54  # Proof found!
% 13.32/13.54  # SZS status Unsatisfiable
% 13.32/13.54  # SZS output start CNFRefutation
% 13.32/13.54  cnf(c_0_0, negated_conjecture, (neq(sk2,nil)|neq(sk2,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 13.32/13.54  cnf(c_0_1, negated_conjecture, (cons(sk5,nil)=sk3|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_14)).
% 13.32/13.54  cnf(c_0_2, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 13.32/13.54  cnf(c_0_3, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 13.32/13.54  cnf(c_0_4, negated_conjecture, (ssItem(sk5)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_13)).
% 13.32/13.54  cnf(c_0_5, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause144)).
% 13.32/13.54  cnf(c_0_6, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
% 13.32/13.54  cnf(c_0_7, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause120)).
% 13.32/13.54  cnf(c_0_8, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 13.32/13.54  cnf(c_0_9, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 13.32/13.54  cnf(c_0_10, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause132)).
% 13.32/13.54  cnf(c_0_11, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause96)).
% 13.32/13.54  cnf(c_0_12, axiom, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause99)).
% 13.32/13.54  cnf(c_0_13, axiom, (nil=X2|tl(app(X2,X1))=app(tl(X2),X1)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause133)).
% 13.32/13.54  cnf(c_0_14, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause51)).
% 13.32/13.54  cnf(c_0_15, axiom, (nil=X2|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause119)).
% 13.32/13.54  cnf(c_0_16, axiom, (rearsegP(X3,X2)|app(X1,X2)!=X3|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause143)).
% 13.32/13.54  cnf(c_0_17, axiom, (frontsegP(app(X1,X3),X2)|~frontsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause137)).
% 13.32/13.54  cnf(c_0_18, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause60)).
% 13.32/13.54  cnf(c_0_19, axiom, (rearsegP(app(X3,X1),X2)|~rearsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause136)).
% 13.32/13.54  cnf(c_0_20, axiom, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause138)).
% 13.32/13.54  cnf(c_0_21, axiom, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause157)).
% 13.32/13.54  cnf(c_0_22, axiom, (app(skaf46(X1,X2),X2)=X1|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause131)).
% 13.32/13.54  cnf(c_0_23, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause179)).
% 13.32/13.54  cnf(c_0_24, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause74)).
% 13.32/13.54  cnf(c_0_25, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause86)).
% 13.32/13.54  cnf(c_0_26, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause169)).
% 13.32/13.54  cnf(c_0_27, axiom, (ssList(skaf46(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause50)).
% 13.32/13.54  cnf(c_0_28, axiom, (segmentP(app(app(X4,X1),X3),X2)|~segmentP(X1,X2)|~ssList(X3)|~ssList(X4)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause172)).
% 13.32/13.54  cnf(c_0_29, axiom, (segmentP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause57)).
% 13.32/13.54  cnf(c_0_30, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause52)).
% 13.32/13.54  cnf(c_0_31, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause53)).
% 13.32/13.54  cnf(c_0_32, axiom, (frontsegP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause61)).
% 13.32/13.54  cnf(c_0_33, negated_conjecture, (memberP(sk4,sk5)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_15)).
% 13.32/13.54  cnf(c_0_34, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause72)).
% 13.32/13.54  cnf(c_0_35, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
% 13.32/13.54  cnf(c_0_36, negated_conjecture, (~neq(sk1,nil)|~segmentP(sk2,sk1)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_16)).
% 13.32/13.54  cnf(c_0_37, axiom, (segmentP(X1,X3)|~segmentP(X1,X2)|~segmentP(X2,X3)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause152)).
% 13.32/13.54  cnf(c_0_38, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause100)).
% 13.32/13.54  cnf(c_0_39, negated_conjecture, (neq(sk2,nil)), inference(cn,[status(thm)],[c_0_0])).
% 13.32/13.54  cnf(c_0_40, negated_conjecture, (cons(sk5,nil)=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_1, c_0_2]), c_0_3]), c_0_39])])).
% 13.32/13.54  cnf(c_0_41, negated_conjecture, (ssItem(sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_3]), c_0_39])])).
% 13.32/13.54  cnf(c_0_42, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_5]), c_0_6])).
% 13.32/13.54  cnf(c_0_43, negated_conjecture, (app(sk1,X1)=cons(sk5,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_40]), c_0_41])])).
% 13.32/13.54  cnf(c_0_44, negated_conjecture, (frontsegP(cons(sk5,X1),sk1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_8])])).
% 13.32/13.54  cnf(c_0_45, negated_conjecture, (frontsegP(sk1,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_40]), c_0_9])])).
% 13.32/13.54  cnf(c_0_46, negated_conjecture, (app(sk1,skaf45(sk1,sk1))=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_45]), c_0_8])])).
% 13.32/13.54  cnf(c_0_47, negated_conjecture, (tl(sk1)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_40]), c_0_41]), c_0_9])])).
% 13.32/13.54  cnf(c_0_48, negated_conjecture, (sk1!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_40]), c_0_41]), c_0_9])])).
% 13.32/13.54  cnf(c_0_49, negated_conjecture, (app(nil,skaf45(sk1,sk1))=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_46]), c_0_47]), c_0_47]), c_0_8]), c_0_14])]), c_0_48])).
% 13.32/13.54  cnf(c_0_50, negated_conjecture, (skaf45(sk1,sk1)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_49]), c_0_14]), c_0_9])])).
% 13.32/13.54  cnf(c_0_51, plain, (rearsegP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_16]), c_0_6])).
% 13.32/13.54  cnf(c_0_52, negated_conjecture, (app(sk1,nil)=sk1), inference(rw,[status(thm)],[c_0_46, c_0_50])).
% 13.32/13.54  cnf(c_0_53, negated_conjecture, (rearsegP(sk1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_9]), c_0_8])])).
% 13.32/13.54  cnf(c_0_54, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_9])])).
% 13.32/13.54  cnf(c_0_55, negated_conjecture, (rearsegP(app(X1,sk1),nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_53]), c_0_9]), c_0_8])])).
% 13.32/13.54  cnf(c_0_56, negated_conjecture, (frontsegP(sk1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_52]), c_0_9]), c_0_8])])).
% 13.32/13.54  cnf(c_0_57, negated_conjecture, (rearsegP(cons(sk5,sk1),nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_43]), c_0_8])])).
% 13.32/13.54  cnf(c_0_58, plain, (memberP(cons(X1,X2),X1)|~ssItem(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_20])).
% 13.32/13.54  cnf(c_0_59, negated_conjecture, (cons(sk5,app(nil,X1))=app(sk1,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_40]), c_0_41]), c_0_9])])).
% 13.32/13.54  cnf(c_0_60, plain, (app(nil,skaf45(X1,nil))=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_18]), c_0_9])])).
% 13.32/13.54  cnf(c_0_61, negated_conjecture, (app(nil,skaf45(sk1,nil))=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_56]), c_0_9]), c_0_8])])).
% 13.32/13.54  cnf(c_0_62, negated_conjecture, (app(skaf46(cons(sk5,sk1),nil),nil)=cons(sk5,sk1)|~ssList(cons(sk5,sk1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_57]), c_0_9])])).
% 13.32/13.54  cnf(c_0_63, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_23])).
% 13.32/13.54  cnf(c_0_64, negated_conjecture, (memberP(sk1,sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_40]), c_0_41]), c_0_9])])).
% 13.32/13.54  cnf(c_0_65, negated_conjecture, (app(sk1,skaf45(X1,nil))=cons(sk5,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_14])])).
% 13.32/13.54  cnf(c_0_66, negated_conjecture, (skaf45(sk1,nil)=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_61]), c_0_14])])).
% 13.32/13.54  cnf(c_0_67, negated_conjecture, (app(skaf46(cons(sk5,sk1),nil),nil)=cons(sk5,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_25]), c_0_41]), c_0_8])])).
% 13.32/13.54  cnf(c_0_68, negated_conjecture, (~ssList(app(app(X1,cons(sk5,X2)),sk1))|~ssList(X2)|~ssList(X1)|~duplicatefreeP(app(app(X1,cons(sk5,X2)),sk1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_40]), c_0_41]), c_0_9])])).
% 13.32/13.54  cnf(c_0_69, negated_conjecture, (app(skaf42(sk1,sk5),cons(sk5,skaf43(sk5,sk1)))=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_64]), c_0_41]), c_0_8])])).
% 13.32/13.54  cnf(c_0_70, negated_conjecture, (app(sk1,sk1)=cons(sk5,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_8])])).
% 13.32/13.54  cnf(c_0_71, negated_conjecture, (ssList(cons(sk5,sk1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_67]), c_0_27]), c_0_9])])).
% 13.32/13.54  cnf(c_0_72, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
% 13.32/13.54  cnf(c_0_73, negated_conjecture, (app(nil,sk1)=sk1), inference(rw,[status(thm)],[c_0_61, c_0_66])).
% 13.32/13.54  cnf(c_0_74, negated_conjecture, (ssList(cons(sk5,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_43]), c_0_8])])).
% 13.32/13.54  cnf(c_0_75, negated_conjecture, (~duplicatefreeP(cons(sk5,sk1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), c_0_71]), c_0_30]), c_0_31]), c_0_70])])).
% 13.32/13.54  cnf(c_0_76, plain, (app(X1,skaf45(X1,X1))=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_10, c_0_32])).
% 13.32/13.54  cnf(c_0_77, negated_conjecture, (memberP(sk2,sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_3]), c_0_3]), c_0_39])])).
% 13.32/13.54  cnf(c_0_78, negated_conjecture, (segmentP(app(sk1,X1),sk1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_9]), c_0_8])])).
% 13.32/13.54  cnf(c_0_79, negated_conjecture, (app(sk1,skaf45(cons(sk5,X1),sk1))=cons(sk5,X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_44]), c_0_8])]), c_0_74])).
% 13.32/13.54  cnf(c_0_80, negated_conjecture, (ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_34]), c_0_71])])).
% 13.32/13.54  cnf(c_0_81, plain, (segmentP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_76]), c_0_14])]), c_0_6])).
% 13.32/13.54  cnf(c_0_82, negated_conjecture, (app(skaf42(sk2,sk5),cons(sk5,skaf43(sk5,sk2)))=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_77]), c_0_41]), c_0_35])])).
% 13.32/13.54  cnf(c_0_83, negated_conjecture, (segmentP(cons(sk5,X1),sk1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_14])])).
% 13.32/13.54  cnf(c_0_84, plain, (ssList(cons(X1,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_80])])).
% 13.32/13.54  cnf(c_0_85, negated_conjecture, (segmentP(sk2,cons(sk5,skaf43(sk5,sk2)))|~ssList(cons(sk5,skaf43(sk5,sk2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_31])])).
% 13.32/13.54  cnf(c_0_86, negated_conjecture, (~neq(sk1,nil)|~segmentP(sk2,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_3]), c_0_39])])).
% 13.32/13.54  cnf(c_0_87, negated_conjecture, (segmentP(X1,sk1)|~segmentP(X1,cons(sk5,X2))|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_83]), c_0_8])]), c_0_84])).
% 13.32/13.54  cnf(c_0_88, negated_conjecture, (segmentP(sk2,cons(sk5,skaf43(sk5,sk2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_84]), c_0_30])])).
% 13.32/13.54  cnf(c_0_89, negated_conjecture, (~segmentP(sk2,sk1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_38]), c_0_8]), c_0_9])]), c_0_48])).
% 13.32/13.54  cnf(c_0_90, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_35]), c_0_30])]), c_0_89]), ['proof']).
% 13.32/13.54  # SZS output end CNFRefutation
% 13.32/13.54  # Training examples: 0 positive, 0 negative
% 13.32/13.58  
% 13.32/13.58  # -------------------------------------------------
% 13.32/13.58  # User time                : 12.957 s
% 13.32/13.58  # System time              : 0.394 s
% 13.32/13.58  # Total time               : 13.352 s
% 13.32/13.58  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
