<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC324+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC324+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC324+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC324%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC324+1.p">SWC324+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n144.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:30 EDT 2016

% Result     : Theorem 0.07s
% Output     : CNFRefutation 0.07s
% Verified   : 
% Statistics : Number of formulae       :   35 ( 166 expanded)
%              Number of clauses        :   23 (  66 expanded)
%              Number of leaves         :    6 (  44 expanded)
%              Depth                    :   10
%              Number of atoms          :  150 (1376 expanded)
%              Number of equality atoms :   67 ( 683 expanded)
%              Maximal formula depth    :   20 (   5 average)
%              Maximal clause size      :   16 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ~ ssList(X4)
                  | X2 != X4
                  | X1 != X3
                  | ? [X5] :
                      ( ssList(X5)
                      & ? [X6] :
                          ( ssList(X6)
                          & app(X5,X6) = X2
                          & app(X6,X5) = X1 ) )
                  | ? [X7] :
                      ( ssItem(X7)
                      & ? [X8] :
                          ( ssList(X8)
                          & app(X8,cons(X7,nil)) != X3
                          & app(cons(X7,nil),X8) = X4 ) )
                  | ( nil != X3
                    & nil = X4 ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil = X1
        | ? [X2] :
            ( ssList(X2)
            & ? [X3] :
                ( ssItem(X3)
                & cons(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax20)).

<A NAME="c_0_6"></A>fof(c_0_6,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ~ ssList(X4)
                    | X2 != X4
                    | X1 != X3
                    | ? [X5] :
                        ( ssList(X5)
                        & ? [X6] :
                            ( ssList(X6)
                            & app(X5,X6) = X2
                            & app(X6,X5) = X1 ) )
                    | ? [X7] :
                        ( ssItem(X7)
                        & ? [X8] :
                            ( ssList(X8)
                            & app(X8,cons(X7,nil)) != X3
                            & app(cons(X7,nil),X8) = X4 ) )
                    | ( nil != X3
                      & nil = X4 ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_7"></A>fof(c_0_7,negated_conjecture,(
    ! [X13,X14,X15,X16] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ( ~ ssList(X13)
        | ~ ssList(X14)
        | app(X13,X14) != esk49_0
        | app(X14,X13) != esk48_0 )
      & ( ~ ssItem(X15)
        | ~ ssList(X16)
        | app(X16,cons(X15,nil)) = esk50_0
        | app(cons(X15,nil),X16) != esk51_0 )
      & ( nil = esk50_0
        | nil != esk51_0 ) ) ),
    inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( app(X1,X2) != esk48_0
    | app(X2,X1) != esk49_0
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_11"></A>fof(c_0_11,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( app(X1,X2) != esk51_0
    | app(X2,X1) != esk50_0
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_9]),<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_14"></A>fof(c_0_14,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_15"></A>fof(c_0_15,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( app(X1,cons(X2,nil)) != esk50_0
    | cons(X2,X1) != esk51_0
    | ~ ssList(cons(X2,nil))
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,plain,
    ( nil = app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != X1
    | nil != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( app(X2,cons(X1,nil)) = esk50_0
    | app(cons(X1,nil),X2) != esk51_0
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( app(X1,cons(X2,nil)) != esk50_0
    | cons(X2,X1) != esk51_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_17]),c_0_18])])).

<A NAME="c_0_22"></A>fof(c_0_22,plain,(
    ! [X4] :
      ( ( ssList(esk44_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( ssItem(esk45_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( cons(esk45_1(X4),esk44_1(X4)) = X4
        | nil = X4
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_19])]),c_0_18])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( nil = esk50_0
    | nil != esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( cons(X1,X2) != esk51_0
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_13]),<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,plain,
    ( nil = X1
    | cons(esk45_1(X1),esk44_1(X1)) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( esk51_0 != nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_24]),c_0_24]),c_0_18])]),<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( ~ ssList(esk44_1(esk51_0))
    | ~ ssItem(esk45_1(esk51_0)) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_27])]),c_0_28])]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( nil = X1
    | ssList(esk44_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( ~ ssItem(esk45_1(esk51_0)) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31]),c_0_28])]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,plain,
    ( nil = X1
    | ssItem(esk45_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_33]),c_0_28])]),<A HREF="#c_0_29">c_0_29</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC324+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.24  % Computer   : n144.star.cs.uiowa.edu
% 0.02/0.24  % Model      : x86_64 x86_64
% 0.02/0.24  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.24  % Memory     : 32218.75MB
% 0.02/0.24  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.24  % CPULimit   : 300
% 0.02/0.24  % DateTime   : Tue May 10 20:41:52 CDT 2016
% 0.02/0.24  % CPUTime    : 
% 0.02/0.24  # Version: 1.9.1pre011
% 0.02/0.24  # No SInE strategy applied
% 0.02/0.24  # Trying AutoSched0 for 151 seconds
% 0.07/0.31  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.07/0.31  # and selection function SelectCQPrecWNTNp.
% 0.07/0.31  #
% 0.07/0.31  # Preprocessing time       : 0.014 s
% 0.07/0.31  # Presaturation interreduction done
% 0.07/0.31  
% 0.07/0.31  # Proof found!
% 0.07/0.31  # SZS status Theorem
% 0.07/0.31  # SZS output start CNFRefutation
% 0.07/0.31  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~(ssList(X4))|X2!=X4)|X1!=X3)|?[X5]:(ssList(X5)&?[X6]:((ssList(X6)&app(X5,X6)=X2)&app(X6,X5)=X1)))|?[X7]:(ssItem(X7)&?[X8]:((ssList(X8)&app(X8,cons(X7,nil))!=X3)&app(cons(X7,nil),X8)=X4)))|(nil!=X3&nil=X4)))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.07/0.31  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.07/0.31  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.07/0.31  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.07/0.31  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.07/0.31  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax20)).
% 0.07/0.31  fof(c_0_6, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~ssList(X4)|X2!=X4)|X1!=X3)|?[X5]:(ssList(X5)&?[X6]:((ssList(X6)&app(X5,X6)=X2)&app(X6,X5)=X1)))|?[X7]:(ssItem(X7)&?[X8]:((ssList(X8)&app(X8,cons(X7,nil))!=X3)&app(cons(X7,nil),X8)=X4)))|(nil!=X3&nil=X4))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.07/0.31  fof(c_0_7, negated_conjecture, (![X13]:![X14]:![X15]:![X16]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(~ssList(X13)|((~ssList(X14)|app(X13,X14)!=esk49_0)|app(X14,X13)!=esk48_0)))&(~ssItem(X15)|((~ssList(X16)|app(X16,cons(X15,nil))=esk50_0)|app(cons(X15,nil),X16)!=esk51_0)))&(nil=esk50_0|nil!=esk51_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
% 0.07/0.31  cnf(c_0_8, negated_conjecture, (app(X1,X2)!=esk48_0|app(X2,X1)!=esk49_0|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.07/0.31  cnf(c_0_9, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.07/0.31  cnf(c_0_10, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.07/0.31  fof(c_0_11, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.07/0.31  cnf(c_0_12, negated_conjecture, (app(X1,X2)!=esk51_0|app(X2,X1)!=esk50_0|~ssList(X1)|~ssList(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_9]), c_0_10])).
% 0.07/0.31  cnf(c_0_13, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.07/0.31  fof(c_0_14, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.07/0.31  fof(c_0_15, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
% 0.07/0.31  cnf(c_0_16, negated_conjecture, (app(X1,cons(X2,nil))!=esk50_0|cons(X2,X1)!=esk51_0|~ssList(cons(X2,nil))|~ssList(X1)|~ssItem(X2)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
% 0.07/0.31  cnf(c_0_17, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.07/0.31  cnf(c_0_18, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.07/0.31  cnf(c_0_19, plain, (nil=app(X1,X2)|~ssList(X1)|~ssList(X2)|nil!=X1|nil!=X2), inference(split_conjunct,[status(thm)],[c_0_15])).
% 0.07/0.31  cnf(c_0_20, negated_conjecture, (app(X2,cons(X1,nil))=esk50_0|app(cons(X1,nil),X2)!=esk51_0|~ssList(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.07/0.31  cnf(c_0_21, negated_conjecture, (app(X1,cons(X2,nil))!=esk50_0|cons(X2,X1)!=esk51_0|~ssList(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])])).
% 0.07/0.31  fof(c_0_22, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
% 0.07/0.31  cnf(c_0_23, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.07/0.31  cnf(c_0_24, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_19])]), c_0_18])])).
% 0.07/0.31  cnf(c_0_25, negated_conjecture, (nil=esk50_0|nil!=esk51_0), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.07/0.31  cnf(c_0_26, negated_conjecture, (cons(X1,X2)!=esk51_0|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_13]), c_0_21])).
% 0.07/0.31  cnf(c_0_27, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.07/0.31  cnf(c_0_28, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_23, c_0_9])).
% 0.07/0.31  cnf(c_0_29, negated_conjecture, (esk51_0!=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_24]), c_0_24]), c_0_18])]), c_0_25])).
% 0.07/0.31  cnf(c_0_30, negated_conjecture, (~ssList(esk44_1(esk51_0))|~ssItem(esk45_1(esk51_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27])]), c_0_28])]), c_0_29])).
% 0.07/0.31  cnf(c_0_31, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.07/0.31  cnf(c_0_32, negated_conjecture, (~ssItem(esk45_1(esk51_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_28])]), c_0_29])).
% 0.07/0.31  cnf(c_0_33, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.07/0.31  cnf(c_0_34, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_28])]), c_0_29]), ['proof']).
% 0.07/0.31  # SZS output end CNFRefutation
% 0.07/0.31  # Training examples: 0 positive, 0 negative
% 0.07/0.31  
% 0.07/0.31  # -------------------------------------------------
% 0.07/0.31  # User time                : 0.065 s
% 0.07/0.31  # System time              : 0.003 s
% 0.07/0.31  # Total time               : 0.069 s
% 0.07/0.31  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
