<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC160-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC160-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC160-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC160-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC160-1.p">SWC160-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n114.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:53 EDT 2016

% Result     : Unsatisfiable 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of clauses        :   48 ( 186 expanded)
%              Number of leaves         :   22 ( 117 expanded)
%              Depth                    :   14
%              Number of atoms          :  122 ( 514 expanded)
%              Number of equality atoms :   17 (  39 expanded)
%              Maximal clause size      :    9 (   3 average)
%              Maximal term depth       :    6 (   2 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_1"></A>cnf(c_0_1,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_3"></A>cnf(c_0_3,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_4"></A>cnf(c_0_4,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( app(app(app(app(sk8,cons(sk6,nil)),sk9),cons(sk7,nil)),sk10) = sk1 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_16)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ssList(sk8) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( ssList(sk9) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_14)).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ssItem(sk6) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( ssItem(sk7) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( ssList(sk10) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_15)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( lt(X2,X4)
    | app(app(X1,cons(X2,X3)),cons(X4,X5)) != X6
    | ~ ssList(X5)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X4)
    | ~ ssItem(X2)
    | ~ strictorderedP(X6)
    | ~ ssList(X6) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause181)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( lt(X1,X3)
    | ~ leq(X1,X2)
    | ~ lt(X2,X3)
    | ~ ssItem(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause147)).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( leq(sk7,sk6) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_17)).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( strictorderedP(sk3) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( ~ lt(X1,X1)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause63)).

<A NAME="c_0_22"></A>cnf(c_0_22,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,plain,
    ( ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(app(app(X3,cons(X2,X4)),cons(X2,X5)))
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,plain,
    ( ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(app(cons(X2,X3),cons(X2,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_2]),c_0_3])]),<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( app(app(app(app(sk8,cons(sk6,nil)),sk9),cons(sk7,nil)),sk10) = sk3 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,plain,
    ( ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(cons(X2,cons(X2,X3)))
    | ~ ssList(X3) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_7]),c_0_3])]),<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( app(app(app(sk8,app(cons(sk6,nil),sk9)),cons(sk7,nil)),sk10) = sk3
    | ~ ssList(cons(sk6,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_8]),c_0_9]),c_0_10])])).

<A NAME="c_0_28"></A>cnf(c_0_28,plain,
    ( ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_4]),<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( app(app(app(sk8,app(cons(sk6,nil),sk9)),cons(sk7,nil)),sk10) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_4]),c_0_11]),c_0_3])])).

<A NAME="c_0_30"></A>cnf(c_0_30,plain,
    ( ssItem(X1)
    | ~ ssItem(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( app(app(app(sk8,cons(sk6,sk9)),cons(sk7,nil)),sk10) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_7]),c_0_11]),c_0_10])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( app(app(sk8,cons(sk6,sk9)),app(cons(sk7,nil),sk10)) = sk3
    | ~ ssList(app(sk8,cons(sk6,sk9)))
    | ~ ssList(cons(sk7,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_31]),c_0_14])])).

<A NAME="c_0_34"></A>cnf(c_0_34,plain,
    ( ssList(cons(X1,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,c_0_32])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( app(app(sk8,cons(sk6,sk9)),app(cons(sk7,nil),sk10)) = sk3
    | ~ ssList(app(sk8,cons(sk6,sk9))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_34]),c_0_3])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( app(app(sk8,cons(sk6,sk9)),app(cons(sk7,nil),sk10)) = sk3
    | ~ ssList(cons(sk6,sk9)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_15]),c_0_9])])).

<A NAME="c_0_37"></A>cnf(c_0_37,plain,
    ( lt(X1,X2)
    | ~ ssItem(X2)
    | ~ ssItem(X1)
    | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( app(app(sk8,cons(sk6,sk9)),app(cons(sk7,nil),sk10)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_34]),c_0_10])])).

<A NAME="c_0_39"></A>cnf(c_0_39,plain,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_32])])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( lt(sk7,X1)
    | ~ lt(sk6,X1)
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_18]),c_0_11]),c_0_13])])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( lt(X1,X2)
    | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_32]),c_0_32])])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( app(app(sk8,cons(sk6,sk9)),cons(sk7,sk10)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_39]),c_0_14])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( lt(sk7,X1)
    | ~ lt(sk6,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_32])])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( lt(sk6,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_42]),c_0_43]),c_0_14]),c_0_10]),c_0_9]),c_0_20])])).

<A NAME="c_0_46"></A>cnf(c_0_46,plain,
    ( ~ lt(X1,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_32])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_45]),<A HREF="#c_0_46">c_0_46</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC160-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n114.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:36:23 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.29  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.02/0.29  # and selection function SelectCQArNTEqFirst.
% 0.02/0.29  #
% 0.02/0.29  # Preprocessing time       : 0.010 s
% 0.02/0.29  # Presaturation interreduction done
% 0.02/0.29  
% 0.02/0.29  # Proof found!
% 0.02/0.29  # SZS status Unsatisfiable
% 0.02/0.29  # SZS output start CNFRefutation
% 0.02/0.29  cnf(c_0_0, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
% 0.02/0.29  cnf(c_0_1, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
% 0.02/0.29  cnf(c_0_2, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause74)).
% 0.02/0.29  cnf(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.02/0.29  cnf(c_0_4, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.02/0.29  cnf(c_0_5, negated_conjecture, (app(app(app(app(sk8,cons(sk6,nil)),sk9),cons(sk7,nil)),sk10)=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_16)).
% 0.02/0.29  cnf(c_0_6, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.02/0.29  cnf(c_0_7, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
% 0.02/0.29  cnf(c_0_8, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
% 0.02/0.29  cnf(c_0_9, negated_conjecture, (ssList(sk8)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
% 0.02/0.29  cnf(c_0_10, negated_conjecture, (ssList(sk9)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_14)).
% 0.02/0.29  cnf(c_0_11, negated_conjecture, (ssItem(sk6)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_11)).
% 0.02/0.29  cnf(c_0_12, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause53)).
% 0.02/0.29  cnf(c_0_13, negated_conjecture, (ssItem(sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_12)).
% 0.02/0.29  cnf(c_0_14, negated_conjecture, (ssList(sk10)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_15)).
% 0.02/0.29  cnf(c_0_15, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.02/0.29  cnf(c_0_16, axiom, (lt(X2,X4)|app(app(X1,cons(X2,X3)),cons(X4,X5))!=X6|~ssList(X5)|~ssList(X3)|~ssList(X1)|~ssItem(X4)|~ssItem(X2)|~strictorderedP(X6)|~ssList(X6)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause181)).
% 0.02/0.29  cnf(c_0_17, axiom, (lt(X1,X3)|~leq(X1,X2)|~lt(X2,X3)|~ssItem(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause147)).
% 0.02/0.29  cnf(c_0_18, negated_conjecture, (leq(sk7,sk6)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_17)).
% 0.02/0.29  cnf(c_0_19, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.02/0.29  cnf(c_0_20, negated_conjecture, (strictorderedP(sk3)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
% 0.02/0.29  cnf(c_0_21, axiom, (~lt(X1,X1)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause63)).
% 0.02/0.29  cnf(c_0_22, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_0])).
% 0.02/0.29  cnf(c_0_23, plain, (ssItem(X1)|~ssItem(X2)|~ssList(app(app(X3,cons(X2,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)), inference(spm,[status(thm)],[c_0_22, c_0_1])).
% 0.02/0.29  cnf(c_0_24, plain, (ssItem(X1)|~ssItem(X2)|~ssList(app(cons(X2,X3),cons(X2,X4)))|~ssList(X4)|~ssList(X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_2]), c_0_3])]), c_0_4])).
% 0.02/0.29  cnf(c_0_25, negated_conjecture, (app(app(app(app(sk8,cons(sk6,nil)),sk9),cons(sk7,nil)),sk10)=sk3), inference(rw,[status(thm)],[c_0_5, c_0_6])).
% 0.02/0.29  cnf(c_0_26, plain, (ssItem(X1)|~ssItem(X2)|~ssList(cons(X2,cons(X2,X3)))|~ssList(X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_7]), c_0_3])]), c_0_4])).
% 0.02/0.29  cnf(c_0_27, negated_conjecture, (app(app(app(sk8,app(cons(sk6,nil),sk9)),cons(sk7,nil)),sk10)=sk3|~ssList(cons(sk6,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_8]), c_0_9]), c_0_10])])).
% 0.02/0.29  cnf(c_0_28, plain, (ssItem(X1)|~ssItem(X2)|~ssList(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_4]), c_0_4])).
% 0.02/0.29  cnf(c_0_29, negated_conjecture, (app(app(app(sk8,app(cons(sk6,nil),sk9)),cons(sk7,nil)),sk10)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_4]), c_0_11]), c_0_3])])).
% 0.02/0.29  cnf(c_0_30, plain, (ssItem(X1)|~ssItem(X2)), inference(spm,[status(thm)],[c_0_28, c_0_12])).
% 0.02/0.29  cnf(c_0_31, negated_conjecture, (app(app(app(sk8,cons(sk6,sk9)),cons(sk7,nil)),sk10)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_7]), c_0_11]), c_0_10])])).
% 0.02/0.29  cnf(c_0_32, negated_conjecture, (ssItem(X1)), inference(spm,[status(thm)],[c_0_30, c_0_13])).
% 0.02/0.29  cnf(c_0_33, negated_conjecture, (app(app(sk8,cons(sk6,sk9)),app(cons(sk7,nil),sk10))=sk3|~ssList(app(sk8,cons(sk6,sk9)))|~ssList(cons(sk7,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_31]), c_0_14])])).
% 0.02/0.29  cnf(c_0_34, plain, (ssList(cons(X1,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_32])])).
% 0.02/0.29  cnf(c_0_35, negated_conjecture, (app(app(sk8,cons(sk6,sk9)),app(cons(sk7,nil),sk10))=sk3|~ssList(app(sk8,cons(sk6,sk9)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_3])])).
% 0.02/0.29  cnf(c_0_36, negated_conjecture, (app(app(sk8,cons(sk6,sk9)),app(cons(sk7,nil),sk10))=sk3|~ssList(cons(sk6,sk9))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_15]), c_0_9])])).
% 0.02/0.29  cnf(c_0_37, plain, (lt(X1,X2)|~ssItem(X2)|~ssItem(X1)|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))), inference(er,[status(thm)],[c_0_16])).
% 0.02/0.29  cnf(c_0_38, negated_conjecture, (app(app(sk8,cons(sk6,sk9)),app(cons(sk7,nil),sk10))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_34]), c_0_10])])).
% 0.02/0.29  cnf(c_0_39, plain, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_32])])).
% 0.02/0.29  cnf(c_0_40, negated_conjecture, (lt(sk7,X1)|~lt(sk6,X1)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_11]), c_0_13])])).
% 0.02/0.29  cnf(c_0_41, plain, (lt(X1,X2)|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_32]), c_0_32])])).
% 0.02/0.29  cnf(c_0_42, negated_conjecture, (app(app(sk8,cons(sk6,sk9)),cons(sk7,sk10))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_14])])).
% 0.02/0.29  cnf(c_0_43, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_19, c_0_6])).
% 0.02/0.29  cnf(c_0_44, negated_conjecture, (lt(sk7,X1)|~lt(sk6,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_32])])).
% 0.02/0.29  cnf(c_0_45, negated_conjecture, (lt(sk6,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_14]), c_0_10]), c_0_9]), c_0_20])])).
% 0.02/0.29  cnf(c_0_46, plain, (~lt(X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_32])])).
% 0.02/0.29  cnf(c_0_47, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46]), ['proof']).
% 0.02/0.29  # SZS output end CNFRefutation
% 0.02/0.29  # Training examples: 0 positive, 0 negative
% 0.02/0.29  
% 0.02/0.29  # -------------------------------------------------
% 0.02/0.29  # User time                : 0.063 s
% 0.02/0.29  # System time              : 0.005 s
% 0.02/0.29  # Total time               : 0.068 s
% 0.02/0.29  # Maximum resident set size: 1868 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
