<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC309+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC309+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC309+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC309%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC309+1.p">SWC309+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n144.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:26 EDT 2016

% Result     : Theorem 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of formulae       :   33 (  78 expanded)
%              Number of clauses        :   23 (  31 expanded)
%              Number of leaves         :    5 (  18 expanded)
%              Depth                    :   11
%              Number of atoms          :  135 ( 732 expanded)
%              Number of equality atoms :   54 ( 334 expanded)
%              Maximal formula depth    :   21 (   5 average)
%              Maximal clause size      :   17 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ~ neq(X2,nil)
                    | ? [X5] :
                        ( ssItem(X5)
                        & ? [X6] :
                            ( ssList(X6)
                            & app(cons(X5,nil),X6) = X2
                            & app(X6,cons(X5,nil)) = X1 ) )
                    | ? [X7] :
                        ( ssItem(X7)
                        & ? [X8] :
                            ( ssList(X8)
                            & app(X8,cons(X7,nil)) != X3
                            & app(cons(X7,nil),X8) = X4 ) )
                    | ( nil != X3
                      & nil = X4 ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil = X1
        | ? [X2] :
            ( ssList(X2)
            & ? [X3] :
                ( ssItem(X3)
                & cons(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax20)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_5"></A>fof(c_0_5,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ~ neq(X2,nil)
                      | ? [X5] :
                          ( ssItem(X5)
                          & ? [X6] :
                              ( ssList(X6)
                              & app(cons(X5,nil),X6) = X2
                              & app(X6,cons(X5,nil)) = X1 ) )
                      | ? [X7] :
                          ( ssItem(X7)
                          & ? [X8] :
                              ( ssList(X8)
                              & app(X8,cons(X7,nil)) != X3
                              & app(cons(X7,nil),X8) = X4 ) )
                      | ( nil != X3
                        & nil = X4 ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_6"></A>fof(c_0_6,negated_conjecture,(
    ! [X13,X14,X15,X16] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & neq(esk49_0,nil)
      & ( ~ ssItem(X13)
        | ~ ssList(X14)
        | app(cons(X13,nil),X14) != esk49_0
        | app(X14,cons(X13,nil)) != esk48_0 )
      & ( ~ ssItem(X15)
        | ~ ssList(X16)
        | app(X16,cons(X15,nil)) = esk50_0
        | app(cons(X15,nil),X16) != esk51_0 )
      & ( nil = esk50_0
        | nil != esk51_0 ) ) ),
    inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( app(X1,cons(X2,nil)) != esk48_0
    | app(cons(X2,nil),X1) != esk49_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( app(X2,cons(X1,nil)) = esk50_0
    | app(cons(X1,nil),X2) != esk51_0
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( app(X1,cons(X2,nil)) != esk50_0
    | app(cons(X2,nil),X1) != esk51_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_8]),<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_12"></A>fof(c_0_12,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( app(cons(X1,nil),X2) != esk51_0
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_15"></A>fof(c_0_15,plain,(
    ! [X4] :
      ( ( ssList(esk44_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( ssItem(esk45_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( cons(esk45_1(X4),esk44_1(X4)) = X4
        | nil = X4
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( cons(X1,X2) != esk51_0
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,plain,
    ( nil = X1
    | cons(esk45_1(X1),esk44_1(X1)) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( esk51_0 = nil
    | ~ ssList(esk44_1(esk51_0))
    | ~ ssItem(esk45_1(esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_18])]),c_0_19])])).

<A NAME="c_0_21"></A>cnf(c_0_21,plain,
    ( nil = X1
    | ssList(esk44_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_22"></A>fof(c_0_22,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( neq(esk49_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( esk51_0 = nil
    | ~ ssItem(esk45_1(esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_21]),c_0_19])])).

<A NAME="c_0_25"></A>cnf(c_0_25,plain,
    ( nil = X1
    | ssItem(esk45_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,plain,
    ( ~ ssList(X1)
    | ~ ssList(X2)
    | X1 != X2
    | ~ neq(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( neq(esk51_0,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_25]),c_0_19])])).

<A NAME="c_0_29"></A>cnf(c_0_29,plain,
    ( ~ ssList(X1)
    | ~ neq(X1,X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( neq(nil,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_30]),c_0_31])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC309+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n144.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:40:23 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.26  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.02/0.26  # and selection function SelectCQPrecWNTNp.
% 0.02/0.26  #
% 0.02/0.26  # Preprocessing time       : 0.014 s
% 0.02/0.26  # Presaturation interreduction done
% 0.02/0.26  
% 0.02/0.26  # Proof found!
% 0.02/0.26  # SZS status Theorem
% 0.02/0.26  # SZS output start CNFRefutation
% 0.02/0.26  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~(neq(X2,nil)))|?[X5]:(ssItem(X5)&?[X6]:((ssList(X6)&app(cons(X5,nil),X6)=X2)&app(X6,cons(X5,nil))=X1)))|?[X7]:(ssItem(X7)&?[X8]:((ssList(X8)&app(X8,cons(X7,nil))!=X3)&app(cons(X7,nil),X8)=X4)))|(nil!=X3&nil=X4))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.02/0.26  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.02/0.26  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax20)).
% 0.02/0.26  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.02/0.26  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.02/0.26  fof(c_0_5, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~neq(X2,nil))|?[X5]:(ssItem(X5)&?[X6]:((ssList(X6)&app(cons(X5,nil),X6)=X2)&app(X6,cons(X5,nil))=X1)))|?[X7]:(ssItem(X7)&?[X8]:((ssList(X8)&app(X8,cons(X7,nil))!=X3)&app(cons(X7,nil),X8)=X4)))|(nil!=X3&nil=X4)))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.02/0.26  fof(c_0_6, negated_conjecture, (![X13]:![X14]:![X15]:![X16]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((((esk49_0=esk51_0&esk48_0=esk50_0)&neq(esk49_0,nil))&(~ssItem(X13)|((~ssList(X14)|app(cons(X13,nil),X14)!=esk49_0)|app(X14,cons(X13,nil))!=esk48_0)))&(~ssItem(X15)|((~ssList(X16)|app(X16,cons(X15,nil))=esk50_0)|app(cons(X15,nil),X16)!=esk51_0)))&(nil=esk50_0|nil!=esk51_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
% 0.02/0.26  cnf(c_0_7, negated_conjecture, (app(X1,cons(X2,nil))!=esk48_0|app(cons(X2,nil),X1)!=esk49_0|~ssList(X1)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
% 0.02/0.26  cnf(c_0_8, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_6])).
% 0.02/0.26  cnf(c_0_9, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_6])).
% 0.02/0.26  cnf(c_0_10, negated_conjecture, (app(X2,cons(X1,nil))=esk50_0|app(cons(X1,nil),X2)!=esk51_0|~ssList(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
% 0.02/0.26  cnf(c_0_11, negated_conjecture, (app(X1,cons(X2,nil))!=esk50_0|app(cons(X2,nil),X1)!=esk51_0|~ssList(X1)|~ssItem(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_8]), c_0_9])).
% 0.02/0.26  fof(c_0_12, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.02/0.26  cnf(c_0_13, negated_conjecture, (app(cons(X1,nil),X2)!=esk51_0|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[c_0_10, c_0_11])).
% 0.02/0.26  cnf(c_0_14, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.02/0.26  fof(c_0_15, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
% 0.02/0.26  cnf(c_0_16, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
% 0.02/0.26  cnf(c_0_17, negated_conjecture, (cons(X1,X2)!=esk51_0|~ssList(X2)|~ssItem(X1)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
% 0.02/0.26  cnf(c_0_18, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
% 0.02/0.26  cnf(c_0_19, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_16, c_0_9])).
% 0.02/0.26  cnf(c_0_20, negated_conjecture, (esk51_0=nil|~ssList(esk44_1(esk51_0))|~ssItem(esk45_1(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18])]), c_0_19])])).
% 0.02/0.26  cnf(c_0_21, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
% 0.02/0.26  fof(c_0_22, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
% 0.02/0.26  cnf(c_0_23, negated_conjecture, (neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
% 0.02/0.26  cnf(c_0_24, negated_conjecture, (esk51_0=nil|~ssItem(esk45_1(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_19])])).
% 0.02/0.26  cnf(c_0_25, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
% 0.02/0.26  cnf(c_0_26, plain, (~ssList(X1)|~ssList(X2)|X1!=X2|~neq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.02/0.26  cnf(c_0_27, negated_conjecture, (neq(esk51_0,nil)), inference(rw,[status(thm)],[c_0_23, c_0_9])).
% 0.02/0.26  cnf(c_0_28, negated_conjecture, (esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_19])])).
% 0.02/0.26  cnf(c_0_29, plain, (~ssList(X1)|~neq(X1,X1)), inference(er,[status(thm)],[c_0_26])).
% 0.02/0.26  cnf(c_0_30, negated_conjecture, (neq(nil,nil)), inference(rw,[status(thm)],[c_0_27, c_0_28])).
% 0.02/0.26  cnf(c_0_31, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.26  cnf(c_0_32, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])]), ['proof']).
% 0.02/0.26  # SZS output end CNFRefutation
% 0.02/0.26  # Training examples: 0 positive, 0 negative
% 0.02/0.26  
% 0.02/0.26  # -------------------------------------------------
% 0.02/0.26  # User time                : 0.029 s
% 0.02/0.26  # System time              : 0.002 s
% 0.02/0.26  # Total time               : 0.032 s
% 0.02/0.26  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
