<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC153+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC153+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC153+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC153%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC153+1.p">SWC153+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n117.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:51 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :   35 ( 634 expanded)
%              Number of clauses        :   30 ( 198 expanded)
%              Number of leaves         :    2 ( 147 expanded)
%              Depth                    :   11
%              Number of atoms          :  188 (15723 expanded)
%              Number of equality atoms :   26 (2071 expanded)
%              Maximal formula depth    :   31 (   5 average)
%              Maximal clause size      :   48 (   2 average)
%              Maximal term depth       :    6 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X4 != X2
                    | X3 != X1
                    | ? [X5] :
                        ( ssItem(X5)
                        & ? [X6] :
                            ( ssItem(X6)
                            & ? [X7] :
                                ( ssList(X7)
                                & ? [X8] :
                                    ( ssList(X8)
                                    & ? [X9] :
                                        ( ssList(X9)
                                        & app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9) = X3
                                        & leq(X6,X5)
                                        & ( ~ leq(X5,X6)
                                          | ? [X10] :
                                              ( ssItem(X10)
                                              & memberP(X8,X10)
                                              & ( ~ leq(X5,X10)
                                                | ~ leq(X10,X6) ) ) ) ) ) ) ) )
                    | ! [X11] :
                        ( ssItem(X11)
                       => ! [X12] :
                            ( ssItem(X12)
                           => ! [X13] :
                                ( ssList(X13)
                               => ! [X14] :
                                    ( ssList(X14)
                                   => ! [X15] :
                                        ( ssList(X15)
                                       => ( app(app(app(app(X13,cons(X11,nil)),X14),cons(X12,nil)),X15) != X1
                                          | ~ leq(X12,X11)
                                          | ( ! [X16] :
                                                ( ssItem(X16)
                                               => ( ~ memberP(X14,X16)
                                                  | ( leq(X16,X12)
                                                    & leq(X11,X16) ) ) )
                                            & leq(X11,X12) ) ) ) ) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( ( leq(X1,X2)
              & leq(X2,X1) )
           => X1 = X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax29)).

<A NAME="c_0_2"></A>fof(c_0_2,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X4 != X2
                      | X3 != X1
                      | ? [X5] :
                          ( ssItem(X5)
                          & ? [X6] :
                              ( ssItem(X6)
                              & ? [X7] :
                                  ( ssList(X7)
                                  & ? [X8] :
                                      ( ssList(X8)
                                      & ? [X9] :
                                          ( ssList(X9)
                                          & app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9) = X3
                                          & leq(X6,X5)
                                          & ( ~ leq(X5,X6)
                                            | ? [X10] :
                                                ( ssItem(X10)
                                                & memberP(X8,X10)
                                                & ( ~ leq(X5,X10)
                                                  | ~ leq(X10,X6) ) ) ) ) ) ) ) )
                      | ! [X11] :
                          ( ssItem(X11)
                         => ! [X12] :
                              ( ssItem(X12)
                             => ! [X13] :
                                  ( ssList(X13)
                                 => ! [X14] :
                                      ( ssList(X14)
                                     => ! [X15] :
                                          ( ssList(X15)
                                         => ( app(app(app(app(X13,cons(X11,nil)),X14),cons(X12,nil)),X15) != X1
                                            | ~ leq(X12,X11)
                                            | ( ! [X16] :
                                                  ( ssItem(X16)
                                                 => ( ~ memberP(X14,X16)
                                                    | ( leq(X16,X12)
                                                      & leq(X11,X16) ) ) )
                                              & leq(X11,X12) ) ) ) ) ) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_3"></A>fof(c_0_3,negated_conjecture,(
    ! [X21,X22,X23,X24,X25,X26] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk51_0 = esk49_0
      & esk50_0 = esk48_0
      & ( leq(X21,X22)
        | ~ ssList(X25)
        | app(app(app(app(X23,cons(X21,nil)),X24),cons(X22,nil)),X25) != esk50_0
        | ~ leq(X22,X21)
        | ~ ssList(X24)
        | ~ ssList(X23)
        | ~ ssItem(X22)
        | ~ ssItem(X21) )
      & ( leq(X21,X26)
        | ~ ssItem(X26)
        | ~ memberP(X24,X26)
        | ~ ssList(X25)
        | app(app(app(app(X23,cons(X21,nil)),X24),cons(X22,nil)),X25) != esk50_0
        | ~ leq(X22,X21)
        | ~ ssList(X24)
        | ~ ssList(X23)
        | ~ ssItem(X22)
        | ~ ssItem(X21) )
      & ( leq(X26,X22)
        | ~ ssItem(X26)
        | ~ memberP(X24,X26)
        | ~ ssList(X25)
        | app(app(app(app(X23,cons(X21,nil)),X24),cons(X22,nil)),X25) != esk50_0
        | ~ leq(X22,X21)
        | ~ ssList(X24)
        | ~ ssList(X23)
        | ~ ssItem(X22)
        | ~ ssItem(X21) )
      & ssItem(esk52_0)
      & ssItem(esk53_0)
      & ssList(esk54_0)
      & ssList(esk55_0)
      & ssList(esk56_0)
      & app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk48_0
      & leq(esk53_0,esk52_0)
      & ( ssItem(esk57_0)
        | ~ leq(esk52_0,esk53_0) )
      & ( memberP(esk55_0,esk57_0)
        | ~ leq(esk52_0,esk53_0) )
      & ( ~ leq(esk57_0,esk53_0)
        | ~ leq(esk52_0,esk57_0)
        | ~ leq(esk52_0,esk53_0) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).

<A NAME="c_0_4"></A>fof(c_0_4,plain,(
    ! [X3,X4] :
      ( ~ ssItem(X3)
      | ~ ssItem(X4)
      | ~ leq(X3,X4)
      | ~ leq(X4,X3)
      | X3 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( esk50_0 = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_7"></A>cnf(c_0_7,plain,
    ( X1 = X2
    | ~ leq(X2,X1)
    | ~ leq(X1,X2)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( leq(esk53_0,esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ssItem(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( ssItem(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( leq(X1,X2)
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ ssList(X4)
    | ~ leq(X2,X1)
    | app(app(app(app(X3,cons(X1,nil)),X4),cons(X2,nil)),X5) != esk50_0
    | ~ ssList(X5) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( ssList(esk56_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( ssList(esk55_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( ssList(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( esk53_0 = esk52_0
    | ~ leq(esk52_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_8]),c_0_9]),c_0_10])])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( leq(esk52_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_12]),c_0_8]),c_0_13]),c_0_14]),c_0_15]),c_0_9]),c_0_10])])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( esk53_0 = esk52_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_17])])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( ~ leq(esk52_0,esk53_0)
    | ~ leq(esk52_0,esk57_0)
    | ~ leq(esk57_0,esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( leq(X6,X2)
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ ssList(X4)
    | ~ leq(X2,X1)
    | app(app(app(app(X3,cons(X1,nil)),X4),cons(X2,nil)),X5) != esk50_0
    | ~ ssList(X5)
    | ~ memberP(X4,X6)
    | ~ ssItem(X6) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk52_0,nil)),esk56_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( leq(esk52_0,esk52_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( memberP(esk55_0,esk57_0)
    | ~ leq(esk52_0,esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( ssItem(esk57_0)
    | ~ leq(esk52_0,esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( ~ leq(esk52_0,esk57_0)
    | ~ leq(esk57_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_17])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( leq(X1,esk52_0)
    | ~ memberP(esk55_0,X1)
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_21]),c_0_22]),c_0_13]),c_0_14]),c_0_15]),c_0_10])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( memberP(esk55_0,esk57_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_17])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( ssItem(esk57_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_17])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( leq(X1,X6)
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ ssList(X4)
    | ~ leq(X2,X1)
    | app(app(app(app(X3,cons(X1,nil)),X4),cons(X2,nil)),X5) != esk50_0
    | ~ ssList(X5)
    | ~ memberP(X4,X6)
    | ~ ssItem(X6) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( ~ leq(esk52_0,esk57_0)
    | ~ leq(esk57_0,esk52_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( leq(esk57_0,esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_27]),c_0_28])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( leq(esk52_0,X1)
    | ~ memberP(esk55_0,X1)
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_21]),c_0_22]),c_0_13]),c_0_14]),c_0_15]),c_0_10])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( ~ leq(esk52_0,esk57_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_27]),c_0_28])]),<A HREF="#c_0_33">c_0_33</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC153+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.22  % Computer   : n117.star.cs.uiowa.edu
% 0.03/0.22  % Model      : x86_64 x86_64
% 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.22  % Memory     : 32218.75MB
% 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.22  % CPULimit   : 300
% 0.03/0.22  % DateTime   : Tue May 10 20:36:07 CDT 2016
% 0.03/0.22  % CPUTime    : 
% 0.03/0.22  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.32  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.32  # and selection function SelectCQPrecWNTNp.
% 0.06/0.32  #
% 0.06/0.32  # Preprocessing time       : 0.014 s
% 0.06/0.32  # Presaturation interreduction done
% 0.06/0.32  
% 0.06/0.32  # Proof found!
% 0.06/0.32  # SZS status Theorem
% 0.06/0.32  # SZS output start CNFRefutation
% 0.06/0.32  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X4!=X2|X3!=X1)|?[X5]:(ssItem(X5)&?[X6]:(ssItem(X6)&?[X7]:(ssList(X7)&?[X8]:(ssList(X8)&?[X9]:(((ssList(X9)&app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9)=X3)&leq(X6,X5))&(~(leq(X5,X6))|?[X10]:((ssItem(X10)&memberP(X8,X10))&(~(leq(X5,X10))|~(leq(X10,X6)))))))))))|![X11]:(ssItem(X11)=>![X12]:(ssItem(X12)=>![X13]:(ssList(X13)=>![X14]:(ssList(X14)=>![X15]:(ssList(X15)=>((app(app(app(app(X13,cons(X11,nil)),X14),cons(X12,nil)),X15)!=X1|~(leq(X12,X11)))|(![X16]:(ssItem(X16)=>(~(memberP(X14,X16))|(leq(X16,X12)&leq(X11,X16))))&leq(X11,X12)))))))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.06/0.32  fof(c_0_1, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>((leq(X1,X2)&leq(X2,X1))=>X1=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax29)).
% 0.06/0.32  fof(c_0_2, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X4!=X2|X3!=X1)|?[X5]:(ssItem(X5)&?[X6]:(ssItem(X6)&?[X7]:(ssList(X7)&?[X8]:(ssList(X8)&?[X9]:(((ssList(X9)&app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9)=X3)&leq(X6,X5))&(~leq(X5,X6)|?[X10]:((ssItem(X10)&memberP(X8,X10))&(~leq(X5,X10)|~leq(X10,X6))))))))))|![X11]:(ssItem(X11)=>![X12]:(ssItem(X12)=>![X13]:(ssList(X13)=>![X14]:(ssList(X14)=>![X15]:(ssList(X15)=>((app(app(app(app(X13,cons(X11,nil)),X14),cons(X12,nil)),X15)!=X1|~leq(X12,X11))|(![X16]:(ssItem(X16)=>(~memberP(X14,X16)|(leq(X16,X12)&leq(X11,X16))))&leq(X11,X12))))))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.32  fof(c_0_3, negated_conjecture, 
% (![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((esk51_0=esk49_0&esk50_0=esk48_0)&((((((leq(X21,X22)|((~ssList(X25)|app(app(app(app(X23,cons(X21,nil)),X24),cons(X22,nil)),X25)!=esk50_0)|~leq(X22,X21)))|~ssList(X24))|~ssList(X23))|~ssItem(X22))|~ssItem(X21))&(((((((leq(X21,X26)|(~ssItem(X26)|~memberP(X24,X26)))|((~ssList(X25)|app(app(app(app(X23,cons(X21,nil)),X24),cons(X22,nil)),X25)!=esk50_0)|~leq(X22,X21)))|~ssList(X24))|~ssList(X23))|~ssItem(X22))|~ssItem(X21))&((((((leq(X26,X22)|(~ssItem(X26)|~memberP(X24,X26)))|((~ssList(X25)|app(app(app(app(X23,cons(X21,nil)),X24),cons(X22,nil)),X25)!=esk50_0)|~leq(X22,X21)))|~ssList(X24))|~ssList(X23))|~ssItem(X22))|~ssItem(X21)))))&(ssItem(esk52_0)&(ssItem(esk53_0)&(ssList(esk54_0)&(ssList(esk55_0)&(ssList(esk56_0)&((app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk48_0&leq(esk53_0,esk52_0))&((ssItem(esk57_0)|~leq(esk52_0,esk53_0))&((memberP(esk55_0,esk57_0)|~leq(esk52_0
% ,esk53_0))&((~leq(esk57_0,esk53_0)|~leq(esk52_0,esk57_0))|~leq(esk52_0,esk53_0)))))))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])])).
% 0.06/0.32  fof(c_0_4, plain, (![X3]:![X4]:(~ssItem(X3)|(~ssItem(X4)|((~leq(X3,X4)|~leq(X4,X3))|X3=X4)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.06/0.32  cnf(c_0_5, negated_conjecture, (app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_6, negated_conjecture, (esk50_0=esk48_0), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_7, plain, (X1=X2|~leq(X2,X1)|~leq(X1,X2)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.06/0.32  cnf(c_0_8, negated_conjecture, (leq(esk53_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_9, negated_conjecture, (ssItem(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_10, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_11, negated_conjecture, (leq(X1,X2)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~ssList(X4)|~leq(X2,X1)|app(app(app(app(X3,cons(X1,nil)),X4),cons(X2,nil)),X5)!=esk50_0|~ssList(X5)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_12, negated_conjecture, (app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk50_0), inference(rw,[status(thm)],[c_0_5, c_0_6])).
% 0.06/0.32  cnf(c_0_13, negated_conjecture, (ssList(esk56_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_14, negated_conjecture, (ssList(esk55_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_15, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_16, negated_conjecture, (esk53_0=esk52_0|~leq(esk52_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9]), c_0_10])])).
% 0.06/0.32  cnf(c_0_17, negated_conjecture, (leq(esk52_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_8]), c_0_13]), c_0_14]), c_0_15]), c_0_9]), c_0_10])])).
% 0.06/0.32  cnf(c_0_18, negated_conjecture, (esk53_0=esk52_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17])])).
% 0.06/0.32  cnf(c_0_19, negated_conjecture, (~leq(esk52_0,esk53_0)|~leq(esk52_0,esk57_0)|~leq(esk57_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_20, negated_conjecture, (leq(X6,X2)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~ssList(X4)|~leq(X2,X1)|app(app(app(app(X3,cons(X1,nil)),X4),cons(X2,nil)),X5)!=esk50_0|~ssList(X5)|~memberP(X4,X6)|~ssItem(X6)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_21, negated_conjecture, (app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk52_0,nil)),esk56_0)=esk50_0), inference(rw,[status(thm)],[c_0_12, c_0_18])).
% 0.06/0.32  cnf(c_0_22, negated_conjecture, (leq(esk52_0,esk52_0)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
% 0.06/0.32  cnf(c_0_23, negated_conjecture, (memberP(esk55_0,esk57_0)|~leq(esk52_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_24, negated_conjecture, (ssItem(esk57_0)|~leq(esk52_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_25, negated_conjecture, (~leq(esk52_0,esk57_0)|~leq(esk57_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_17])])).
% 0.06/0.32  cnf(c_0_26, negated_conjecture, (leq(X1,esk52_0)|~memberP(esk55_0,X1)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_13]), c_0_14]), c_0_15]), c_0_10])])).
% 0.06/0.32  cnf(c_0_27, negated_conjecture, (memberP(esk55_0,esk57_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_17])])).
% 0.06/0.32  cnf(c_0_28, negated_conjecture, (ssItem(esk57_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_17])])).
% 0.06/0.32  cnf(c_0_29, negated_conjecture, (leq(X1,X6)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~ssList(X4)|~leq(X2,X1)|app(app(app(app(X3,cons(X1,nil)),X4),cons(X2,nil)),X5)!=esk50_0|~ssList(X5)|~memberP(X4,X6)|~ssItem(X6)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.32  cnf(c_0_30, negated_conjecture, (~leq(esk52_0,esk57_0)|~leq(esk57_0,esk52_0)), inference(rw,[status(thm)],[c_0_25, c_0_18])).
% 0.06/0.32  cnf(c_0_31, negated_conjecture, (leq(esk57_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
% 0.06/0.32  cnf(c_0_32, negated_conjecture, (leq(esk52_0,X1)|~memberP(esk55_0,X1)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_21]), c_0_22]), c_0_13]), c_0_14]), c_0_15]), c_0_10])])).
% 0.06/0.32  cnf(c_0_33, negated_conjecture, (~leq(esk52_0,esk57_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31])])).
% 0.06/0.32  cnf(c_0_34, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_27]), c_0_28])]), c_0_33]), ['proof']).
% 0.06/0.32  # SZS output end CNFRefutation
% 0.06/0.32  # Training examples: 0 positive, 0 negative
% 0.06/0.32  
% 0.06/0.32  # -------------------------------------------------
% 0.06/0.32  # User time                : 0.093 s
% 0.06/0.32  # System time              : 0.003 s
% 0.06/0.32  # Total time               : 0.096 s
% 0.06/0.32  # Maximum resident set size: 1896 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
