<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC410+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC410+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC410+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC410%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC410+1.p">SWC410+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n142.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:50 EDT 2016

% Result     : Theorem 2.72s
% Output     : CNFRefutation 2.72s
% Verified   : 
% Statistics : Number of formulae       :  180 (14722 expanded)
%              Number of clauses        :  127 (5929 expanded)
%              Number of leaves         :   26 (3737 expanded)
%              Depth                    :   31
%              Number of atoms          :  652 (157149 expanded)
%              Number of equality atoms :  202 (57064 expanded)
%              Maximal formula depth    :   30 (   5 average)
%              Maximal clause size      :   25 (   3 average)
%              Maximal term depth       :    6 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ! [X5] :
                        ( ssItem(X5)
                       => ( ~ memberP(X2,X5)
                          | memberP(X1,X5) ) )
                    | ( nil != X3
                      & nil = X4 )
                    | ( neq(X4,nil)
                      & ( ~ neq(X3,nil)
                        | ? [X6] :
                            ( ssList(X6)
                            & X4 != X6
                            & ? [X7] :
                                ( ssList(X7)
                                & ? [X8] :
                                    ( ssList(X8)
                                    & tl(X3) = X7
                                    & app(X7,X8) = X6
                                    & ? [X9] :
                                        ( ssItem(X9)
                                        & cons(X9,nil) = X8
                                        & hd(X3) = X9
                                        & neq(nil,X3) )
                                    & neq(nil,X3) ) ) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssList(tl(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssItem(hd(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax22)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => cons(hd(X1),tl(X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax78)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax86)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => hd(app(X1,X2)) = hd(X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax85)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => nil != cons(X2,X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax21)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( memberP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(X3,cons(X2,X4)) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax3)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => hd(cons(X2,X1)) = X2 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax23)).

<A NAME="c_0_17"></A>fof(c_0_17,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_18"></A>fof(c_0_18,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( rearsegP(X1,X2)
               => rearsegP(app(X3,X1),X2) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax50)).

<A NAME="c_0_19"></A>fof(c_0_19,axiom,(
    ! [X1] :
      ( ssList(X1)
     => rearsegP(X1,nil) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax51)).

<A NAME="c_0_20"></A>fof(c_0_20,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(cons(X2,X3),X1)
              &lt;=> ( X1 = X2
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax37)).

<A NAME="c_0_21"></A>fof(c_0_21,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( rearsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax6)).

<A NAME="c_0_22"></A>fof(c_0_22,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_23"></A>fof(c_0_23,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(app(X2,X3),X1)
              &lt;=> ( memberP(X2,X1)
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax36)).

<A NAME="c_0_24"></A>fof(c_0_24,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax38)).

<A NAME="c_0_25"></A>fof(c_0_25,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( rearsegP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax52)).

<A NAME="c_0_26"></A>fof(c_0_26,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ! [X5] :
                          ( ssItem(X5)
                         => ( ~ memberP(X2,X5)
                            | memberP(X1,X5) ) )
                      | ( nil != X3
                        & nil = X4 )
                      | ( neq(X4,nil)
                        & ( ~ neq(X3,nil)
                          | ? [X6] :
                              ( ssList(X6)
                              & X4 != X6
                              & ? [X7] :
                                  ( ssList(X7)
                                  & ? [X8] :
                                      ( ssList(X8)
                                      & tl(X3) = X7
                                      & app(X7,X8) = X6
                                      & ? [X9] :
                                          ( ssItem(X9)
                                          & cons(X9,nil) = X8
                                          & hd(X3) = X9
                                          & neq(nil,X3) )
                                      & neq(nil,X3) ) ) ) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_27"></A>fof(c_0_27,negated_conjecture,(
    ! [X15,X16,X17,X18] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ssItem(esk52_0)
      & memberP(esk49_0,esk52_0)
      & ~ memberP(esk48_0,esk52_0)
      & ( nil = esk50_0
        | nil != esk51_0 )
      & ( neq(esk50_0,nil)
        | ~ neq(esk51_0,nil) )
      & ( ~ ssList(X15)
        | esk51_0 = X15
        | ~ ssList(X16)
        | ~ ssList(X17)
        | tl(esk50_0) != X16
        | app(X16,X17) != X15
        | ~ ssItem(X18)
        | cons(X18,nil) != X17
        | hd(esk50_0) != X18
        | ~ neq(nil,esk50_0)
        | ~ neq(nil,esk50_0)
        | ~ neq(esk51_0,nil) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])])])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( esk51_0 = X4
    | ~ neq(esk51_0,nil)
    | ~ neq(nil,esk50_0)
    | ~ neq(nil,esk50_0)
    | hd(esk50_0) != X1
    | cons(X1,nil) != X2
    | ~ ssItem(X1)
    | app(X3,X2) != X4
    | tl(esk50_0) != X3
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X4) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_29"></A>fof(c_0_29,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( esk51_0 = X4
    | hd(esk50_0) != X1
    | tl(esk50_0) != X3
    | app(X3,X2) != X4
    | cons(X1,nil) != X2
    | ~ ssItem(X1)
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ neq(nil,esk50_0)
    | ~ neq(esk51_0,nil) ),
    inference(cn,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_32"></A>fof(c_0_32,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssList(tl(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( app(tl(esk50_0),cons(hd(esk50_0),nil)) = esk51_0
    | ~ ssList(cons(hd(esk50_0),nil))
    | ~ ssList(tl(esk50_0))
    | ~ neq(nil,esk50_0)
    | ~ neq(esk51_0,nil)
    | ~ ssItem(hd(esk50_0)) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_30])])])]),<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_38"></A>fof(c_0_38,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( app(tl(esk50_0),cons(hd(esk50_0),nil)) = esk51_0
    | esk50_0 = nil
    | ~ ssList(cons(hd(esk50_0),nil))
    | ~ neq(nil,esk50_0)
    | ~ neq(esk51_0,nil)
    | ~ ssItem(hd(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_36]),c_0_37])])).

<A NAME="c_0_40"></A>cnf(c_0_40,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_42"></A>fof(c_0_42,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssItem(hd(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( app(tl(esk50_0),cons(hd(esk50_0),nil)) = esk51_0
    | esk50_0 = nil
    | ~ neq(nil,esk50_0)
    | ~ neq(esk51_0,nil)
    | ~ ssItem(hd(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_40]),c_0_41])])).

<A NAME="c_0_44"></A>cnf(c_0_44,plain,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_45"></A>fof(c_0_45,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).

<A NAME="c_0_46"></A>fof(c_0_46,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | cons(hd(X2),tl(X2)) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

<A NAME="c_0_47"></A>fof(c_0_47,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

<A NAME="c_0_48"></A>fof(c_0_48,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( app(tl(esk50_0),cons(hd(esk50_0),nil)) = esk51_0
    | esk50_0 = nil
    | ~ neq(nil,esk50_0)
    | ~ neq(esk51_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_44]),c_0_37])])).

<A NAME="c_0_50"></A>cnf(c_0_50,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_54"></A>cnf(c_0_54,plain,
    ( tl(app(X1,X2)) = app(tl(X1),X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_55"></A>cnf(c_0_55,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( app(tl(esk50_0),cons(hd(esk50_0),nil)) = esk51_0
    | esk50_0 = nil
    | ~ neq(esk51_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_50]),c_0_37]),c_0_41])])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,<A HREF="#c_0_52">c_0_52</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( nil = esk50_0
    | nil != esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,plain,
    ( cons(hd(app(X1,X2)),app(tl(X1),X2)) = app(X1,X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_54]),c_0_31]),<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( app(tl(esk50_0),cons(hd(esk50_0),nil)) = esk51_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_50]),c_0_41]),c_0_57])]),<A HREF="#c_0_58">c_0_58</A>])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( cons(hd(app(esk50_0,cons(hd(esk50_0),nil))),esk51_0) = app(esk50_0,cons(hd(esk50_0),nil))
    | esk50_0 = nil
    | ~ ssList(cons(hd(esk50_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,c_0_60]),c_0_37])])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( cons(hd(app(esk50_0,cons(hd(esk50_0),nil))),esk51_0) = app(esk50_0,cons(hd(esk50_0),nil))
    | esk50_0 = nil
    | ~ ssItem(hd(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_40]),c_0_41])])).

<A NAME="c_0_63"></A>fof(c_0_63,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | hd(app(X3,X4)) = hd(X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( cons(hd(app(esk50_0,cons(hd(esk50_0),nil))),esk51_0) = app(esk50_0,cons(hd(esk50_0),nil))
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_44]),c_0_37])])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( hd(app(X1,X2)) = hd(X1)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_66"></A>fof(c_0_66,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | nil != cons(X4,X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( app(esk50_0,cons(hd(esk50_0),nil)) = cons(hd(esk50_0),esk51_0)
    | esk50_0 = nil
    | ~ ssList(cons(hd(esk50_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,c_0_65]),c_0_37])])).

<A NAME="c_0_68"></A>cnf(c_0_68,plain,
    ( nil != cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>])).

<A NAME="c_0_69"></A>fof(c_0_69,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).

<A NAME="c_0_70"></A>fof(c_0_70,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).

<A NAME="c_0_71"></A>fof(c_0_71,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk3_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ssList(esk4_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6))) = X5
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(X9,cons(X6,X10)) != X5
        | memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])])).

<A NAME="c_0_72"></A>cnf(c_0_72,negated_conjecture,
    ( memberP(esk49_0,esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,negated_conjecture,
    ( app(esk50_0,cons(hd(esk50_0),nil)) = cons(hd(esk50_0),esk51_0)
    | esk50_0 = nil
    | ~ ssItem(hd(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,c_0_40]),c_0_41])])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( esk50_0 = nil
    | app(esk50_0,cons(hd(esk50_0),nil)) != nil
    | ~ ssItem(hd(app(esk50_0,cons(hd(esk50_0),nil)))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_64]),c_0_57])])).

<A NAME="c_0_75"></A>cnf(c_0_75,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,plain,
    ( app(esk3_2(X1,X2),cons(X2,esk4_2(X1,X2))) = X1
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ memberP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>])).

<A NAME="c_0_78"></A>cnf(c_0_78,negated_conjecture,
    ( memberP(esk51_0,esk52_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,<A HREF="#c_0_52">c_0_52</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( ssItem(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_80"></A>cnf(c_0_80,plain,
    ( ssList(esk3_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ memberP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>])).

<A NAME="c_0_81"></A>cnf(c_0_81,negated_conjecture,
    ( app(esk50_0,cons(hd(esk50_0),nil)) = cons(hd(esk50_0),esk51_0)
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>,c_0_44]),c_0_37])])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( esk50_0 = nil
    | app(esk50_0,cons(hd(esk50_0),nil)) != nil
    | ~ ssList(cons(hd(esk50_0),nil))
    | ~ ssItem(hd(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,c_0_65]),c_0_37])])).

<A NAME="c_0_83"></A>fof(c_0_83,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssItem(X6)
      | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).

<A NAME="c_0_84"></A>cnf(c_0_84,plain,
    ( app(nil,app(X1,X2)) = app(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_76]),c_0_41])])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( app(esk3_2(esk51_0,esk52_0),cons(esk52_0,esk4_2(esk51_0,esk52_0))) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_78]),c_0_57]),c_0_79])])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( ssList(esk3_2(esk51_0,esk52_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,c_0_78]),c_0_57]),c_0_79])])).

<A NAME="c_0_87"></A>cnf(c_0_87,plain,
    ( ssList(esk4_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ memberP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( hd(cons(hd(esk50_0),esk51_0)) = hd(esk50_0)
    | esk50_0 = nil
    | ~ ssList(cons(hd(esk50_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_81]),c_0_37])])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( esk50_0 = nil
    | ssList(cons(hd(esk50_0),esk51_0))
    | ~ ssList(cons(hd(esk50_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_81]),c_0_37])])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( esk50_0 = nil
    | app(esk50_0,cons(hd(esk50_0),nil)) != nil
    | ~ ssItem(hd(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,c_0_40]),c_0_41])])).

<A NAME="c_0_91"></A>fof(c_0_91,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | hd(cons(X4,X3)) = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).

<A NAME="c_0_92"></A>fof(c_0_92,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).

<A NAME="c_0_93"></A>cnf(c_0_93,plain,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>])).

<A NAME="c_0_94"></A>cnf(c_0_94,negated_conjecture,
    ( app(nil,esk51_0) = esk51_0
    | ~ ssList(cons(esk52_0,esk4_2(esk51_0,esk52_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,c_0_85]),c_0_86])])).

<A NAME="c_0_95"></A>cnf(c_0_95,negated_conjecture,
    ( ssList(esk4_2(esk51_0,esk52_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,c_0_78]),c_0_57]),c_0_79])])).

<A NAME="c_0_96"></A>cnf(c_0_96,negated_conjecture,
    ( hd(cons(hd(esk50_0),esk51_0)) = hd(esk50_0)
    | esk50_0 = nil
    | ~ ssItem(hd(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>,c_0_40]),c_0_41])])).

<A NAME="c_0_97"></A>cnf(c_0_97,negated_conjecture,
    ( esk50_0 = nil
    | ssList(cons(hd(esk50_0),esk51_0))
    | ~ ssItem(hd(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>,c_0_40]),c_0_41])])).

<A NAME="c_0_98"></A>cnf(c_0_98,negated_conjecture,
    ( esk50_0 = nil
    | app(esk50_0,cons(hd(esk50_0),nil)) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,c_0_44]),c_0_37])])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>])).

<A NAME="c_0_101"></A>cnf(c_0_101,plain,
    ( hd(cons(X1,app(X2,X3))) = hd(cons(X1,X2))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_93]),c_0_40]),<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_102"></A>cnf(c_0_102,negated_conjecture,
    ( app(nil,esk51_0) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>,c_0_40]),c_0_95]),c_0_79])])).

<A NAME="c_0_103"></A>cnf(c_0_103,negated_conjecture,
    ( hd(cons(hd(esk50_0),esk51_0)) = hd(esk50_0)
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,c_0_44]),c_0_37])])).

<A NAME="c_0_104"></A>cnf(c_0_104,negated_conjecture,
    ( esk50_0 = nil
    | ssList(cons(hd(esk50_0),esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_44]),c_0_37])])).

<A NAME="c_0_105"></A>cnf(c_0_105,negated_conjecture,
    ( esk50_0 = nil
    | cons(hd(esk50_0),esk51_0) != nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_98">c_0_98</A>,<A HREF="#c_0_81">c_0_81</A>])).

<A NAME="c_0_106"></A>fof(c_0_106,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ rearsegP(X4,X5)
      | rearsegP(app(X6,X4),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).

<A NAME="c_0_107"></A>fof(c_0_107,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | rearsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).

<A NAME="c_0_108"></A>fof(c_0_108,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(cons(X5,X6),X4)
        | X4 = X5
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( X4 != X5
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])).

<A NAME="c_0_109"></A>cnf(c_0_109,plain,
    ( cons(X1,tl(cons(X1,X2))) = cons(X1,X2)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_99]),c_0_40]),<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_110"></A>cnf(c_0_110,plain,
    ( cons(hd(cons(X1,X2)),X2) = cons(X1,X2)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_100]),c_0_40]),<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_111"></A>cnf(c_0_111,negated_conjecture,
    ( hd(cons(X1,esk51_0)) = hd(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_101">c_0_101</A>,c_0_102]),c_0_57]),c_0_41])])).

<A NAME="c_0_112"></A>cnf(c_0_112,negated_conjecture,
    ( esk50_0 = nil
    | ssItem(hd(esk50_0)) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_103]),c_0_104]),<A HREF="#c_0_105">c_0_105</A>])).

<A NAME="c_0_113"></A>cnf(c_0_113,plain,
    ( rearsegP(app(X1,X2),X3)
    | ~ rearsegP(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>])).

<A NAME="c_0_114"></A>cnf(c_0_114,plain,
    ( rearsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,plain,
    ( memberP(cons(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ memberP(X3,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>])).

<A NAME="c_0_116"></A>cnf(c_0_116,plain,
    ( nil = app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != X1
    | nil != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_117"></A>cnf(c_0_117,plain,
    ( cons(hd(cons(X1,X2)),tl(cons(X1,X2))) = cons(X1,X2)
    | ~ ssList(X2)
    | ~ ssItem(hd(cons(X1,X2)))
    | ~ ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>,<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_118"></A>cnf(c_0_118,negated_conjecture,
    ( hd(cons(hd(esk50_0),nil)) = hd(esk50_0)
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_103">c_0_103</A>,c_0_111]),<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_119"></A>fof(c_0_119,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk7_2(X4,X5))
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(esk7_2(X4,X5),X5) = X4
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X7,X5) != X4
        | rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).

<A NAME="c_0_120"></A>cnf(c_0_120,plain,
    ( rearsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_113">c_0_113</A>,c_0_114]),c_0_41])])).

<A NAME="c_0_121"></A>cnf(c_0_121,negated_conjecture,
    ( memberP(cons(X1,esk51_0),esk52_0)
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_115">c_0_115</A>,c_0_78]),c_0_57]),c_0_79])])).

<A NAME="c_0_122"></A>cnf(c_0_122,plain,
    ( ssList(cons(X1,app(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_93]),<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_123"></A>cnf(c_0_123,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_116])]),c_0_41])])).

<A NAME="c_0_124"></A>cnf(c_0_124,negated_conjecture,
    ( cons(hd(esk50_0),tl(cons(hd(esk50_0),nil))) = cons(hd(esk50_0),nil)
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_117">c_0_117</A>,c_0_118]),c_0_41])]),<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_125"></A>fof(c_0_125,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])).

<A NAME="c_0_126"></A>cnf(c_0_126,plain,
    ( app(esk7_2(X1,X2),X2) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>])).

<A NAME="c_0_127"></A>cnf(c_0_127,negated_conjecture,
    ( rearsegP(esk51_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_120">c_0_120</A>,c_0_102]),c_0_57]),c_0_41])])).

<A NAME="c_0_128"></A>cnf(c_0_128,plain,
    ( ssList(esk7_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>])).

<A NAME="c_0_129"></A>fof(c_0_129,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(app(X5,X6),X4)
        | memberP(X5,X4)
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X5,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).

<A NAME="c_0_130"></A>cnf(c_0_130,negated_conjecture,
    ( esk50_0 = nil
    | memberP(app(esk50_0,cons(hd(esk50_0),nil)),esk52_0)
    | ~ ssItem(hd(app(esk50_0,cons(hd(esk50_0),nil)))) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_121">c_0_121</A>,<A HREF="#c_0_64">c_0_64</A>])).

<A NAME="c_0_131"></A>cnf(c_0_131,plain,
    ( ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_122">c_0_122</A>,c_0_123]),c_0_41])])).

<A NAME="c_0_132"></A>cnf(c_0_132,negated_conjecture,
    ( ~ memberP(esk48_0,esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_133"></A>cnf(c_0_133,negated_conjecture,
    ( esk50_0 = nil
    | ssList(cons(hd(esk50_0),nil))
    | ~ ssList(tl(cons(hd(esk50_0),nil))) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_124]),<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_134"></A>cnf(c_0_134,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_125">c_0_125</A>])).

<A NAME="c_0_135"></A>cnf(c_0_135,negated_conjecture,
    ( app(esk7_2(esk51_0,nil),nil) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_126">c_0_126</A>,c_0_127]),c_0_41]),c_0_57])])).

<A NAME="c_0_136"></A>cnf(c_0_136,negated_conjecture,
    ( ssList(esk7_2(esk51_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_128">c_0_128</A>,c_0_127]),c_0_41]),c_0_57])])).

<A NAME="c_0_137"></A>cnf(c_0_137,plain,
    ( memberP(cons(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | X1 != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>])).

<A NAME="c_0_138"></A>cnf(c_0_138,plain,
    ( memberP(X3,X1)
    | memberP(X2,X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(app(X2,X3),X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_129">c_0_129</A>])).

<A NAME="c_0_139"></A>cnf(c_0_139,negated_conjecture,
    ( esk50_0 = nil
    | memberP(app(esk50_0,cons(hd(esk50_0),nil)),esk52_0) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_130">c_0_130</A>,c_0_65]),c_0_37])]),c_0_131]),<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_140"></A>cnf(c_0_140,negated_conjecture,
    ( ~ memberP(esk50_0,esk52_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_132">c_0_132</A>,<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_141"></A>cnf(c_0_141,negated_conjecture,
    ( esk50_0 = nil
    | ssList(cons(hd(esk50_0),nil)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_133">c_0_133</A>,c_0_100]),c_0_41])]),<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_142"></A>fof(c_0_142,plain,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_143"></A>cnf(c_0_143,plain,
    ( memberP(app(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(X3,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_129">c_0_129</A>])).

<A NAME="c_0_144"></A>cnf(c_0_144,plain,
    ( ssList(app(X1,app(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_75]),<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_145"></A>cnf(c_0_145,negated_conjecture,
    ( esk7_2(esk51_0,nil) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_134">c_0_134</A>,c_0_135]),c_0_136])])).

<A NAME="c_0_146"></A>cnf(c_0_146,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_137">c_0_137</A>])).

<A NAME="c_0_147"></A>cnf(c_0_147,plain,
    ( memberP(X3,X1)
    | X1 = X2
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ memberP(cons(X2,X3),X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>])).

<A NAME="c_0_148"></A>cnf(c_0_148,negated_conjecture,
    ( esk50_0 = nil
    | memberP(cons(hd(esk50_0),nil),esk52_0) ),
    inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_138">c_0_138</A>,c_0_139]),c_0_37]),c_0_79])]),c_0_140]),<A HREF="#c_0_141">c_0_141</A>])).

<A NAME="c_0_149"></A>fof(c_0_149,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | ~ memberP(nil,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_142])])).

<A NAME="c_0_150"></A>cnf(c_0_150,negated_conjecture,
    ( memberP(app(X1,esk51_0),esk52_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_143">c_0_143</A>,c_0_78]),c_0_57]),c_0_79])])).

<A NAME="c_0_151"></A>cnf(c_0_151,negated_conjecture,
    ( ssList(app(X1,esk51_0))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_144">c_0_144</A>,c_0_102]),c_0_57]),c_0_41])])).

<A NAME="c_0_152"></A>cnf(c_0_152,negated_conjecture,
    ( app(esk51_0,nil) = esk51_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_135">c_0_135</A>,<A HREF="#c_0_145">c_0_145</A>])).

<A NAME="c_0_153"></A>cnf(c_0_153,plain,
    ( nil = X1
    | memberP(X1,hd(X1))
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_146">c_0_146</A>,c_0_53]),c_0_44]),<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_154"></A>cnf(c_0_154,negated_conjecture,
    ( hd(esk50_0) = esk52_0
    | esk50_0 = nil
    | memberP(nil,esk52_0) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_147">c_0_147</A>,c_0_148]),c_0_41]),c_0_79])]),<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_155"></A>fof(c_0_155,plain,(
    ! [X2] :
      ( ( ~ rearsegP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | rearsegP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).

<A NAME="c_0_156"></A>cnf(c_0_156,plain,
    ( app(esk7_2(X1,nil),nil) = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_126">c_0_126</A>,c_0_114]),c_0_41])])).

<A NAME="c_0_157"></A>cnf(c_0_157,plain,
    ( ssList(esk7_2(X1,nil))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_128">c_0_128</A>,c_0_114]),c_0_41])])).

<A NAME="c_0_158"></A>cnf(c_0_158,plain,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_149">c_0_149</A>])).

<A NAME="c_0_159"></A>cnf(c_0_159,negated_conjecture,
    ( memberP(app(X1,app(X2,esk51_0)),esk52_0)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_143">c_0_143</A>,c_0_150]),c_0_79])]),<A HREF="#c_0_151">c_0_151</A>])).

<A NAME="c_0_160"></A>cnf(c_0_160,negated_conjecture,
    ( app(esk51_0,app(nil,X1)) = app(esk51_0,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_152]),c_0_41]),c_0_57])])).

<A NAME="c_0_161"></A>cnf(c_0_161,negated_conjecture,
    ( neq(esk50_0,nil)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_162"></A>cnf(c_0_162,negated_conjecture,
    ( esk50_0 = nil
    | memberP(nil,esk52_0) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_153">c_0_153</A>,c_0_154]),c_0_37])]),<A HREF="#c_0_140">c_0_140</A>])).

<A NAME="c_0_163"></A>cnf(c_0_163,plain,
    ( rearsegP(nil,X1)
    | ~ ssList(X1)
    | nil != X1 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_155">c_0_155</A>])).

<A NAME="c_0_164"></A>cnf(c_0_164,plain,
    ( memberP(esk7_2(X1,nil),X2)
    | ~ memberP(X1,X2)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_138">c_0_138</A>,c_0_156]),c_0_41])]),c_0_157]),<A HREF="#c_0_158">c_0_158</A>])).

<A NAME="c_0_165"></A>cnf(c_0_165,negated_conjecture,
    ( memberP(app(esk51_0,esk51_0),esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_159">c_0_159</A>,c_0_160]),c_0_57]),c_0_41])])).

<A NAME="c_0_166"></A>cnf(c_0_166,plain,
    ( ~ ssList(X1)
    | ~ ssList(X2)
    | X1 != X2
    | ~ neq(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>])).

<A NAME="c_0_167"></A>cnf(c_0_167,negated_conjecture,
    ( esk51_0 = nil
    | neq(esk50_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_161">c_0_161</A>,c_0_50]),c_0_41]),c_0_57])])).

<A NAME="c_0_168"></A>cnf(c_0_168,negated_conjecture,
    ( esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_158">c_0_158</A>,c_0_162]),c_0_79])])).

<A NAME="c_0_169"></A>cnf(c_0_169,plain,
    ( rearsegP(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_163]),c_0_41])])).

<A NAME="c_0_170"></A>cnf(c_0_170,negated_conjecture,
    ( memberP(esk7_2(app(esk51_0,esk51_0),nil),esk52_0)
    | ~ ssList(app(esk51_0,esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_164">c_0_164</A>,c_0_165]),c_0_79])])).

<A NAME="c_0_171"></A>cnf(c_0_171,plain,
    ( ~ ssList(X1)
    | ~ neq(X1,X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_166">c_0_166</A>])).

<A NAME="c_0_172"></A>cnf(c_0_172,negated_conjecture,
    ( esk51_0 = nil
    | neq(nil,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_167">c_0_167</A>,<A HREF="#c_0_168">c_0_168</A>])).

<A NAME="c_0_173"></A>cnf(c_0_173,plain,
    ( app(esk7_2(nil,nil),nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_126">c_0_126</A>,c_0_169]),c_0_41])])).

<A NAME="c_0_174"></A>cnf(c_0_174,plain,
    ( ssList(esk7_2(nil,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_128">c_0_128</A>,c_0_169]),c_0_41])])).

<A NAME="c_0_175"></A>cnf(c_0_175,negated_conjecture,
    ( memberP(esk7_2(app(esk51_0,esk51_0),nil),esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_170">c_0_170</A>,c_0_31]),c_0_57])])).

<A NAME="c_0_176"></A>cnf(c_0_176,negated_conjecture,
    ( esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_171">c_0_171</A>,c_0_172]),c_0_41])])).

<A NAME="c_0_177"></A>cnf(c_0_177,plain,
    ( esk7_2(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_173]),c_0_41]),c_0_174])])).

<A NAME="c_0_178"></A>cnf(c_0_178,negated_conjecture,
    ( ~ memberP(nil,esk52_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_140">c_0_140</A>,<A HREF="#c_0_168">c_0_168</A>])).

<A NAME="c_0_179"></A>cnf(c_0_179,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_175">c_0_175</A>,c_0_176]),c_0_176]),c_0_123]),c_0_177]),<A HREF="#c_0_178">c_0_178</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC410+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n142.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:43:38 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 2.72/2.90  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 2.72/2.90  # and selection function SelectCQPrecWNTNp.
% 2.72/2.90  #
% 2.72/2.90  # Preprocessing time       : 0.014 s
% 2.72/2.90  # Presaturation interreduction done
% 2.72/2.90  
% 2.72/2.90  # Proof found!
% 2.72/2.90  # SZS status Theorem
% 2.72/2.90  # SZS output start CNFRefutation
% 2.72/2.90  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssItem(X5)=>(~(memberP(X2,X5))|memberP(X1,X5))))|(nil!=X3&nil=X4))|(neq(X4,nil)&(~(neq(X3,nil))|?[X6]:((ssList(X6)&X4!=X6)&?[X7]:(ssList(X7)&?[X8]:((((ssList(X8)&tl(X3)=X7)&app(X7,X8)=X6)&?[X9]:(((ssItem(X9)&cons(X9,nil)=X8)&hd(X3)=X9)&neq(nil,X3)))&neq(nil,X3)))))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 2.72/2.90  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 2.72/2.90  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
% 2.72/2.90  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
% 2.72/2.90  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 2.72/2.90  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax22)).
% 2.72/2.90  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
% 2.72/2.90  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax78)).
% 2.72/2.90  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax86)).
% 2.72/2.90  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 2.72/2.90  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>hd(app(X1,X2))=hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax85)).
% 2.72/2.90  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>nil!=cons(X2,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax21)).
% 2.72/2.90  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 2.72/2.90  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
% 2.72/2.90  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>(memberP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(X3,cons(X2,X4))=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax3)).
% 2.72/2.90  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
% 2.72/2.90  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax23)).
% 2.72/2.90  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
% 2.72/2.90  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(rearsegP(X1,X2)=>rearsegP(app(X3,X1),X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax50)).
% 2.72/2.90  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>rearsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax51)).
% 2.72/2.90  fof(c_0_20, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)&lt;=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax37)).
% 2.72/2.90  fof(c_0_21, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax6)).
% 2.72/2.90  fof(c_0_22, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
% 2.72/2.90  fof(c_0_23, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(memberP(app(X2,X3),X1)&lt;=>(memberP(X2,X1)|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax36)).
% 2.72/2.90  fof(c_0_24, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax38)).
% 2.72/2.90  fof(c_0_25, axiom, (![X1]:(ssList(X1)=>(rearsegP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax52)).
% 2.72/2.90  fof(c_0_26, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssItem(X5)=>(~memberP(X2,X5)|memberP(X1,X5))))|(nil!=X3&nil=X4))|(neq(X4,nil)&(~neq(X3,nil)|?[X6]:((ssList(X6)&X4!=X6)&?[X7]:(ssList(X7)&?[X8]:((((ssList(X8)&tl(X3)=X7)&app(X7,X8)=X6)&?[X9]:(((ssItem(X9)&cons(X9,nil)=X8)&hd(X3)=X9)&neq(nil,X3)))&neq(nil,X3))))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 2.72/2.90  fof(c_0_27, negated_conjecture, (![X15]:![X16]:![X17]:![X18]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(ssItem(esk52_0)&(memberP(esk49_0,esk52_0)&~memberP(esk48_0,esk52_0))))&(nil=esk50_0|nil!=esk51_0))&((neq(esk50_0,nil)|~neq(esk51_0,nil))&(((~ssList(X15)|esk51_0=X15)|(~ssList(X16)|((((~ssList(X17)|tl(esk50_0)!=X16)|app(X16,X17)!=X15)|(((~ssItem(X18)|cons(X18,nil)!=X17)|hd(esk50_0)!=X18)|~neq(nil,esk50_0)))|~neq(nil,esk50_0))))|~neq(esk51_0,nil))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])])])])).
% 2.72/2.90  cnf(c_0_28, negated_conjecture, (esk51_0=X4|~neq(esk51_0,nil)|~neq(nil,esk50_0)|~neq(nil,esk50_0)|hd(esk50_0)!=X1|cons(X1,nil)!=X2|~ssItem(X1)|app(X3,X2)!=X4|tl(esk50_0)!=X3|~ssList(X2)|~ssList(X3)|~ssList(X4)), inference(split_conjunct,[status(thm)],[c_0_27])).
% 2.72/2.90  fof(c_0_29, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 2.72/2.90  cnf(c_0_30, negated_conjecture, (esk51_0=X4|hd(esk50_0)!=X1|tl(esk50_0)!=X3|app(X3,X2)!=X4|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X4)|~ssList(X3)|~ssList(X2)|~neq(nil,esk50_0)|~neq(esk51_0,nil)), inference(cn,[status(thm)],[c_0_28])).
% 2.72/2.90  cnf(c_0_31, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
% 2.72/2.90  fof(c_0_32, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
% 2.72/2.90  cnf(c_0_33, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
% 2.72/2.90  cnf(c_0_34, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_27])).
% 2.72/2.90  cnf(c_0_35, negated_conjecture, (app(tl(esk50_0),cons(hd(esk50_0),nil))=esk51_0|~ssList(cons(hd(esk50_0),nil))|~ssList(tl(esk50_0))|~neq(nil,esk50_0)|~neq(esk51_0,nil)|~ssItem(hd(esk50_0))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_30])])])]), c_0_31])).
% 2.72/2.90  cnf(c_0_36, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 2.72/2.90  cnf(c_0_37, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_33, c_0_34])).
% 2.72/2.90  fof(c_0_38, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
% 2.72/2.90  cnf(c_0_39, negated_conjecture, (app(tl(esk50_0),cons(hd(esk50_0),nil))=esk51_0|esk50_0=nil|~ssList(cons(hd(esk50_0),nil))|~neq(nil,esk50_0)|~neq(esk51_0,nil)|~ssItem(hd(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])])).
% 2.72/2.90  cnf(c_0_40, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
% 2.72/2.90  cnf(c_0_41, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 2.72/2.90  fof(c_0_42, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
% 2.72/2.90  cnf(c_0_43, negated_conjecture, (app(tl(esk50_0),cons(hd(esk50_0),nil))=esk51_0|esk50_0=nil|~neq(nil,esk50_0)|~neq(esk51_0,nil)|~ssItem(hd(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])])).
% 2.72/2.90  cnf(c_0_44, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
% 2.72/2.90  fof(c_0_45, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
% 2.72/2.90  fof(c_0_46, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
% 2.72/2.90  fof(c_0_47, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
% 2.72/2.90  fof(c_0_48, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
% 2.72/2.90  cnf(c_0_49, negated_conjecture, (app(tl(esk50_0),cons(hd(esk50_0),nil))=esk51_0|esk50_0=nil|~neq(nil,esk50_0)|~neq(esk51_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_37])])).
% 2.72/2.90  cnf(c_0_50, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
% 2.72/2.90  cnf(c_0_51, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
% 2.72/2.90  cnf(c_0_52, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_27])).
% 2.72/2.90  cnf(c_0_53, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
% 2.72/2.90  cnf(c_0_54, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
% 2.72/2.90  cnf(c_0_55, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
% 2.72/2.90  cnf(c_0_56, negated_conjecture, (app(tl(esk50_0),cons(hd(esk50_0),nil))=esk51_0|esk50_0=nil|~neq(esk51_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_37]), c_0_41])])).
% 2.72/2.90  cnf(c_0_57, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_51, c_0_52])).
% 2.72/2.90  cnf(c_0_58, negated_conjecture, (nil=esk50_0|nil!=esk51_0), inference(split_conjunct,[status(thm)],[c_0_27])).
% 2.72/2.90  cnf(c_0_59, plain, (cons(hd(app(X1,X2)),app(tl(X1),X2))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_31]), c_0_55])).
% 2.72/2.90  cnf(c_0_60, negated_conjecture, (app(tl(esk50_0),cons(hd(esk50_0),nil))=esk51_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_50]), c_0_41]), c_0_57])]), c_0_58])).
% 2.72/2.90  cnf(c_0_61, negated_conjecture, (cons(hd(app(esk50_0,cons(hd(esk50_0),nil))),esk51_0)=app(esk50_0,cons(hd(esk50_0),nil))|esk50_0=nil|~ssList(cons(hd(esk50_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_37])])).
% 2.72/2.90  cnf(c_0_62, negated_conjecture, (cons(hd(app(esk50_0,cons(hd(esk50_0),nil))),esk51_0)=app(esk50_0,cons(hd(esk50_0),nil))|esk50_0=nil|~ssItem(hd(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_40]), c_0_41])])).
% 2.72/2.90  fof(c_0_63, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|hd(app(X3,X4))=hd(X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
% 2.72/2.90  cnf(c_0_64, negated_conjecture, (cons(hd(app(esk50_0,cons(hd(esk50_0),nil))),esk51_0)=app(esk50_0,cons(hd(esk50_0),nil))|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_44]), c_0_37])])).
% 2.72/2.90  cnf(c_0_65, plain, (hd(app(X1,X2))=hd(X1)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
% 2.72/2.90  fof(c_0_66, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|nil!=cons(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
% 2.72/2.90  cnf(c_0_67, negated_conjecture, (app(esk50_0,cons(hd(esk50_0),nil))=cons(hd(esk50_0),esk51_0)|esk50_0=nil|~ssList(cons(hd(esk50_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_37])])).
% 2.72/2.90  cnf(c_0_68, plain, (nil!=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
% 2.72/2.90  fof(c_0_69, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
% 2.72/2.90  fof(c_0_70, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).
% 2.72/2.90  fof(c_0_71, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk3_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&((((ssList(esk4_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&(((app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6)))=X5|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(X9,cons(X6,X10))!=X5))|memberP(X5,X6))|~ssItem(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])])).
% 2.72/2.90  cnf(c_0_72, negated_conjecture, (memberP(esk49_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
% 2.72/2.90  cnf(c_0_73, negated_conjecture, (app(esk50_0,cons(hd(esk50_0),nil))=cons(hd(esk50_0),esk51_0)|esk50_0=nil|~ssItem(hd(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_40]), c_0_41])])).
% 2.72/2.90  cnf(c_0_74, negated_conjecture, (esk50_0=nil|app(esk50_0,cons(hd(esk50_0),nil))!=nil|~ssItem(hd(app(esk50_0,cons(hd(esk50_0),nil))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_64]), c_0_57])])).
% 2.72/2.90  cnf(c_0_75, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
% 2.72/2.90  cnf(c_0_76, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
% 2.72/2.90  cnf(c_0_77, plain, (app(esk3_2(X1,X2),cons(X2,esk4_2(X1,X2)))=X1|~ssList(X1)|~ssItem(X2)|~memberP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
% 2.72/2.90  cnf(c_0_78, negated_conjecture, (memberP(esk51_0,esk52_0)), inference(rw,[status(thm)],[c_0_72, c_0_52])).
% 2.72/2.90  cnf(c_0_79, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
% 2.72/2.90  cnf(c_0_80, plain, (ssList(esk3_2(X1,X2))|~ssList(X1)|~ssItem(X2)|~memberP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
% 2.72/2.90  cnf(c_0_81, negated_conjecture, (app(esk50_0,cons(hd(esk50_0),nil))=cons(hd(esk50_0),esk51_0)|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_44]), c_0_37])])).
% 2.72/2.90  cnf(c_0_82, negated_conjecture, (esk50_0=nil|app(esk50_0,cons(hd(esk50_0),nil))!=nil|~ssList(cons(hd(esk50_0),nil))|~ssItem(hd(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_65]), c_0_37])])).
% 2.72/2.90  fof(c_0_83, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
% 2.72/2.90  cnf(c_0_84, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_41])])).
% 2.72/2.90  cnf(c_0_85, negated_conjecture, (app(esk3_2(esk51_0,esk52_0),cons(esk52_0,esk4_2(esk51_0,esk52_0)))=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_57]), c_0_79])])).
% 2.72/2.90  cnf(c_0_86, negated_conjecture, (ssList(esk3_2(esk51_0,esk52_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_78]), c_0_57]), c_0_79])])).
% 2.72/2.90  cnf(c_0_87, plain, (ssList(esk4_2(X1,X2))|~ssList(X1)|~ssItem(X2)|~memberP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
% 2.72/2.90  cnf(c_0_88, negated_conjecture, (hd(cons(hd(esk50_0),esk51_0))=hd(esk50_0)|esk50_0=nil|~ssList(cons(hd(esk50_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_81]), c_0_37])])).
% 2.72/2.90  cnf(c_0_89, negated_conjecture, (esk50_0=nil|ssList(cons(hd(esk50_0),esk51_0))|~ssList(cons(hd(esk50_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_81]), c_0_37])])).
% 2.72/2.90  cnf(c_0_90, negated_conjecture, (esk50_0=nil|app(esk50_0,cons(hd(esk50_0),nil))!=nil|~ssItem(hd(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_40]), c_0_41])])).
% 2.72/2.90  fof(c_0_91, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
% 2.72/2.90  fof(c_0_92, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
% 2.72/2.90  cnf(c_0_93, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_83])).
% 2.72/2.90  cnf(c_0_94, negated_conjecture, (app(nil,esk51_0)=esk51_0|~ssList(cons(esk52_0,esk4_2(esk51_0,esk52_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86])])).
% 2.72/2.90  cnf(c_0_95, negated_conjecture, (ssList(esk4_2(esk51_0,esk52_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_78]), c_0_57]), c_0_79])])).
% 2.72/2.90  cnf(c_0_96, negated_conjecture, (hd(cons(hd(esk50_0),esk51_0))=hd(esk50_0)|esk50_0=nil|~ssItem(hd(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_40]), c_0_41])])).
% 2.72/2.90  cnf(c_0_97, negated_conjecture, (esk50_0=nil|ssList(cons(hd(esk50_0),esk51_0))|~ssItem(hd(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_40]), c_0_41])])).
% 2.72/2.90  cnf(c_0_98, negated_conjecture, (esk50_0=nil|app(esk50_0,cons(hd(esk50_0),nil))!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_44]), c_0_37])])).
% 2.72/2.90  cnf(c_0_99, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_91])).
% 2.72/2.90  cnf(c_0_100, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_92])).
% 2.72/2.90  cnf(c_0_101, plain, (hd(cons(X1,app(X2,X3)))=hd(cons(X1,X2))|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_93]), c_0_40]), c_0_68])).
% 2.72/2.90  cnf(c_0_102, negated_conjecture, (app(nil,esk51_0)=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_40]), c_0_95]), c_0_79])])).
% 2.72/2.90  cnf(c_0_103, negated_conjecture, (hd(cons(hd(esk50_0),esk51_0))=hd(esk50_0)|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_44]), c_0_37])])).
% 2.72/2.90  cnf(c_0_104, negated_conjecture, (esk50_0=nil|ssList(cons(hd(esk50_0),esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_44]), c_0_37])])).
% 2.72/2.90  cnf(c_0_105, negated_conjecture, (esk50_0=nil|cons(hd(esk50_0),esk51_0)!=nil), inference(spm,[status(thm)],[c_0_98, c_0_81])).
% 2.72/2.90  fof(c_0_106, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~rearsegP(X4,X5)|rearsegP(app(X6,X4),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
% 2.72/2.90  fof(c_0_107, plain, (![X2]:(~ssList(X2)|rearsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).
% 2.72/2.90  fof(c_0_108, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])).
% 2.72/2.90  cnf(c_0_109, plain, (cons(X1,tl(cons(X1,X2)))=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_99]), c_0_40]), c_0_68])).
% 2.72/2.90  cnf(c_0_110, plain, (cons(hd(cons(X1,X2)),X2)=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_100]), c_0_40]), c_0_68])).
% 2.72/2.90  cnf(c_0_111, negated_conjecture, (hd(cons(X1,esk51_0))=hd(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_57]), c_0_41])])).
% 2.72/2.90  cnf(c_0_112, negated_conjecture, (esk50_0=nil|ssItem(hd(esk50_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_103]), c_0_104]), c_0_105])).
% 2.72/2.90  cnf(c_0_113, plain, (rearsegP(app(X1,X2),X3)|~rearsegP(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_106])).
% 2.72/2.90  cnf(c_0_114, plain, (rearsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_107])).
% 2.72/2.90  cnf(c_0_115, plain, (memberP(cons(X2,X3),X1)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~memberP(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
% 2.72/2.90  cnf(c_0_116, plain, (nil=app(X1,X2)|~ssList(X1)|~ssList(X2)|nil!=X1|nil!=X2), inference(split_conjunct,[status(thm)],[c_0_48])).
% 2.72/2.90  cnf(c_0_117, plain, (cons(hd(cons(X1,X2)),tl(cons(X1,X2)))=cons(X1,X2)|~ssList(X2)|~ssItem(hd(cons(X1,X2)))|~ssItem(X1)), inference(spm,[status(thm)],[c_0_109, c_0_110])).
% 2.72/2.90  cnf(c_0_118, negated_conjecture, (hd(cons(hd(esk50_0),nil))=hd(esk50_0)|esk50_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_111]), c_0_112])).
% 2.72/2.90  fof(c_0_119, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).
% 2.72/2.90  cnf(c_0_120, plain, (rearsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_41])])).
% 2.72/2.90  cnf(c_0_121, negated_conjecture, (memberP(cons(X1,esk51_0),esk52_0)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_78]), c_0_57]), c_0_79])])).
% 2.72/2.90  cnf(c_0_122, plain, (ssList(cons(X1,app(X2,X3)))|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_93]), c_0_40])).
% 2.72/2.90  cnf(c_0_123, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_116])]), c_0_41])])).
% 2.72/2.90  cnf(c_0_124, negated_conjecture, (cons(hd(esk50_0),tl(cons(hd(esk50_0),nil)))=cons(hd(esk50_0),nil)|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_41])]), c_0_112])).
% 2.72/2.90  fof(c_0_125, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])).
% 2.72/2.90  cnf(c_0_126, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_119])).
% 2.72/2.90  cnf(c_0_127, negated_conjecture, (rearsegP(esk51_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_102]), c_0_57]), c_0_41])])).
% 2.72/2.90  cnf(c_0_128, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_119])).
% 2.72/2.90  fof(c_0_129, plain, (![X4]:![X5]:![X6]:(((((~memberP(app(X5,X6),X4)|(memberP(X5,X4)|memberP(X6,X4)))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&(((((~memberP(X5,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).
% 2.72/2.90  cnf(c_0_130, negated_conjecture, (esk50_0=nil|memberP(app(esk50_0,cons(hd(esk50_0),nil)),esk52_0)|~ssItem(hd(app(esk50_0,cons(hd(esk50_0),nil))))), inference(spm,[status(thm)],[c_0_121, c_0_64])).
% 2.72/2.90  cnf(c_0_131, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_41])])).
% 2.72/2.90  cnf(c_0_132, negated_conjecture, (~memberP(esk48_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
% 2.72/2.90  cnf(c_0_133, negated_conjecture, (esk50_0=nil|ssList(cons(hd(esk50_0),nil))|~ssList(tl(cons(hd(esk50_0),nil)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_124]), c_0_112])).
% 2.72/2.90  cnf(c_0_134, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_125])).
% 2.72/2.90  cnf(c_0_135, negated_conjecture, (app(esk7_2(esk51_0,nil),nil)=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_41]), c_0_57])])).
% 2.72/2.90  cnf(c_0_136, negated_conjecture, (ssList(esk7_2(esk51_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_127]), c_0_41]), c_0_57])])).
% 2.72/2.90  cnf(c_0_137, plain, (memberP(cons(X2,X3),X1)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_108])).
% 2.72/2.90  cnf(c_0_138, plain, (memberP(X3,X1)|memberP(X2,X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(app(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_129])).
% 2.72/2.90  cnf(c_0_139, negated_conjecture, (esk50_0=nil|memberP(app(esk50_0,cons(hd(esk50_0),nil)),esk52_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_65]), c_0_37])]), c_0_131]), c_0_112])).
% 2.72/2.90  cnf(c_0_140, negated_conjecture, (~memberP(esk50_0,esk52_0)), inference(rw,[status(thm)],[c_0_132, c_0_34])).
% 2.72/2.90  cnf(c_0_141, negated_conjecture, (esk50_0=nil|ssList(cons(hd(esk50_0),nil))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_100]), c_0_41])]), c_0_112])).
% 2.72/2.90  fof(c_0_142, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_24])).
% 2.72/2.90  cnf(c_0_143, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_129])).
% 2.72/2.90  cnf(c_0_144, plain, (ssList(app(X1,app(X2,X3)))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_75]), c_0_31])).
% 2.72/2.90  cnf(c_0_145, negated_conjecture, (esk7_2(esk51_0,nil)=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_136])])).
% 2.72/2.90  cnf(c_0_146, plain, (memberP(cons(X1,X2),X1)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_137])).
% 2.72/2.90  cnf(c_0_147, plain, (memberP(X3,X1)|X1=X2|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~memberP(cons(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
% 2.72/2.90  cnf(c_0_148, negated_conjecture, (esk50_0=nil|memberP(cons(hd(esk50_0),nil),esk52_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_37]), c_0_79])]), c_0_140]), c_0_141])).
% 2.72/2.90  fof(c_0_149, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_142])])).
% 2.72/2.90  cnf(c_0_150, negated_conjecture, (memberP(app(X1,esk51_0),esk52_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_78]), c_0_57]), c_0_79])])).
% 2.72/2.90  cnf(c_0_151, negated_conjecture, (ssList(app(X1,esk51_0))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_102]), c_0_57]), c_0_41])])).
% 2.72/2.90  cnf(c_0_152, negated_conjecture, (app(esk51_0,nil)=esk51_0), inference(rw,[status(thm)],[c_0_135, c_0_145])).
% 2.72/2.90  cnf(c_0_153, plain, (nil=X1|memberP(X1,hd(X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_53]), c_0_44]), c_0_36])).
% 2.72/2.90  cnf(c_0_154, negated_conjecture, (hd(esk50_0)=esk52_0|esk50_0=nil|memberP(nil,esk52_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_148]), c_0_41]), c_0_79])]), c_0_112])).
% 2.72/2.90  fof(c_0_155, plain, (![X2]:(((~rearsegP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|rearsegP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
% 2.72/2.90  cnf(c_0_156, plain, (app(esk7_2(X1,nil),nil)=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_114]), c_0_41])])).
% 2.72/2.90  cnf(c_0_157, plain, (ssList(esk7_2(X1,nil))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_114]), c_0_41])])).
% 2.72/2.90  cnf(c_0_158, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_149])).
% 2.72/2.90  cnf(c_0_159, negated_conjecture, (memberP(app(X1,app(X2,esk51_0)),esk52_0)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_150]), c_0_79])]), c_0_151])).
% 2.72/2.90  cnf(c_0_160, negated_conjecture, (app(esk51_0,app(nil,X1))=app(esk51_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_152]), c_0_41]), c_0_57])])).
% 2.72/2.90  cnf(c_0_161, negated_conjecture, (neq(esk50_0,nil)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_27])).
% 2.72/2.90  cnf(c_0_162, negated_conjecture, (esk50_0=nil|memberP(nil,esk52_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154]), c_0_37])]), c_0_140])).
% 2.72/2.90  cnf(c_0_163, plain, (rearsegP(nil,X1)|~ssList(X1)|nil!=X1), inference(split_conjunct,[status(thm)],[c_0_155])).
% 2.72/2.90  cnf(c_0_164, plain, (memberP(esk7_2(X1,nil),X2)|~memberP(X1,X2)|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_156]), c_0_41])]), c_0_157]), c_0_158])).
% 2.72/2.90  cnf(c_0_165, negated_conjecture, (memberP(app(esk51_0,esk51_0),esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_57]), c_0_41])])).
% 2.72/2.90  cnf(c_0_166, plain, (~ssList(X1)|~ssList(X2)|X1!=X2|~neq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
% 2.72/2.90  cnf(c_0_167, negated_conjecture, (esk51_0=nil|neq(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_50]), c_0_41]), c_0_57])])).
% 2.72/2.90  cnf(c_0_168, negated_conjecture, (esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_162]), c_0_79])])).
% 2.72/2.90  cnf(c_0_169, plain, (rearsegP(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_163]), c_0_41])])).
% 2.72/2.90  cnf(c_0_170, negated_conjecture, (memberP(esk7_2(app(esk51_0,esk51_0),nil),esk52_0)|~ssList(app(esk51_0,esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_79])])).
% 2.72/2.90  cnf(c_0_171, plain, (~ssList(X1)|~neq(X1,X1)), inference(er,[status(thm)],[c_0_166])).
% 2.72/2.90  cnf(c_0_172, negated_conjecture, (esk51_0=nil|neq(nil,nil)), inference(rw,[status(thm)],[c_0_167, c_0_168])).
% 2.72/2.90  cnf(c_0_173, plain, (app(esk7_2(nil,nil),nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_169]), c_0_41])])).
% 2.72/2.90  cnf(c_0_174, plain, (ssList(esk7_2(nil,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_169]), c_0_41])])).
% 2.72/2.90  cnf(c_0_175, negated_conjecture, (memberP(esk7_2(app(esk51_0,esk51_0),nil),esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_31]), c_0_57])])).
% 2.72/2.90  cnf(c_0_176, negated_conjecture, (esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_41])])).
% 2.72/2.90  cnf(c_0_177, plain, (esk7_2(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_173]), c_0_41]), c_0_174])])).
% 2.72/2.90  cnf(c_0_178, negated_conjecture, (~memberP(nil,esk52_0)), inference(rw,[status(thm)],[c_0_140, c_0_168])).
% 2.72/2.90  cnf(c_0_179, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_175, c_0_176]), c_0_176]), c_0_123]), c_0_177]), c_0_178]), ['proof']).
% 2.72/2.90  # SZS output end CNFRefutation
% 2.72/2.90  # Training examples: 0 positive, 0 negative
% 2.72/2.91  
% 2.72/2.91  # -------------------------------------------------
% 2.72/2.91  # User time                : 2.627 s
% 2.72/2.91  # System time              : 0.052 s
% 2.72/2.91  # Total time               : 2.679 s
% 2.72/2.91  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
