<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC364-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC364-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC364-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC364-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC364-1.p">SWC364-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n179.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:39 EDT 2016

% Result     : Unsatisfiable 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of clauses        :   27 (  38 expanded)
%              Number of leaves         :   16 (  25 expanded)
%              Depth                    :    7
%              Number of atoms          :   57 (  70 expanded)
%              Number of equality atoms :    9 (  12 expanded)
%              Maximal clause size      :    6 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( neq(sk2,nil)
    | neq(sk2,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_1"></A>cnf(c_0_1,axiom,
    ( rearsegP(X3,X2)
    | app(X1,X2) != X3
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause143)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( app(cons(sk5,nil),sk3) = sk4
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( ssItem(sk5)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( app(skaf46(X1,X2),X2) = X1
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause131)).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( ssList(skaf46(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause50)).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( ~ segmentP(sk2,sk1)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_14)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( neq(sk2,nil) ),
    inference(cn,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_1]),<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( app(cons(sk5,nil),sk1) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>,c_0_4]),c_0_5]),c_0_5]),c_0_16])])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( rearsegP(sk2,sk1)
    | ~ ssList(cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_18]),c_0_6])])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( ssItem(sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_5]),c_0_16])])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( rearsegP(sk2,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_8]),c_0_20]),c_0_9])])).

<A NAME="c_0_22"></A>cnf(c_0_22,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( app(skaf46(sk2,sk1),sk1) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_21]),c_0_6]),c_0_12])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( segmentP(app(sk2,X1),sk1)
    | ~ ssList(app(sk2,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_23]),c_0_6]),c_0_13])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( ~ segmentP(sk2,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_5]),c_0_16])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_15]),c_0_12]),c_0_9])]),<A HREF="#c_0_25">c_0_25</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC364-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n179.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:42:08 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.27  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.02/0.27  # and selection function SelectCQArNTEqFirst.
% 0.02/0.27  #
% 0.02/0.27  # Preprocessing time       : 0.009 s
% 0.02/0.27  # Presaturation interreduction done
% 0.02/0.27  
% 0.02/0.27  # Proof found!
% 0.02/0.27  # SZS status Unsatisfiable
% 0.02/0.27  # SZS output start CNFRefutation
% 0.02/0.27  cnf(c_0_0, negated_conjecture, (neq(sk2,nil)|neq(sk2,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 0.02/0.27  cnf(c_0_1, axiom, (rearsegP(X3,X2)|app(X1,X2)!=X3|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause143)).
% 0.02/0.27  cnf(c_0_2, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.02/0.27  cnf(c_0_3, negated_conjecture, (app(cons(sk5,nil),sk3)=sk4|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
% 0.02/0.27  cnf(c_0_4, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.02/0.27  cnf(c_0_5, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 0.02/0.27  cnf(c_0_6, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.02/0.27  cnf(c_0_7, negated_conjecture, (ssItem(sk5)|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_12)).
% 0.02/0.27  cnf(c_0_8, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.02/0.27  cnf(c_0_9, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.02/0.27  cnf(c_0_10, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause173)).
% 0.02/0.27  cnf(c_0_11, axiom, (app(skaf46(X1,X2),X2)=X1|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause131)).
% 0.02/0.27  cnf(c_0_12, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 0.02/0.27  cnf(c_0_13, axiom, (ssList(skaf46(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause50)).
% 0.02/0.27  cnf(c_0_14, negated_conjecture, (~segmentP(sk2,sk1)|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_14)).
% 0.02/0.27  cnf(c_0_15, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause73)).
% 0.02/0.27  cnf(c_0_16, negated_conjecture, (neq(sk2,nil)), inference(cn,[status(thm)],[c_0_0])).
% 0.02/0.27  cnf(c_0_17, plain, (rearsegP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_1]), c_0_2])).
% 0.02/0.27  cnf(c_0_18, negated_conjecture, (app(cons(sk5,nil),sk1)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_4]), c_0_5]), c_0_5]), c_0_16])])).
% 0.02/0.27  cnf(c_0_19, negated_conjecture, (rearsegP(sk2,sk1)|~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_6])])).
% 0.02/0.27  cnf(c_0_20, negated_conjecture, (ssItem(sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_5]), c_0_16])])).
% 0.02/0.27  cnf(c_0_21, negated_conjecture, (rearsegP(sk2,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_8]), c_0_20]), c_0_9])])).
% 0.02/0.27  cnf(c_0_22, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(er,[status(thm)],[c_0_10])).
% 0.02/0.27  cnf(c_0_23, negated_conjecture, (app(skaf46(sk2,sk1),sk1)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_21]), c_0_6]), c_0_12])])).
% 0.02/0.27  cnf(c_0_24, negated_conjecture, (segmentP(app(sk2,X1),sk1)|~ssList(app(sk2,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_6]), c_0_13])])).
% 0.02/0.27  cnf(c_0_25, negated_conjecture, (~segmentP(sk2,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_5]), c_0_16])])).
% 0.02/0.27  cnf(c_0_26, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_15]), c_0_12]), c_0_9])]), c_0_25]), ['proof']).
% 0.02/0.27  # SZS output end CNFRefutation
% 0.02/0.27  # Training examples: 0 positive, 0 negative
% 0.02/0.27  
% 0.02/0.27  # -------------------------------------------------
% 0.02/0.27  # User time                : 0.043 s
% 0.02/0.27  # System time              : 0.002 s
% 0.02/0.27  # Total time               : 0.045 s
% 0.02/0.27  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
