<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC288-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC288-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC288-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC288-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC288-1.p">SWC288-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n153.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:21 EDT 2016

% Result     : Unsatisfiable 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of clauses        :    9 (  12 expanded)
%              Number of leaves         :    6 (   8 expanded)
%              Depth                    :    3
%              Number of atoms          :   12 (  15 expanded)
%              Number of equality atoms :    5 (   6 expanded)
%              Maximal clause size      :    2 (   1 average)
%              Maximal term depth       :    2 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( ~ strictorderedP(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( strictorderedP(cons(X1,nil))
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause66)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( cons(sk5,nil) = sk3
    | nil = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_18)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssItem(sk5)
    | nil = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( strictorderedP(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause3)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ~ strictorderedP(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( sk3 = nil ),
    inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_3]),c_0_6]),<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>,c_0_7]),c_0_5])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC288-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.22  % Computer   : n153.star.cs.uiowa.edu
% 0.03/0.22  % Model      : x86_64 x86_64
% 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.22  % Memory     : 32218.75MB
% 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.22  % CPULimit   : 300
% 0.03/0.22  % DateTime   : Tue May 10 20:40:22 CDT 2016
% 0.03/0.22  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.03/0.25  # and selection function SelectCQArNTEqFirst.
% 0.03/0.25  #
% 0.03/0.25  # Preprocessing time       : 0.009 s
% 0.03/0.25  # Presaturation interreduction done
% 0.03/0.25  
% 0.03/0.25  # Proof found!
% 0.03/0.25  # SZS status Unsatisfiable
% 0.03/0.25  # SZS output start CNFRefutation
% 0.03/0.25  cnf(c_0_0, negated_conjecture, (~strictorderedP(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 0.03/0.25  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 0.03/0.25  cnf(c_0_2, axiom, (strictorderedP(cons(X1,nil))|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause66)).
% 0.03/0.25  cnf(c_0_3, negated_conjecture, (cons(sk5,nil)=sk3|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_18)).
% 0.03/0.25  cnf(c_0_4, negated_conjecture, (ssItem(sk5)|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
% 0.03/0.25  cnf(c_0_5, axiom, (strictorderedP(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause3)).
% 0.03/0.25  cnf(c_0_6, negated_conjecture, (~strictorderedP(sk3)), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 0.03/0.25  cnf(c_0_7, negated_conjecture, (sk3=nil), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_3]), c_0_6]), c_0_4])).
% 0.03/0.25  cnf(c_0_8, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_7]), c_0_5])]), ['proof']).
% 0.03/0.25  # SZS output end CNFRefutation
% 0.03/0.25  # Training examples: 0 positive, 0 negative
% 0.03/0.25  
% 0.03/0.25  # -------------------------------------------------
% 0.03/0.25  # User time                : 0.020 s
% 0.03/0.25  # System time              : 0.002 s
% 0.03/0.25  # Total time               : 0.022 s
% 0.03/0.25  # Maximum resident set size: 1868 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
