<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC398+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC398+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC398+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC398%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC398+1.p">SWC398+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n137.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:47 EDT 2016

% Result     : Theorem 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of formulae       :   19 (  44 expanded)
%              Number of clauses        :   14 (  15 expanded)
%              Number of leaves         :    2 (  10 expanded)
%              Depth                    :    6
%              Number of atoms          :   98 ( 555 expanded)
%              Number of equality atoms :   24 ( 193 expanded)
%              Maximal formula depth    :   24 (   6 average)
%              Maximal clause size      :   19 (   1 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ! [X5] :
                        ( ssList(X5)
                       => ( app(X3,X5) != X4
                          | ~ equalelemsP(X3)
                          | ? [X6] :
                              ( ssItem(X6)
                              & ? [X7] :
                                  ( ssList(X7)
                                  & app(cons(X6,nil),X7) = X5
                                  & ? [X8] :
                                      ( ssList(X8)
                                      & app(X8,cons(X6,nil)) = X3 ) ) ) ) )
                    | ! [X9] :
                        ( ssItem(X9)
                       => ( ~ memberP(X1,X9)
                          | memberP(X2,X9) ) )
                    | ( nil != X4
                      & nil = X3 ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(app(X2,X3),X1)
              &lt;=> ( memberP(X2,X1)
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax36)).

<A NAME="c_0_2"></A>fof(c_0_2,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ! [X5] :
                          ( ssList(X5)
                         => ( app(X3,X5) != X4
                            | ~ equalelemsP(X3)
                            | ? [X6] :
                                ( ssItem(X6)
                                & ? [X7] :
                                    ( ssList(X7)
                                    & app(cons(X6,nil),X7) = X5
                                    & ? [X8] :
                                        ( ssList(X8)
                                        & app(X8,cons(X6,nil)) = X3 ) ) ) ) )
                      | ! [X9] :
                          ( ssItem(X9)
                         => ( ~ memberP(X1,X9)
                            | memberP(X2,X9) ) )
                      | ( nil != X4
                        & nil = X3 ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_3"></A>fof(c_0_3,negated_conjecture,(
    ! [X15,X16,X17] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ssList(esk52_0)
      & app(esk50_0,esk52_0) = esk51_0
      & equalelemsP(esk50_0)
      & ( ~ ssItem(X15)
        | ~ ssList(X16)
        | app(cons(X15,nil),X16) != esk52_0
        | ~ ssList(X17)
        | app(X17,cons(X15,nil)) != esk50_0 )
      & ssItem(esk53_0)
      & memberP(esk48_0,esk53_0)
      & ~ memberP(esk49_0,esk53_0)
      & ( nil = esk51_0
        | nil != esk50_0 ) ) ),
    inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).

<A NAME="c_0_4"></A>fof(c_0_4,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(app(X5,X6),X4)
        | memberP(X5,X4)
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X5,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( memberP(esk48_0,esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_8"></A>cnf(c_0_8,plain,
    ( memberP(app(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(X2,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( memberP(esk50_0,esk53_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ssItem(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( ~ memberP(esk49_0,esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( memberP(app(esk50_0,X1),esk53_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_9]),c_0_10]),c_0_11])])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( app(esk50_0,esk52_0) = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ssList(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( ~ memberP(esk51_0,esk53_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_15]),c_0_16])]),<A HREF="#c_0_17">c_0_17</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC398+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n137.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:43:37 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.26  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.02/0.26  # and selection function SelectCQPrecWNTNp.
% 0.02/0.26  #
% 0.02/0.26  # Preprocessing time       : 0.014 s
% 0.02/0.26  # Presaturation interreduction done
% 0.02/0.26  
% 0.02/0.26  # Proof found!
% 0.02/0.26  # SZS status Theorem
% 0.02/0.26  # SZS output start CNFRefutation
% 0.02/0.26  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>((app(X3,X5)!=X4|~(equalelemsP(X3)))|?[X6]:(ssItem(X6)&?[X7]:((ssList(X7)&app(cons(X6,nil),X7)=X5)&?[X8]:(ssList(X8)&app(X8,cons(X6,nil))=X3))))))|![X9]:(ssItem(X9)=>(~(memberP(X1,X9))|memberP(X2,X9))))|(nil!=X4&nil=X3))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.02/0.26  fof(c_0_1, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(memberP(app(X2,X3),X1)&lt;=>(memberP(X2,X1)|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax36)).
% 0.02/0.26  fof(c_0_2, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>((app(X3,X5)!=X4|~equalelemsP(X3))|?[X6]:(ssItem(X6)&?[X7]:((ssList(X7)&app(cons(X6,nil),X7)=X5)&?[X8]:(ssList(X8)&app(X8,cons(X6,nil))=X3))))))|![X9]:(ssItem(X9)=>(~memberP(X1,X9)|memberP(X2,X9))))|(nil!=X4&nil=X3)))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.02/0.26  fof(c_0_3, negated_conjecture, (![X15]:![X16]:![X17]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(ssList(esk52_0)&((app(esk50_0,esk52_0)=esk51_0&equalelemsP(esk50_0))&(~ssItem(X15)|((~ssList(X16)|app(cons(X15,nil),X16)!=esk52_0)|(~ssList(X17)|app(X17,cons(X15,nil))!=esk50_0))))))&(ssItem(esk53_0)&(memberP(esk48_0,esk53_0)&~memberP(esk49_0,esk53_0))))&(nil=esk51_0|nil!=esk50_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
% 0.02/0.26  fof(c_0_4, plain, (![X4]:![X5]:![X6]:(((((~memberP(app(X5,X6),X4)|(memberP(X5,X4)|memberP(X6,X4)))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&(((((~memberP(X5,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
% 0.02/0.26  cnf(c_0_5, negated_conjecture, (memberP(esk48_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.02/0.26  cnf(c_0_6, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.02/0.26  cnf(c_0_7, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.02/0.26  cnf(c_0_8, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.26  cnf(c_0_9, negated_conjecture, (memberP(esk50_0,esk53_0)), inference(rw,[status(thm)],[c_0_5, c_0_6])).
% 0.02/0.26  cnf(c_0_10, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_7, c_0_6])).
% 0.02/0.26  cnf(c_0_11, negated_conjecture, (ssItem(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.02/0.26  cnf(c_0_12, negated_conjecture, (~memberP(esk49_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.02/0.26  cnf(c_0_13, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.02/0.26  cnf(c_0_14, negated_conjecture, (memberP(app(esk50_0,X1),esk53_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10]), c_0_11])])).
% 0.02/0.26  cnf(c_0_15, negated_conjecture, (app(esk50_0,esk52_0)=esk51_0), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.02/0.26  cnf(c_0_16, negated_conjecture, (ssList(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.02/0.26  cnf(c_0_17, negated_conjecture, (~memberP(esk51_0,esk53_0)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
% 0.02/0.26  cnf(c_0_18, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])]), c_0_17]), ['proof']).
% 0.02/0.26  # SZS output end CNFRefutation
% 0.02/0.26  # Training examples: 0 positive, 0 negative
% 0.02/0.26  
% 0.02/0.26  # -------------------------------------------------
% 0.02/0.26  # User time                : 0.031 s
% 0.02/0.26  # System time              : 0.006 s
% 0.02/0.26  # Total time               : 0.037 s
% 0.02/0.26  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
