<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC125-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC125-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC125-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC125-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC125-1.p">SWC125-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n169.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:45 EDT 2016

% Result     : Unsatisfiable 2.21s
% Output     : CNFRefutation 2.21s
% Verified   : 
% Statistics : Number of clauses        :  116 ( 806 expanded)
%              Number of leaves         :   52 ( 540 expanded)
%              Depth                    :   19
%              Number of atoms          :  308 (1770 expanded)
%              Number of equality atoms :   91 ( 504 expanded)
%              Maximal clause size      :    7 (   3 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( ~ neq(sk1,nil)
    | ~ segmentP(sk2,sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( segmentP(sk4,sk3) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause99)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( cons(skaf83(X1),skaf82(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause109)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( ssItem(skaf83(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause12)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( ssList(skaf82(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause13)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( app(app(skaf47(X1,X2),X2),skaf48(X2,X1)) = X1
    | ~ segmentP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause168)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( rearsegP(app(X3,X1),X2)
    | ~ rearsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause136)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( rearsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause58)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( ssList(skaf47(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause49)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( nil = X2
    | tl(app(X2,X1)) = app(tl(X2),X1)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause133)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( nil = X1
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause118)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( ssList(skaf48(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause48)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause75)).

<A NAME="c_0_23"></A>cnf(c_0_23,axiom,
    ( memberP(cons(X2,X3),X1)
    | X1 != X2
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause138)).

<A NAME="c_0_24"></A>cnf(c_0_24,axiom,
    ( app(skaf46(X1,X2),X2) = X1
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause131)).

<A NAME="c_0_25"></A>cnf(c_0_25,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_26"></A>cnf(c_0_26,axiom,
    ( cons(X1,X2) != nil
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause98)).

<A NAME="c_0_27"></A>cnf(c_0_27,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_28"></A>cnf(c_0_28,axiom,
    ( ssList(skaf46(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause50)).

<A NAME="c_0_29"></A>cnf(c_0_29,axiom,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause157)).

<A NAME="c_0_30"></A>cnf(c_0_30,axiom,
    ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
    | ~ memberP(X1,X2)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause169)).

<A NAME="c_0_31"></A>cnf(c_0_31,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_32"></A>cnf(c_0_32,axiom,
    ( ssList(skaf43(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause52)).

<A NAME="c_0_33"></A>cnf(c_0_33,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_34"></A>cnf(c_0_34,axiom,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause104)).

<A NAME="c_0_35"></A>cnf(c_0_35,axiom,
    ( nil = X2
    | hd(app(X2,X1)) = hd(X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause123)).

<A NAME="c_0_36"></A>cnf(c_0_36,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_37"></A>cnf(c_0_37,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_38"></A>cnf(c_0_38,axiom,
    ( X4 = X2
    | cons(X1,X2) != cons(X3,X4)
    | ~ ssItem(X3)
    | ~ ssItem(X1)
    | ~ ssList(X4)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause171)).

<A NAME="c_0_39"></A>cnf(c_0_39,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_40"></A>cnf(c_0_40,axiom,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause60)).

<A NAME="c_0_41"></A>cnf(c_0_41,axiom,
    ( app(X1,X2) = nil
    | nil != X1
    | nil != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause145)).

<A NAME="c_0_42"></A>cnf(c_0_42,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_43"></A>cnf(c_0_43,axiom,
    ( frontsegP(cons(X3,X1),cons(X4,X2))
    | ~ frontsegP(X1,X2)
    | X3 != X4
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssItem(X4)
    | ~ ssItem(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause178)).

<A NAME="c_0_44"></A>cnf(c_0_44,axiom,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause76)).

<A NAME="c_0_45"></A>cnf(c_0_45,axiom,
    ( ssItem(skaf44(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause47)).

<A NAME="c_0_46"></A>cnf(c_0_46,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( ~ ssList(X1)
    | ~ neq(sk3,X1)
    | ~ segmentP(sk4,X1)
    | ~ segmentP(X1,sk3)
    | ~ strictorderedP(X1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_48"></A>cnf(c_0_48,axiom,
    ( segmentP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause56)).

<A NAME="c_0_49"></A>cnf(c_0_49,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_50"></A>cnf(c_0_50,axiom,
    ( strictorderedP(cons(X1,nil))
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause66)).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( neq(sk2,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( ~ neq(sk3,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,c_0_1]),c_0_2]),c_0_3]),c_0_2])])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( sk3 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_5]),c_0_53]),c_0_6])])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( segmentP(sk4,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>,<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,plain,
    ( nil = X1
    | skaf82(X1) != X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_9]),c_0_10]),c_0_11])])).

<A NAME="c_0_58"></A>cnf(c_0_58,plain,
    ( skaf82(X1) = tl(X1)
    | nil = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_9]),c_0_10]),c_0_11])])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( app(app(skaf47(sk4,nil),nil),skaf48(nil,sk4)) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_55]),c_0_6]),c_0_56])])).

<A NAME="c_0_60"></A>cnf(c_0_60,plain,
    ( nil = X1
    | tl(X1) != X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,<A HREF="#c_0_58">c_0_58</A>])).

<A NAME="c_0_61"></A>cnf(c_0_61,plain,
    ( rearsegP(app(X1,X2),nil)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_15]),c_0_6])])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( app(skaf47(sk4,nil),skaf48(nil,sk4)) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,c_0_16]),c_0_17])])).

<A NAME="c_0_63"></A>cnf(c_0_63,plain,
    ( nil = X1
    | app(tl(X1),X2) != app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>,c_0_18]),c_0_19]),<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( rearsegP(sk4,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_62]),c_0_17]),c_0_21])])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( nil = X1
    | tl(X1) != app(X1,nil)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>,c_0_16]),c_0_6])]),<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( app(skaf46(sk4,nil),nil) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_64]),c_0_6]),c_0_56])])).

<A NAME="c_0_68"></A>cnf(c_0_68,plain,
    ( app(cons(X1,X2),nil) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_12]),c_0_25]),<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,plain,
    ( nil = X1
    | memberP(X1,skaf83(X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>,c_0_9]),c_0_10]),c_0_11])])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( skaf46(sk4,nil) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_67]),c_0_28])])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( cons(X1,app(X2,nil)) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_29]),c_0_6])])).

<A NAME="c_0_73"></A>cnf(c_0_73,plain,
    ( nil = X1
    | ~ ssList(app(app(X2,cons(skaf83(X1),X3)),X1))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(app(X2,cons(skaf83(X1),X3)),X1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>,c_0_9]),c_0_10]),c_0_11])])).

<A NAME="c_0_74"></A>cnf(c_0_74,plain,
    ( app(skaf42(X1,skaf83(X1)),cons(skaf83(X1),skaf43(skaf83(X1),X1))) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_70]),c_0_10])])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( app(sk4,nil) = sk4 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,<A HREF="#c_0_71">c_0_71</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,plain,
    ( cons(X1,app(X2,app(X3,nil))) != app(X2,X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_31]),c_0_6])]),<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,plain,
    ( nil = X1
    | ~ ssList(app(X1,X1))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(X1,X1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>,c_0_74]),c_0_32]),c_0_33])])).

<A NAME="c_0_78"></A>cnf(c_0_78,plain,
    ( cons(hd(X1),tl(app(X1,X2))) = app(X1,X2)
    | nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_35]),c_0_19]),<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,plain,
    ( nil = X1
    | ssList(app(tl(X1),X2))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_18]),c_0_19]),<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( app(tl(sk4),nil) = tl(sk4)
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_75]),c_0_56]),c_0_6])])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( cons(X1,app(X2,nil)) != app(nil,X2)
    | ~ ssItem(X1)
    | ~ ssList(app(X2,nil))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>,c_0_36]),c_0_6])])).

<A NAME="c_0_82"></A>cnf(c_0_82,plain,
    ( nil = X1
    | ssItem(X2)
    | ~ ssList(app(X1,X1))
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,<A HREF="#c_0_37">c_0_37</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,plain,
    ( skaf82(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_9]),c_0_10]),c_0_11])])]),c_0_25]),<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( cons(hd(sk4),tl(sk4)) = sk4
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_75]),c_0_56]),c_0_6])])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( sk4 = nil
    | ssList(tl(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_79">c_0_79</A>,c_0_80]),c_0_56]),c_0_6])])).

<A NAME="c_0_86"></A>cnf(c_0_86,plain,
    ( cons(X1,X2) != app(nil,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_87"></A>cnf(c_0_87,plain,
    ( app(nil,skaf45(X1,nil)) = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_40]),c_0_6])])).

<A NAME="c_0_88"></A>cnf(c_0_88,plain,
    ( nil = X1
    | ssItem(X2)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,plain,
    ( ssList(cons(X1,app(X2,X3)))
    | ~ ssItem(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_29]),<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_90"></A>cnf(c_0_90,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_41])]),c_0_6])])).

<A NAME="c_0_91"></A>cnf(c_0_91,negated_conjecture,
    ( skaf82(sk4) = tl(sk4)
    | sk4 = nil
    | ~ ssItem(hd(sk4)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_84]),<A HREF="#c_0_85">c_0_85</A>])).

<A NAME="c_0_92"></A>cnf(c_0_92,plain,
    ( cons(X1,skaf45(X2,nil)) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>,c_0_87]),c_0_42])])).

<A NAME="c_0_93"></A>cnf(c_0_93,plain,
    ( skaf45(X1,X2) = nil
    | ssItem(X3) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>,<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_94"></A>cnf(c_0_94,plain,
    ( ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>,c_0_90]),c_0_6])])).

<A NAME="c_0_95"></A>cnf(c_0_95,plain,
    ( frontsegP(cons(X1,X2),cons(X1,X3))
    | ~ frontsegP(X2,X3)
    | ~ ssItem(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>])).

<A NAME="c_0_96"></A>cnf(c_0_96,negated_conjecture,
    ( skaf82(sk4) = tl(sk4)
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_44]),c_0_56])])).

<A NAME="c_0_97"></A>cnf(c_0_97,plain,
    ( ssItem(X1)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_93])]),<A HREF="#c_0_94">c_0_94</A>])).

<A NAME="c_0_98"></A>cnf(c_0_98,plain,
    ( frontsegP(cons(X1,X2),cons(X1,nil))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_40]),c_0_6])])).

<A NAME="c_0_99"></A>cnf(c_0_99,negated_conjecture,
    ( cons(skaf83(sk4),tl(sk4)) = sk4
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_96]),c_0_56])])).

<A NAME="c_0_100"></A>cnf(c_0_100,plain,
    ( ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,<A HREF="#c_0_45">c_0_45</A>])).

<A NAME="c_0_101"></A>cnf(c_0_101,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_102"></A>cnf(c_0_102,negated_conjecture,
    ( sk4 = nil
    | frontsegP(sk4,cons(skaf83(sk4),nil)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_98">c_0_98</A>,c_0_99]),c_0_10])]),<A HREF="#c_0_85">c_0_85</A>])).

<A NAME="c_0_103"></A>cnf(c_0_103,plain,
    ( ssList(cons(X1,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>,c_0_100])])).

<A NAME="c_0_104"></A>cnf(c_0_104,negated_conjecture,
    ( ~ neq(nil,X1)
    | ~ segmentP(sk4,X1)
    | ~ ssList(X1)
    | ~ strictorderedP(X1) ),
    inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_54]),c_0_54]),<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_105"></A>cnf(c_0_105,plain,
    ( segmentP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_101">c_0_101</A>,c_0_36]),c_0_6])]),<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_106"></A>cnf(c_0_106,negated_conjecture,
    ( app(cons(skaf83(sk4),nil),skaf45(sk4,cons(skaf83(sk4),nil))) = sk4
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_102]),c_0_56])]),c_0_103])])).

<A NAME="c_0_107"></A>cnf(c_0_107,plain,
    ( cons(X1,nil) != nil
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_49]),c_0_6])])).

<A NAME="c_0_108"></A>cnf(c_0_108,negated_conjecture,
    ( X1 = nil
    | ~ segmentP(sk4,X1)
    | ~ ssList(X1)
    | ~ strictorderedP(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>,c_0_5]),c_0_6])])).

<A NAME="c_0_109"></A>cnf(c_0_109,negated_conjecture,
    ( sk4 = nil
    | segmentP(sk4,cons(skaf83(sk4),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_105">c_0_105</A>,c_0_106]),c_0_42]),c_0_103])])).

<A NAME="c_0_110"></A>cnf(c_0_110,plain,
    ( strictorderedP(cons(X1,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_100])])).

<A NAME="c_0_111"></A>cnf(c_0_111,plain,
    ( cons(X1,nil) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>,c_0_100])])).

<A NAME="c_0_112"></A>cnf(c_0_112,negated_conjecture,
    ( neq(sk4,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_113"></A>cnf(c_0_113,negated_conjecture,
    ( sk4 = nil ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>,c_0_109]),c_0_103]),c_0_110])]),<A HREF="#c_0_111">c_0_111</A>])).

<A NAME="c_0_114"></A>cnf(c_0_114,negated_conjecture,
    ( ~ neq(nil,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_113]),<A HREF="#c_0_114">c_0_114</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC125-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n169.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:35:38 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 2.21/2.40  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 2.21/2.40  # and selection function SelectCQArNTEqFirst.
% 2.21/2.40  #
% 2.21/2.40  # Preprocessing time       : 0.010 s
% 2.21/2.40  # Presaturation interreduction done
% 2.21/2.40  
% 2.21/2.40  # Proof found!
% 2.21/2.40  # SZS status Unsatisfiable
% 2.21/2.40  # SZS output start CNFRefutation
% 2.21/2.40  cnf(c_0_0, negated_conjecture, (~neq(sk1,nil)|~segmentP(sk2,sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_11)).
% 2.21/2.40  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 2.21/2.40  cnf(c_0_2, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 2.21/2.40  cnf(c_0_3, negated_conjecture, (segmentP(sk4,sk3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
% 2.21/2.40  cnf(c_0_4, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 2.21/2.40  cnf(c_0_5, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause100)).
% 2.21/2.40  cnf(c_0_6, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 2.21/2.40  cnf(c_0_7, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
% 2.21/2.40  cnf(c_0_8, axiom, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause99)).
% 2.21/2.40  cnf(c_0_9, axiom, (cons(skaf83(X1),skaf82(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause109)).
% 2.21/2.40  cnf(c_0_10, axiom, (ssItem(skaf83(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause12)).
% 2.21/2.40  cnf(c_0_11, axiom, (ssList(skaf82(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause13)).
% 2.21/2.40  cnf(c_0_12, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause96)).
% 2.21/2.40  cnf(c_0_13, axiom, (app(app(skaf47(X1,X2),X2),skaf48(X2,X1))=X1|~segmentP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause168)).
% 2.21/2.40  cnf(c_0_14, axiom, (rearsegP(app(X3,X1),X2)|~rearsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause136)).
% 2.21/2.40  cnf(c_0_15, axiom, (rearsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause58)).
% 2.21/2.40  cnf(c_0_16, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause73)).
% 2.21/2.40  cnf(c_0_17, axiom, (ssList(skaf47(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause49)).
% 2.21/2.40  cnf(c_0_18, axiom, (nil=X2|tl(app(X2,X1))=app(tl(X2),X1)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause133)).
% 2.21/2.40  cnf(c_0_19, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
% 2.21/2.40  cnf(c_0_20, axiom, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause118)).
% 2.21/2.40  cnf(c_0_21, axiom, (ssList(skaf48(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause48)).
% 2.21/2.40  cnf(c_0_22, axiom, (ssList(tl(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause75)).
% 2.21/2.40  cnf(c_0_23, axiom, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause138)).
% 2.21/2.40  cnf(c_0_24, axiom, (app(skaf46(X1,X2),X2)=X1|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause131)).
% 2.21/2.40  cnf(c_0_25, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause86)).
% 2.21/2.40  cnf(c_0_26, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause98)).
% 2.21/2.40  cnf(c_0_27, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause179)).
% 2.21/2.40  cnf(c_0_28, axiom, (ssList(skaf46(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause50)).
% 2.21/2.40  cnf(c_0_29, axiom, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause157)).
% 2.21/2.40  cnf(c_0_30, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause169)).
% 2.21/2.40  cnf(c_0_31, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause149)).
% 2.21/2.40  cnf(c_0_32, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause52)).
% 2.21/2.40  cnf(c_0_33, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause53)).
% 2.21/2.40  cnf(c_0_34, axiom, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause104)).
% 2.21/2.40  cnf(c_0_35, axiom, (nil=X2|hd(app(X2,X1))=hd(X2)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause123)).
% 2.21/2.40  cnf(c_0_36, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause74)).
% 2.21/2.40  cnf(c_0_37, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause72)).
% 2.21/2.40  cnf(c_0_38, axiom, (X4=X2|cons(X1,X2)!=cons(X3,X4)|~ssItem(X3)|~ssItem(X1)|~ssList(X4)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause171)).
% 2.21/2.40  cnf(c_0_39, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause132)).
% 2.21/2.40  cnf(c_0_40, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause60)).
% 2.21/2.40  cnf(c_0_41, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause145)).
% 2.21/2.40  cnf(c_0_42, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause51)).
% 2.21/2.40  cnf(c_0_43, axiom, (frontsegP(cons(X3,X1),cons(X4,X2))|~frontsegP(X1,X2)|X3!=X4|~ssList(X2)|~ssList(X1)|~ssItem(X4)|~ssItem(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause178)).
% 2.21/2.40  cnf(c_0_44, axiom, (ssItem(hd(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause76)).
% 2.21/2.40  cnf(c_0_45, axiom, (ssItem(skaf44(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause47)).
% 2.21/2.40  cnf(c_0_46, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause173)).
% 2.21/2.40  cnf(c_0_47, negated_conjecture, (~ssList(X1)|~neq(sk3,X1)|~segmentP(sk4,X1)|~segmentP(X1,sk3)|~strictorderedP(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
% 2.21/2.40  cnf(c_0_48, axiom, (segmentP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause56)).
% 2.21/2.40  cnf(c_0_49, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause120)).
% 2.21/2.40  cnf(c_0_50, axiom, (strictorderedP(cons(X1,nil))|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause66)).
% 2.21/2.40  cnf(c_0_51, negated_conjecture, (neq(sk2,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 2.21/2.40  cnf(c_0_52, negated_conjecture, (~neq(sk3,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_0, c_0_1]), c_0_2]), c_0_3]), c_0_2])])).
% 2.21/2.40  cnf(c_0_53, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_4, c_0_2])).
% 2.21/2.40  cnf(c_0_54, negated_conjecture, (sk3=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_5]), c_0_53]), c_0_6])])).
% 2.21/2.40  cnf(c_0_55, negated_conjecture, (segmentP(sk4,nil)), inference(rw,[status(thm)],[c_0_3, c_0_54])).
% 2.21/2.40  cnf(c_0_56, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_7, c_0_1])).
% 2.21/2.40  cnf(c_0_57, plain, (nil=X1|skaf82(X1)!=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10]), c_0_11])])).
% 2.21/2.40  cnf(c_0_58, plain, (skaf82(X1)=tl(X1)|nil=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_9]), c_0_10]), c_0_11])])).
% 2.21/2.40  cnf(c_0_59, negated_conjecture, (app(app(skaf47(sk4,nil),nil),skaf48(nil,sk4))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_55]), c_0_6]), c_0_56])])).
% 2.21/2.40  cnf(c_0_60, plain, (nil=X1|tl(X1)!=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
% 2.21/2.40  cnf(c_0_61, plain, (rearsegP(app(X1,X2),nil)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_6])])).
% 2.21/2.40  cnf(c_0_62, negated_conjecture, (app(skaf47(sk4,nil),skaf48(nil,sk4))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_16]), c_0_17])])).
% 2.21/2.40  cnf(c_0_63, plain, (nil=X1|app(tl(X1),X2)!=app(X1,X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_18]), c_0_19]), c_0_20])).
% 2.21/2.40  cnf(c_0_64, negated_conjecture, (rearsegP(sk4,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_17]), c_0_21])])).
% 2.21/2.40  cnf(c_0_65, plain, (nil=X1|tl(X1)!=app(X1,nil)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_16]), c_0_6])]), c_0_22])).
% 2.21/2.40  cnf(c_0_66, plain, (memberP(cons(X1,X2),X1)|~ssItem(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_23])).
% 2.21/2.40  cnf(c_0_67, negated_conjecture, (app(skaf46(sk4,nil),nil)=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_64]), c_0_6]), c_0_56])])).
% 2.21/2.40  cnf(c_0_68, plain, (app(cons(X1,X2),nil)!=X2|~ssItem(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_12]), c_0_25]), c_0_26])).
% 2.21/2.40  cnf(c_0_69, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_27])).
% 2.21/2.40  cnf(c_0_70, plain, (nil=X1|memberP(X1,skaf83(X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_9]), c_0_10]), c_0_11])])).
% 2.21/2.40  cnf(c_0_71, negated_conjecture, (skaf46(sk4,nil)=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_67]), c_0_28])])).
% 2.21/2.40  cnf(c_0_72, plain, (cons(X1,app(X2,nil))!=X2|~ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_29]), c_0_6])])).
% 2.21/2.40  cnf(c_0_73, plain, (nil=X1|~ssList(app(app(X2,cons(skaf83(X1),X3)),X1))|~ssList(X3)|~ssList(X2)|~ssList(X1)|~duplicatefreeP(app(app(X2,cons(skaf83(X1),X3)),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_9]), c_0_10]), c_0_11])])).
% 2.21/2.40  cnf(c_0_74, plain, (app(skaf42(X1,skaf83(X1)),cons(skaf83(X1),skaf43(skaf83(X1),X1)))=X1|nil=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_70]), c_0_10])])).
% 2.21/2.40  cnf(c_0_75, negated_conjecture, (app(sk4,nil)=sk4), inference(rw,[status(thm)],[c_0_67, c_0_71])).
% 2.21/2.40  cnf(c_0_76, plain, (cons(X1,app(X2,app(X3,nil)))!=app(X2,X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_31]), c_0_6])]), c_0_19])).
% 2.21/2.40  cnf(c_0_77, plain, (nil=X1|~ssList(app(X1,X1))|~ssList(X1)|~duplicatefreeP(app(X1,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_32]), c_0_33])])).
% 2.21/2.40  cnf(c_0_78, plain, (cons(hd(X1),tl(app(X1,X2)))=app(X1,X2)|nil=X1|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_19]), c_0_20])).
% 2.21/2.40  cnf(c_0_79, plain, (nil=X1|ssList(app(tl(X1),X2))|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_18]), c_0_19]), c_0_20])).
% 2.21/2.40  cnf(c_0_80, negated_conjecture, (app(tl(sk4),nil)=tl(sk4)|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_75]), c_0_56]), c_0_6])])).
% 2.21/2.40  cnf(c_0_81, plain, (cons(X1,app(X2,nil))!=app(nil,X2)|~ssItem(X1)|~ssList(app(X2,nil))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_36]), c_0_6])])).
% 2.21/2.40  cnf(c_0_82, plain, (nil=X1|ssItem(X2)|~ssList(app(X1,X1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_77, c_0_37])).
% 2.21/2.40  cnf(c_0_83, plain, (skaf82(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_9]), c_0_10]), c_0_11])])]), c_0_25]), c_0_26])).
% 2.21/2.40  cnf(c_0_84, negated_conjecture, (cons(hd(sk4),tl(sk4))=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_75]), c_0_56]), c_0_6])])).
% 2.21/2.40  cnf(c_0_85, negated_conjecture, (sk4=nil|ssList(tl(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_56]), c_0_6])])).
% 2.21/2.40  cnf(c_0_86, plain, (cons(X1,X2)!=app(nil,X2)|~ssItem(X1)|~ssList(X2)), inference(spm,[status(thm)],[c_0_81, c_0_16])).
% 2.21/2.40  cnf(c_0_87, plain, (app(nil,skaf45(X1,nil))=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_6])])).
% 2.21/2.40  cnf(c_0_88, plain, (nil=X1|ssItem(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_82, c_0_19])).
% 2.21/2.40  cnf(c_0_89, plain, (ssList(cons(X1,app(X2,X3)))|~ssItem(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_29]), c_0_25])).
% 2.21/2.40  cnf(c_0_90, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_41])]), c_0_6])])).
% 2.21/2.40  cnf(c_0_91, negated_conjecture, (skaf82(sk4)=tl(sk4)|sk4=nil|~ssItem(hd(sk4))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_85])).
% 2.21/2.40  cnf(c_0_92, plain, (cons(X1,skaf45(X2,nil))!=X2|~ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_42])])).
% 2.21/2.40  cnf(c_0_93, plain, (skaf45(X1,X2)=nil|ssItem(X3)), inference(spm,[status(thm)],[c_0_88, c_0_42])).
% 2.21/2.40  cnf(c_0_94, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_6])])).
% 2.21/2.40  cnf(c_0_95, plain, (frontsegP(cons(X1,X2),cons(X1,X3))|~frontsegP(X2,X3)|~ssItem(X1)|~ssList(X3)|~ssList(X2)), inference(er,[status(thm)],[c_0_43])).
% 2.21/2.40  cnf(c_0_96, negated_conjecture, (skaf82(sk4)=tl(sk4)|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_44]), c_0_56])])).
% 2.21/2.40  cnf(c_0_97, plain, (ssItem(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93])]), c_0_94])).
% 2.21/2.40  cnf(c_0_98, plain, (frontsegP(cons(X1,X2),cons(X1,nil))|~ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_40]), c_0_6])])).
% 2.21/2.40  cnf(c_0_99, negated_conjecture, (cons(skaf83(sk4),tl(sk4))=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_96]), c_0_56])])).
% 2.21/2.40  cnf(c_0_100, plain, (ssItem(X1)), inference(spm,[status(thm)],[c_0_97, c_0_45])).
% 2.21/2.40  cnf(c_0_101, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(er,[status(thm)],[c_0_46])).
% 2.21/2.40  cnf(c_0_102, negated_conjecture, (sk4=nil|frontsegP(sk4,cons(skaf83(sk4),nil))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_10])]), c_0_85])).
% 2.21/2.40  cnf(c_0_103, plain, (ssList(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_100])])).
% 2.21/2.40  cnf(c_0_104, negated_conjecture, (~neq(nil,X1)|~segmentP(sk4,X1)|~ssList(X1)|~strictorderedP(X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_54]), c_0_54]), c_0_48])).
% 2.21/2.40  cnf(c_0_105, plain, (segmentP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_36]), c_0_6])]), c_0_19])).
% 2.21/2.40  cnf(c_0_106, negated_conjecture, (app(cons(skaf83(sk4),nil),skaf45(sk4,cons(skaf83(sk4),nil)))=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_102]), c_0_56])]), c_0_103])])).
% 2.21/2.40  cnf(c_0_107, plain, (cons(X1,nil)!=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_49]), c_0_6])])).
% 2.21/2.40  cnf(c_0_108, negated_conjecture, (X1=nil|~segmentP(sk4,X1)|~ssList(X1)|~strictorderedP(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_5]), c_0_6])])).
% 2.21/2.40  cnf(c_0_109, negated_conjecture, (sk4=nil|segmentP(sk4,cons(skaf83(sk4),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_42]), c_0_103])])).
% 2.21/2.40  cnf(c_0_110, plain, (strictorderedP(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_100])])).
% 2.21/2.40  cnf(c_0_111, plain, (cons(X1,nil)!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_100])])).
% 2.21/2.40  cnf(c_0_112, negated_conjecture, (neq(sk4,nil)), inference(rw,[status(thm)],[c_0_51, c_0_1])).
% 2.21/2.40  cnf(c_0_113, negated_conjecture, (sk4=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_103]), c_0_110])]), c_0_111])).
% 2.21/2.40  cnf(c_0_114, negated_conjecture, (~neq(nil,nil)), inference(rw,[status(thm)],[c_0_52, c_0_54])).
% 2.21/2.40  cnf(c_0_115, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_113]), c_0_114]), ['proof']).
% 2.21/2.40  # SZS output end CNFRefutation
% 2.21/2.40  # Training examples: 0 positive, 0 negative
% 2.21/2.40  
% 2.21/2.40  # -------------------------------------------------
% 2.21/2.40  # User time                : 2.129 s
% 2.21/2.40  # System time              : 0.045 s
% 2.21/2.40  # Total time               : 2.174 s
% 2.21/2.40  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
