<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC045+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC045+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC045+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC045%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC045+1.p">SWC045+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n115.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:25 EDT 2016

% Result     : Theorem 162.80s
% Output     : CNFRefutation 162.80s
% Verified   : 
% Statistics : Number of formulae       :  171 ( 736 expanded)
%              Number of clauses        :  114 ( 344 expanded)
%              Number of leaves         :   28 ( 225 expanded)
%              Depth                    :   30
%              Number of atoms          :  690 (2858 expanded)
%              Number of equality atoms :  154 ( 629 expanded)
%              Maximal formula depth    :   21 (   5 average)
%              Maximal clause size      :   18 (   3 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( singletonP(X1)
      &lt;=> ? [X2] :
            ( ssItem(X2)
            & cons(X2,nil) = X1 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax4)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax86)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ~ singletonP(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax39)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ? [X2] :
            ( ssItem(X2)
            & hd(X1) = X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax75)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( rearsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax6)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( frontsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X2,X3) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax5)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( rearsegP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax52)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( segmentP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax7)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( ( frontsegP(X1,X2)
                  & frontsegP(X2,X3) )
               => frontsegP(X1,X3) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax40)).

<A NAME="c_0_17"></A>fof(c_0_17,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,nil) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax45)).

<A NAME="c_0_18"></A>fof(c_0_18,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( segmentP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax58)).

<A NAME="c_0_19"></A>fof(c_0_19,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( ( segmentP(X1,X2)
                  & segmentP(X2,X3) )
               => segmentP(X1,X3) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax53)).

<A NAME="c_0_20"></A>fof(c_0_20,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( frontsegP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax46)).

<A NAME="c_0_21"></A>fof(c_0_21,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax38)).

<A NAME="c_0_22"></A>fof(c_0_22,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(app(X2,X3),X1)
              &lt;=> ( memberP(X2,X1)
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax36)).

<A NAME="c_0_23"></A>fof(c_0_23,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(cons(X2,X3),X1)
              &lt;=> ( X1 = X2
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax37)).

<A NAME="c_0_24"></A>fof(c_0_24,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => hd(cons(X2,X1)) = X2 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax23)).

<A NAME="c_0_25"></A>fof(c_0_25,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( nil != X2
                    | X2 != X4
                    | X1 != X3
                    | nil = X1
                    | ? [X5] :
                        ( ssItem(X5)
                        & ( ( ~ memberP(X3,X5)
                            & ? [X6] :
                                ( ssList(X6)
                                & segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) ) )
                          | ( ! [X6] :
                                ( ssList(X6)
                               => ~ segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) )
                            & memberP(X3,X5) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_26"></A>fof(c_0_26,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

<A NAME="c_0_27"></A>fof(c_0_27,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil = X1
        | ? [X2] :
            ( ssList(X2)
            & ? [X3] :
                ( ssItem(X3)
                & cons(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax20)).

<A NAME="c_0_28"></A>fof(c_0_28,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])).

<A NAME="c_0_29"></A>fof(c_0_29,plain,(
    ! [X3,X5] :
      ( ( ssItem(esk5_1(X3))
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( cons(esk5_1(X3),nil) = X3
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( ~ ssItem(X5)
        | cons(X5,nil) != X3
        | singletonP(X3)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

<A NAME="c_0_30"></A>fof(c_0_30,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_31"></A>fof(c_0_31,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

<A NAME="c_0_32"></A>fof(c_0_32,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_33"></A>cnf(c_0_33,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,plain,
    ( cons(esk5_1(X1),nil) = X1
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( ssItem(esk5_1(X1))
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,plain,
    ( singletonP(X1)
    | ~ ssList(X1)
    | cons(X2,nil) != X1
    | ~ ssItem(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_38"></A>fof(c_0_38,plain,(
    ~ singletonP(nil) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_39"></A>fof(c_0_39,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).

<A NAME="c_0_40"></A>cnf(c_0_40,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,plain,
    ( tl(app(X1,X2)) = app(tl(X1),X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,plain,
    ( tl(X1) = nil
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_34]),c_0_35])]),<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( ~ singletonP(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_46"></A>fof(c_0_46,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).

<A NAME="c_0_47"></A>fof(c_0_47,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_48"></A>cnf(c_0_48,plain,
    ( nil = X2
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_49"></A>cnf(c_0_49,plain,
    ( app(X1,app(nil,X2)) = app(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_41]),c_0_35])])).

<A NAME="c_0_50"></A>cnf(c_0_50,plain,
    ( app(tl(X1),nil) = tl(X1)
    | nil = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_41]),c_0_35])])).

<A NAME="c_0_51"></A>cnf(c_0_51,plain,
    ( tl(cons(X1,nil)) = nil
    | ~ ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,plain,
    ( cons(X1,nil) != nil
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_37]),c_0_35])])).

<A NAME="c_0_53"></A>fof(c_0_53,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).

<A NAME="c_0_54"></A>cnf(c_0_54,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_55"></A>cnf(c_0_55,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( app(nil,X1) = nil
    | app(X2,X1) != nil
    | ~ ssList(app(nil,X1))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,plain,
    ( app(nil,nil) = nil
    | ~ ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_51]),<A HREF="#c_0_52">c_0_52</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>])).

<A NAME="c_0_59"></A>fof(c_0_59,plain,(
    ! [X3] :
      ( ( ssItem(esk46_1(X3))
        | nil = X3
        | ~ ssList(X3) )
      & ( hd(X3) = esk46_1(X3)
        | nil = X3
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).

<A NAME="c_0_60"></A>cnf(c_0_60,plain,
    ( app(nil,app(X1,X2)) = app(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_54]),c_0_35])])).

<A NAME="c_0_61"></A>cnf(c_0_61,plain,
    ( app(X1,app(X2,nil)) = app(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_40]),c_0_35])]),<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_62"></A>cnf(c_0_62,plain,
    ( app(nil,X1) = nil
    | app(X2,X1) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_55]),c_0_35])])).

<A NAME="c_0_63"></A>cnf(c_0_63,plain,
    ( app(nil,nil) = nil
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,c_0_58]),c_0_35])])).

<A NAME="c_0_64"></A>cnf(c_0_64,plain,
    ( nil = X1
    | ssItem(esk46_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( app(nil,X1) = app(X1,nil)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>,c_0_61]),c_0_35])])).

<A NAME="c_0_66"></A>cnf(c_0_66,plain,
    ( app(nil,X1) = nil
    | X1 != nil
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_54]),c_0_35])])).

<A NAME="c_0_67"></A>cnf(c_0_67,plain,
    ( app(nil,nil) = nil
    | nil = X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>,<A HREF="#c_0_64">c_0_64</A>])).

<A NAME="c_0_68"></A>fof(c_0_68,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk7_2(X4,X5))
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(esk7_2(X4,X5),X5) = X4
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X7,X5) != X4
        | rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).

<A NAME="c_0_69"></A>fof(c_0_69,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk6_2(X4,X5))
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(X5,esk6_2(X4,X5)) = X4
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X5,X7) != X4
        | frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])).

<A NAME="c_0_70"></A>cnf(c_0_70,plain,
    ( ssList(app(nil,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_65]),c_0_35])])).

<A NAME="c_0_71"></A>cnf(c_0_71,plain,
    ( app(nil,app(X1,X2)) = app(nil,X2)
    | X1 != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_66]),c_0_35])])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( app(nil,nil) = nil
    | app(X1,X2) = nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_73"></A>fof(c_0_73,plain,(
    ! [X2] :
      ( ( ~ rearsegP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | rearsegP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).

<A NAME="c_0_74"></A>cnf(c_0_74,plain,
    ( rearsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X3,X2) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_75"></A>cnf(c_0_75,plain,
    ( app(X2,esk6_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,plain,
    ( ssList(esk6_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_77"></A>fof(c_0_77,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk8_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ssList(esk9_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(app(X9,X6),X10) != X5
        | segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])])).

<A NAME="c_0_78"></A>cnf(c_0_78,plain,
    ( app(X1,app(X2,app(nil,X3))) = app(app(X1,X2),X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_49]),c_0_55]),<A HREF="#c_0_70">c_0_70</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,plain,
    ( app(nil,app(X1,nil)) = app(nil,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>,c_0_65]),c_0_35])])).

<A NAME="c_0_80"></A>cnf(c_0_80,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(ef,[status(thm)],[c_0_72]),c_0_35])])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ rearsegP(nil,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,plain,
    ( rearsegP(X1,esk6_2(X1,X2))
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,c_0_75])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_83"></A>fof(c_0_83,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ frontsegP(X4,X5)
      | ~ frontsegP(X5,X6)
      | frontsegP(X4,X6) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).

<A NAME="c_0_84"></A>fof(c_0_84,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).

<A NAME="c_0_85"></A>fof(c_0_85,plain,(
    ! [X2] :
      ( ( ~ segmentP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | segmentP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).

<A NAME="c_0_86"></A>fof(c_0_86,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ segmentP(X4,X5)
      | ~ segmentP(X5,X6)
      | segmentP(X4,X6) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).

<A NAME="c_0_87"></A>cnf(c_0_87,plain,
    ( segmentP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(app(X3,X2),X4) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,plain,
    ( app(app(X1,nil),nil) = app(X1,nil)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_79]),c_0_80]),c_0_35])])).

<A NAME="c_0_89"></A>cnf(c_0_89,plain,
    ( esk6_2(nil,X1) = nil
    | ~ frontsegP(nil,X1)
    | ~ ssList(esk6_2(nil,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_82]),c_0_35])])).

<A NAME="c_0_90"></A>cnf(c_0_90,plain,
    ( frontsegP(X1,X2)
    | ~ frontsegP(X3,X2)
    | ~ frontsegP(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>])).

<A NAME="c_0_91"></A>cnf(c_0_91,plain,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_92"></A>fof(c_0_92,plain,(
    ! [X2] :
      ( ( ~ frontsegP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | frontsegP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).

<A NAME="c_0_93"></A>cnf(c_0_93,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ segmentP(nil,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>])).

<A NAME="c_0_94"></A>cnf(c_0_94,plain,
    ( segmentP(X1,X2)
    | ~ segmentP(X3,X2)
    | ~ segmentP(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>])).

<A NAME="c_0_95"></A>cnf(c_0_95,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>])).

<A NAME="c_0_96"></A>cnf(c_0_96,plain,
    ( app(app(nil,X1),nil) = app(nil,X1)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>,<A HREF="#c_0_65">c_0_65</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,plain,
    ( esk6_2(nil,X1) = nil
    | ~ frontsegP(nil,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>,c_0_76]),c_0_35])])).

<A NAME="c_0_98"></A>cnf(c_0_98,plain,
    ( frontsegP(X1,X2)
    | ~ frontsegP(nil,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,c_0_91]),c_0_35])])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( frontsegP(nil,X1)
    | ~ ssList(X1)
    | nil != X1 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,plain,
    ( nil = X1
    | ~ segmentP(nil,X2)
    | ~ segmentP(X2,X1)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_94]),c_0_35])])).

<A NAME="c_0_101"></A>cnf(c_0_101,plain,
    ( segmentP(app(nil,X1),X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_96]),c_0_35])]),<A HREF="#c_0_70">c_0_70</A>])).

<A NAME="c_0_102"></A>cnf(c_0_102,plain,
    ( app(X1,nil) = nil
    | ~ frontsegP(nil,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_97]),c_0_35])])).

<A NAME="c_0_103"></A>cnf(c_0_103,plain,
    ( frontsegP(X1,X2)
    | nil != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_98">c_0_98</A>,<A HREF="#c_0_99">c_0_99</A>])).

<A NAME="c_0_104"></A>fof(c_0_104,plain,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_105"></A>cnf(c_0_105,plain,
    ( nil = X1
    | ~ segmentP(nil,app(nil,X1))
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>,c_0_101]),<A HREF="#c_0_70">c_0_70</A>])).

<A NAME="c_0_106"></A>cnf(c_0_106,plain,
    ( ssList(app(X1,nil))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_65]),c_0_35])])).

<A NAME="c_0_107"></A>cnf(c_0_107,plain,
    ( app(X1,nil) = nil
    | nil != X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>,c_0_103]),c_0_35])])).

<A NAME="c_0_108"></A>fof(c_0_108,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(app(X5,X6),X4)
        | memberP(X5,X4)
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X5,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).

<A NAME="c_0_109"></A>fof(c_0_109,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | ~ memberP(nil,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_104])])).

<A NAME="c_0_110"></A>cnf(c_0_110,plain,
    ( app(X1,nil) = nil
    | ~ segmentP(nil,app(nil,X1))
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_105">c_0_105</A>,c_0_61]),c_0_35])]),<A HREF="#c_0_106">c_0_106</A>])).

<A NAME="c_0_111"></A>cnf(c_0_111,plain,
    ( app(X1,app(nil,X2)) = app(nil,X2)
    | nil != X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_107]),c_0_35])])).

<A NAME="c_0_112"></A>cnf(c_0_112,plain,
    ( memberP(app(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(X3,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>])).

<A NAME="c_0_113"></A>cnf(c_0_113,plain,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>])).

<A NAME="c_0_114"></A>cnf(c_0_114,plain,
    ( app(app(X1,nil),nil) = nil
    | ~ segmentP(nil,app(nil,X1))
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_110">c_0_110</A>,c_0_61]),c_0_35])]),<A HREF="#c_0_106">c_0_106</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,plain,
    ( app(X1,X2) = X2
    | nil != X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_111">c_0_111</A>,<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_116"></A>cnf(c_0_116,plain,
    ( X1 != nil
    | ~ memberP(X1,X2)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_66]),c_0_35])]),<A HREF="#c_0_113">c_0_113</A>])).

<A NAME="c_0_117"></A>cnf(c_0_117,plain,
    ( app(app(X1,nil),nil) = nil
    | ~ segmentP(nil,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_114">c_0_114</A>,c_0_115]),c_0_35])])).

<A NAME="c_0_118"></A>fof(c_0_118,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(cons(X5,X6),X4)
        | X4 = X5
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( X4 != X5
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).

<A NAME="c_0_119"></A>cnf(c_0_119,plain,
    ( app(X1,X2) != nil
    | ~ memberP(X2,X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssItem(X3) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_116">c_0_116</A>,c_0_112]),<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_120"></A>cnf(c_0_120,plain,
    ( memberP(app(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(X2,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>])).

<A NAME="c_0_121"></A>cnf(c_0_121,plain,
    ( app(app(nil,X1),nil) = nil
    | ~ segmentP(nil,X1)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_117">c_0_117</A>,<A HREF="#c_0_65">c_0_65</A>])).

<A NAME="c_0_122"></A>cnf(c_0_122,plain,
    ( memberP(cons(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | X1 != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_118">c_0_118</A>])).

<A NAME="c_0_123"></A>fof(c_0_123,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | hd(cons(X4,X3)) = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).

<A NAME="c_0_124"></A>fof(c_0_124,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( nil != X2
                      | X2 != X4
                      | X1 != X3
                      | nil = X1
                      | ? [X5] :
                          ( ssItem(X5)
                          & ( ( ~ memberP(X3,X5)
                              & ? [X6] :
                                  ( ssList(X6)
                                  & segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) ) )
                            | ( ! [X6] :
                                  ( ssList(X6)
                                 => ~ segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) )
                              & memberP(X3,X5) ) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_25])])).

<A NAME="c_0_125"></A>cnf(c_0_125,plain,
    ( app(X1,app(X2,X3)) != nil
    | ~ memberP(X2,X4)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X4) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>,c_0_120]),<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_126"></A>cnf(c_0_126,plain,
    ( app(nil,app(X1,nil)) = nil
    | ~ segmentP(nil,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_121]),c_0_35])])).

<A NAME="c_0_127"></A>cnf(c_0_127,plain,
    ( app(esk7_2(X1,X2),X2) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_128"></A>cnf(c_0_128,plain,
    ( ssList(esk7_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_129"></A>cnf(c_0_129,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_122">c_0_122</A>])).

<A NAME="c_0_130"></A>cnf(c_0_130,plain,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>])).

<A NAME="c_0_131"></A>fof(c_0_131,negated_conjecture,(
    ! [X11,X12] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & nil = esk49_0
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & nil != esk48_0
      & ( memberP(esk50_0,X11)
        | ~ ssList(X12)
        | ~ segmentP(esk51_0,app(app(cons(X11,nil),X12),cons(X11,nil)))
        | ~ ssItem(X11) )
      & ( ssList(esk52_1(X11))
        | ~ memberP(esk50_0,X11)
        | ~ ssItem(X11) )
      & ( segmentP(esk51_0,app(app(cons(X11,nil),esk52_1(X11)),cons(X11,nil)))
        | ~ memberP(esk50_0,X11)
        | ~ ssItem(X11) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_124])])])])])).

<A NAME="c_0_132"></A>cnf(c_0_132,plain,
    ( ~ segmentP(nil,X1)
    | ~ memberP(X1,X2)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_125">c_0_125</A>,c_0_126]),c_0_35])])).

<A NAME="c_0_133"></A>cnf(c_0_133,plain,
    ( memberP(X1,X2)
    | ~ rearsegP(X1,X3)
    | ~ memberP(X3,X2)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_127]),<A HREF="#c_0_128">c_0_128</A>])).

<A NAME="c_0_134"></A>cnf(c_0_134,plain,
    ( memberP(X1,esk5_1(X1))
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_129">c_0_129</A>,c_0_34]),c_0_35])]),<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_135"></A>cnf(c_0_135,plain,
    ( esk5_1(X1) = hd(X1)
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_130">c_0_130</A>,c_0_34]),c_0_35])]),<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_136"></A>cnf(c_0_136,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>])).

<A NAME="c_0_137"></A>cnf(c_0_137,negated_conjecture,
    ( nil = esk49_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>])).

<A NAME="c_0_138"></A>cnf(c_0_138,plain,
    ( ~ segmentP(nil,X1)
    | ~ rearsegP(X1,X2)
    | ~ memberP(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssItem(X3) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_132">c_0_132</A>,<A HREF="#c_0_133">c_0_133</A>])).

<A NAME="c_0_139"></A>cnf(c_0_139,plain,
    ( memberP(X1,hd(X1))
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_134">c_0_134</A>,<A HREF="#c_0_135">c_0_135</A>])).

<A NAME="c_0_140"></A>cnf(c_0_140,plain,
    ( ssItem(hd(X1))
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,<A HREF="#c_0_135">c_0_135</A>])).

<A NAME="c_0_141"></A>cnf(c_0_141,negated_conjecture,
    ( segmentP(esk51_0,app(app(cons(X1,nil),esk52_1(X1)),cons(X1,nil)))
    | ~ ssItem(X1)
    | ~ memberP(esk50_0,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>])).

<A NAME="c_0_142"></A>cnf(c_0_142,negated_conjecture,
    ( esk51_0 = nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_136">c_0_136</A>,<A HREF="#c_0_137">c_0_137</A>])).

<A NAME="c_0_143"></A>cnf(c_0_143,plain,
    ( ~ segmentP(nil,X1)
    | ~ rearsegP(X1,X2)
    | ~ singletonP(X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_138">c_0_138</A>,c_0_139]),<A HREF="#c_0_140">c_0_140</A>])).

<A NAME="c_0_144"></A>cnf(c_0_144,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_74]),<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_145"></A>cnf(c_0_145,negated_conjecture,
    ( segmentP(nil,app(app(cons(X1,nil),esk52_1(X1)),cons(X1,nil)))
    | ~ memberP(esk50_0,X1)
    | ~ ssItem(X1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_141">c_0_141</A>,<A HREF="#c_0_142">c_0_142</A>])).

<A NAME="c_0_146"></A>cnf(c_0_146,plain,
    ( ~ segmentP(nil,app(X1,X2))
    | ~ singletonP(X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_143">c_0_143</A>,c_0_144]),<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_147"></A>cnf(c_0_147,negated_conjecture,
    ( segmentP(nil,app(app(X1,esk52_1(esk5_1(X1))),X1))
    | ~ singletonP(X1)
    | ~ memberP(esk50_0,esk5_1(X1))
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_145">c_0_145</A>,c_0_34]),<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_148"></A>cnf(c_0_148,negated_conjecture,
    ( ~ singletonP(X1)
    | ~ memberP(esk50_0,esk5_1(X1))
    | ~ ssList(app(X1,esk52_1(esk5_1(X1))))
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_146">c_0_146</A>,<A HREF="#c_0_147">c_0_147</A>])).

<A NAME="c_0_149"></A>fof(c_0_149,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssItem(X6)
      | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).

<A NAME="c_0_150"></A>cnf(c_0_150,negated_conjecture,
    ( ~ singletonP(X1)
    | ~ memberP(esk50_0,esk5_1(X1))
    | ~ ssList(esk52_1(esk5_1(X1)))
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_148">c_0_148</A>,<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_151"></A>cnf(c_0_151,negated_conjecture,
    ( ssList(esk52_1(X1))
    | ~ ssItem(X1)
    | ~ memberP(esk50_0,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>])).

<A NAME="c_0_152"></A>cnf(c_0_152,plain,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_149">c_0_149</A>])).

<A NAME="c_0_153"></A>cnf(c_0_153,negated_conjecture,
    ( ~ singletonP(X1)
    | ~ memberP(esk50_0,esk5_1(X1))
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_150">c_0_150</A>,c_0_151]),<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_154"></A>cnf(c_0_154,plain,
    ( ssList(cons(X1,app(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_152]),<A HREF="#c_0_58">c_0_58</A>])).

<A NAME="c_0_155"></A>cnf(c_0_155,negated_conjecture,
    ( ~ singletonP(X1)
    | ~ memberP(esk50_0,hd(X1))
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_153">c_0_153</A>,<A HREF="#c_0_135">c_0_135</A>])).

<A NAME="c_0_156"></A>cnf(c_0_156,plain,
    ( ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_154">c_0_154</A>,c_0_80]),c_0_35])])).

<A NAME="c_0_157"></A>fof(c_0_157,plain,(
    ! [X4] :
      ( ( ssList(esk44_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( ssItem(esk45_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( cons(esk45_1(X4),esk44_1(X4)) = X4
        | nil = X4
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])).

<A NAME="c_0_158"></A>cnf(c_0_158,negated_conjecture,
    ( ~ memberP(esk50_0,hd(cons(X1,nil)))
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_155">c_0_155</A>,c_0_44]),<A HREF="#c_0_156">c_0_156</A>])).

<A NAME="c_0_159"></A>cnf(c_0_159,plain,
    ( nil = X1
    | cons(esk45_1(X1),esk44_1(X1)) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_157">c_0_157</A>])).

<A NAME="c_0_160"></A>cnf(c_0_160,plain,
    ( nil = X1
    | ssItem(esk45_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_157">c_0_157</A>])).

<A NAME="c_0_161"></A>cnf(c_0_161,plain,
    ( nil = X1
    | ssList(esk44_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_157">c_0_157</A>])).

<A NAME="c_0_162"></A>cnf(c_0_162,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>])).

<A NAME="c_0_163"></A>cnf(c_0_163,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>])).

<A NAME="c_0_164"></A>cnf(c_0_164,negated_conjecture,
    ( nil != esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>])).

<A NAME="c_0_165"></A>cnf(c_0_165,negated_conjecture,
    ( ~ memberP(esk50_0,X1)
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_158">c_0_158</A>,c_0_130]),c_0_35])])).

<A NAME="c_0_166"></A>cnf(c_0_166,plain,
    ( nil = X1
    | memberP(X1,esk45_1(X1))
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_129">c_0_129</A>,c_0_159]),c_0_160]),<A HREF="#c_0_161">c_0_161</A>])).

<A NAME="c_0_167"></A>cnf(c_0_167,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_162">c_0_162</A>,<A HREF="#c_0_163">c_0_163</A>])).

<A NAME="c_0_168"></A>cnf(c_0_168,negated_conjecture,
    ( esk50_0 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_164">c_0_164</A>,<A HREF="#c_0_163">c_0_163</A>])).

<A NAME="c_0_169"></A>cnf(c_0_169,negated_conjecture,
    ( ~ ssItem(esk45_1(esk50_0)) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_165">c_0_165</A>,c_0_166]),c_0_167])]),<A HREF="#c_0_168">c_0_168</A>])).

<A NAME="c_0_170"></A>cnf(c_0_170,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_169">c_0_169</A>,c_0_160]),c_0_167])]),<A HREF="#c_0_168">c_0_168</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC045+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n115.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:34:53 CDT 2016
% 0.02/0.22  % CPUTime    : 
% 0.02/0.22  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 151.10/151.24  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 151.10/151.24  # and selection function SelectCQPrecWNTNp.
% 151.10/151.24  #
% 151.10/151.24  # Preprocessing time       : 0.014 s
% 151.10/151.24  # Presaturation interreduction done
% 151.40/151.57  # No success with AutoSched0
% 151.40/151.57  # Trying AutoSched1 for 73 seconds
% 162.80/162.90  # AutoSched1-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S00
% 162.80/162.90  # and selection function NoSelection.
% 162.80/162.90  #
% 162.80/162.90  # Preprocessing time       : 0.014 s
% 162.80/162.90  # Presaturation interreduction done
% 162.80/162.90  
% 162.80/162.90  # Proof found!
% 162.80/162.90  # SZS status Theorem
% 162.80/162.90  # SZS output start CNFRefutation
% 162.80/162.90  fof(c_0_0, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
% 162.80/162.90  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)&lt;=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax4)).
% 162.80/162.90  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 162.80/162.90  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
% 162.80/162.90  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax86)).
% 162.80/162.90  fof(c_0_5, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 162.80/162.90  fof(c_0_6, axiom, (~(singletonP(nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax39)).
% 162.80/162.90  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 162.80/162.90  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
% 162.80/162.90  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 162.80/162.90  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
% 162.80/162.90  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>?[X2]:(ssItem(X2)&hd(X1)=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax75)).
% 162.80/162.90  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax6)).
% 162.80/162.90  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax5)).
% 162.80/162.90  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>(rearsegP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax52)).
% 162.80/162.90  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax7)).
% 162.80/162.90  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>((frontsegP(X1,X2)&frontsegP(X2,X3))=>frontsegP(X1,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax40)).
% 162.80/162.90  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax45)).
% 162.80/162.90  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>(segmentP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax58)).
% 162.80/162.90  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>((segmentP(X1,X2)&segmentP(X2,X3))=>segmentP(X1,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax53)).
% 162.80/162.90  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>(frontsegP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax46)).
% 162.80/162.90  fof(c_0_21, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax38)).
% 162.80/162.90  fof(c_0_22, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(memberP(app(X2,X3),X1)&lt;=>(memberP(X2,X1)|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax36)).
% 162.80/162.90  fof(c_0_23, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)&lt;=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax37)).
% 162.80/162.90  fof(c_0_24, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax23)).
% 162.80/162.90  fof(c_0_25, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((nil!=X2|X2!=X4)|X1!=X3)|nil=X1)|?[X5]:(ssItem(X5)&((~(memberP(X3,X5))&?[X6]:(ssList(X6)&segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil)))))|(![X6]:(ssList(X6)=>~(segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil)))))&memberP(X3,X5)))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 162.80/162.90  fof(c_0_26, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
% 162.80/162.90  fof(c_0_27, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax20)).
% 162.80/162.90  fof(c_0_28, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])).
% 162.80/162.90  fof(c_0_29, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
% 162.80/162.90  fof(c_0_30, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 162.80/162.90  fof(c_0_31, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
% 162.80/162.90  fof(c_0_32, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 162.80/162.90  cnf(c_0_33, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
% 162.80/162.90  cnf(c_0_34, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
% 162.80/162.90  cnf(c_0_35, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 162.80/162.90  cnf(c_0_36, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
% 162.80/162.90  cnf(c_0_37, plain, (singletonP(X1)|~ssList(X1)|cons(X2,nil)!=X1|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
% 162.80/162.90  fof(c_0_38, plain, (~singletonP(nil)), inference(fof_simplification,[status(thm)],[c_0_6])).
% 162.80/162.90  fof(c_0_39, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
% 162.80/162.90  cnf(c_0_40, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
% 162.80/162.90  cnf(c_0_41, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
% 162.80/162.90  cnf(c_0_42, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 162.80/162.90  cnf(c_0_43, plain, (tl(X1)=nil|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])]), c_0_36])).
% 162.80/162.90  cnf(c_0_44, plain, (singletonP(cons(X1,nil))|~ssList(cons(X1,nil))|~ssItem(X1)), inference(er,[status(thm)],[c_0_37])).
% 162.80/162.90  cnf(c_0_45, plain, (~singletonP(nil)), inference(split_conjunct,[status(thm)],[c_0_38])).
% 162.80/162.90  fof(c_0_46, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
% 162.80/162.90  fof(c_0_47, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 162.80/162.90  cnf(c_0_48, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 162.80/162.90  cnf(c_0_49, plain, (app(X1,app(nil,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_35])])).
% 162.80/162.90  cnf(c_0_50, plain, (app(tl(X1),nil)=tl(X1)|nil=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_41]), c_0_35])])).
% 162.80/162.90  cnf(c_0_51, plain, (tl(cons(X1,nil))=nil|~ssList(cons(X1,nil))|~ssItem(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
% 162.80/162.90  cnf(c_0_52, plain, (cons(X1,nil)!=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_37]), c_0_35])])).
% 162.80/162.90  fof(c_0_53, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
% 162.80/162.90  cnf(c_0_54, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
% 162.80/162.90  cnf(c_0_55, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
% 162.80/162.90  cnf(c_0_56, plain, (app(nil,X1)=nil|app(X2,X1)!=nil|~ssList(app(nil,X1))|~ssList(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
% 162.80/162.90  cnf(c_0_57, plain, (app(nil,nil)=nil|~ssList(cons(X1,nil))|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
% 162.80/162.90  cnf(c_0_58, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
% 162.80/162.90  fof(c_0_59, plain, (![X3]:(((ssItem(esk46_1(X3))|nil=X3)|~ssList(X3))&((hd(X3)=esk46_1(X3)|nil=X3)|~ssList(X3)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).
% 162.80/162.90  cnf(c_0_60, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_54]), c_0_35])])).
% 162.80/162.90  cnf(c_0_61, plain, (app(X1,app(X2,nil))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_40]), c_0_35])]), c_0_55])).
% 162.80/162.90  cnf(c_0_62, plain, (app(nil,X1)=nil|app(X2,X1)!=nil|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_55]), c_0_35])])).
% 162.80/162.90  cnf(c_0_63, plain, (app(nil,nil)=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_35])])).
% 162.80/162.90  cnf(c_0_64, plain, (nil=X1|ssItem(esk46_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
% 162.80/162.90  cnf(c_0_65, plain, (app(nil,X1)=app(X1,nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_35])])).
% 162.80/162.90  cnf(c_0_66, plain, (app(nil,X1)=nil|X1!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_54]), c_0_35])])).
% 162.80/162.90  cnf(c_0_67, plain, (app(nil,nil)=nil|nil=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
% 162.80/162.90  fof(c_0_68, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).
% 162.80/162.90  fof(c_0_69, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])).
% 162.80/162.90  cnf(c_0_70, plain, (ssList(app(nil,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_65]), c_0_35])])).
% 162.80/162.90  cnf(c_0_71, plain, (app(nil,app(X1,X2))=app(nil,X2)|X1!=nil|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_66]), c_0_35])])).
% 162.80/162.90  cnf(c_0_72, plain, (app(nil,nil)=nil|app(X1,X2)=nil|~ssList(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_67, c_0_55])).
% 162.80/162.90  fof(c_0_73, plain, (![X2]:(((~rearsegP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|rearsegP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
% 162.80/162.90  cnf(c_0_74, plain, (rearsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X3,X2)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_68])).
% 162.80/162.90  cnf(c_0_75, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
% 162.80/162.90  cnf(c_0_76, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
% 162.80/162.90  fof(c_0_77, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])])).
% 162.80/162.90  cnf(c_0_78, plain, (app(X1,app(X2,app(nil,X3)))=app(app(X1,X2),X3)|~ssList(X2)|~ssList(X1)|~ssList(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_49]), c_0_55]), c_0_70])).
% 162.80/162.90  cnf(c_0_79, plain, (app(nil,app(X1,nil))=app(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_65]), c_0_35])])).
% 162.80/162.90  cnf(c_0_80, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(ef,[status(thm)],[c_0_72]), c_0_35])])).
% 162.80/162.90  cnf(c_0_81, plain, (nil=X1|~ssList(X1)|~rearsegP(nil,X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
% 162.80/162.90  cnf(c_0_82, plain, (rearsegP(X1,esk6_2(X1,X2))|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75])]), c_0_76])).
% 162.80/162.90  fof(c_0_83, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|((~frontsegP(X4,X5)|~frontsegP(X5,X6))|frontsegP(X4,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
% 162.80/162.90  fof(c_0_84, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).
% 162.80/162.90  fof(c_0_85, plain, (![X2]:(((~segmentP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|segmentP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
% 162.80/162.90  fof(c_0_86, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|((~segmentP(X4,X5)|~segmentP(X5,X6))|segmentP(X4,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
% 162.80/162.90  cnf(c_0_87, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
% 162.80/162.90  cnf(c_0_88, plain, (app(app(X1,nil),nil)=app(X1,nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80]), c_0_35])])).
% 162.80/162.90  cnf(c_0_89, plain, (esk6_2(nil,X1)=nil|~frontsegP(nil,X1)|~ssList(esk6_2(nil,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_35])])).
% 162.80/162.90  cnf(c_0_90, plain, (frontsegP(X1,X2)|~frontsegP(X3,X2)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
% 162.80/162.90  cnf(c_0_91, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
% 162.80/162.90  fof(c_0_92, plain, (![X2]:(((~frontsegP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|frontsegP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
% 162.80/162.90  cnf(c_0_93, plain, (nil=X1|~ssList(X1)|~segmentP(nil,X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
% 162.80/162.90  cnf(c_0_94, plain, (segmentP(X1,X2)|~segmentP(X3,X2)|~segmentP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
% 162.80/162.90  cnf(c_0_95, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_87])).
% 162.80/162.90  cnf(c_0_96, plain, (app(app(nil,X1),nil)=app(nil,X1)|~ssList(X1)), inference(spm,[status(thm)],[c_0_88, c_0_65])).
% 162.80/162.90  cnf(c_0_97, plain, (esk6_2(nil,X1)=nil|~frontsegP(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_76]), c_0_35])])).
% 162.80/162.90  cnf(c_0_98, plain, (frontsegP(X1,X2)|~frontsegP(nil,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_35])])).
% 162.80/162.90  cnf(c_0_99, plain, (frontsegP(nil,X1)|~ssList(X1)|nil!=X1), inference(split_conjunct,[status(thm)],[c_0_92])).
% 162.80/162.90  cnf(c_0_100, plain, (nil=X1|~segmentP(nil,X2)|~segmentP(X2,X1)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_35])])).
% 162.80/162.90  cnf(c_0_101, plain, (segmentP(app(nil,X1),X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_35])]), c_0_70])).
% 162.80/162.90  cnf(c_0_102, plain, (app(X1,nil)=nil|~frontsegP(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_97]), c_0_35])])).
% 162.80/162.90  cnf(c_0_103, plain, (frontsegP(X1,X2)|nil!=X2|~ssList(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
% 162.80/162.90  fof(c_0_104, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_21])).
% 162.80/162.90  cnf(c_0_105, plain, (nil=X1|~segmentP(nil,app(nil,X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_70])).
% 162.80/162.90  cnf(c_0_106, plain, (ssList(app(X1,nil))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_65]), c_0_35])])).
% 162.80/162.90  cnf(c_0_107, plain, (app(X1,nil)=nil|nil!=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_35])])).
% 162.80/162.90  fof(c_0_108, plain, (![X4]:![X5]:![X6]:(((((~memberP(app(X5,X6),X4)|(memberP(X5,X4)|memberP(X6,X4)))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&(((((~memberP(X5,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).
% 162.80/162.90  fof(c_0_109, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_104])])).
% 162.80/162.90  cnf(c_0_110, plain, (app(X1,nil)=nil|~segmentP(nil,app(nil,X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_61]), c_0_35])]), c_0_106])).
% 162.80/162.90  cnf(c_0_111, plain, (app(X1,app(nil,X2))=app(nil,X2)|nil!=X1|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_107]), c_0_35])])).
% 162.80/162.90  cnf(c_0_112, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
% 162.80/162.90  cnf(c_0_113, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
% 162.80/162.90  cnf(c_0_114, plain, (app(app(X1,nil),nil)=nil|~segmentP(nil,app(nil,X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_61]), c_0_35])]), c_0_106])).
% 162.80/162.90  cnf(c_0_115, plain, (app(X1,X2)=X2|nil!=X1|~ssList(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_111, c_0_54])).
% 162.80/162.90  cnf(c_0_116, plain, (X1!=nil|~memberP(X1,X2)|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_66]), c_0_35])]), c_0_113])).
% 162.80/162.90  cnf(c_0_117, plain, (app(app(X1,nil),nil)=nil|~segmentP(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_35])])).
% 162.80/162.90  fof(c_0_118, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).
% 162.80/162.90  cnf(c_0_119, plain, (app(X1,X2)!=nil|~memberP(X2,X3)|~ssList(X2)|~ssList(X1)|~ssItem(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_112]), c_0_55])).
% 162.80/162.90  cnf(c_0_120, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
% 162.80/162.90  cnf(c_0_121, plain, (app(app(nil,X1),nil)=nil|~segmentP(nil,X1)|~ssList(X1)), inference(spm,[status(thm)],[c_0_117, c_0_65])).
% 162.80/162.90  cnf(c_0_122, plain, (memberP(cons(X2,X3),X1)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_118])).
% 162.80/162.90  fof(c_0_123, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
% 162.80/162.90  fof(c_0_124, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((nil!=X2|X2!=X4)|X1!=X3)|nil=X1)|?[X5]:(ssItem(X5)&((~memberP(X3,X5)&?[X6]:(ssList(X6)&segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil)))))|(![X6]:(ssList(X6)=>~segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))))&memberP(X3,X5))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_25])])).
% 162.80/162.90  cnf(c_0_125, plain, (app(X1,app(X2,X3))!=nil|~memberP(X2,X4)|~ssList(X1)|~ssList(X3)|~ssList(X2)|~ssItem(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_55])).
% 162.80/162.90  cnf(c_0_126, plain, (app(nil,app(X1,nil))=nil|~segmentP(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_121]), c_0_35])])).
% 162.80/162.90  cnf(c_0_127, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
% 162.80/162.90  cnf(c_0_128, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
% 162.80/162.90  cnf(c_0_129, plain, (memberP(cons(X1,X2),X1)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_122])).
% 162.80/162.90  cnf(c_0_130, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_123])).
% 162.80/162.90  fof(c_0_131, negated_conjecture, (![X11]:![X12]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((nil=esk49_0&esk49_0=esk51_0)&esk48_0=esk50_0)&nil!=esk48_0)&(((memberP(esk50_0,X11)|(~ssList(X12)|~segmentP(esk51_0,app(app(cons(X11,nil),X12),cons(X11,nil)))))|~ssItem(X11))&(((ssList(esk52_1(X11))|~memberP(esk50_0,X11))|~ssItem(X11))&((segmentP(esk51_0,app(app(cons(X11,nil),esk52_1(X11)),cons(X11,nil)))|~memberP(esk50_0,X11))|~ssItem(X11)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_124])])])])])).
% 162.80/162.90  cnf(c_0_132, plain, (~segmentP(nil,X1)|~memberP(X1,X2)|~ssList(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_35])])).
% 162.80/162.90  cnf(c_0_133, plain, (memberP(X1,X2)|~rearsegP(X1,X3)|~memberP(X3,X2)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_127]), c_0_128])).
% 162.80/162.90  cnf(c_0_134, plain, (memberP(X1,esk5_1(X1))|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_34]), c_0_35])]), c_0_36])).
% 162.80/162.90  cnf(c_0_135, plain, (esk5_1(X1)=hd(X1)|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_34]), c_0_35])]), c_0_36])).
% 162.80/162.90  cnf(c_0_136, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_131])).
% 162.80/162.90  cnf(c_0_137, negated_conjecture, (nil=esk49_0), inference(split_conjunct,[status(thm)],[c_0_131])).
% 162.80/162.90  cnf(c_0_138, plain, (~segmentP(nil,X1)|~rearsegP(X1,X2)|~memberP(X2,X3)|~ssList(X1)|~ssList(X2)|~ssItem(X3)), inference(spm,[status(thm)],[c_0_132, c_0_133])).
% 162.80/162.90  cnf(c_0_139, plain, (memberP(X1,hd(X1))|~singletonP(X1)|~ssList(X1)), inference(spm,[status(thm)],[c_0_134, c_0_135])).
% 162.80/162.90  cnf(c_0_140, plain, (ssItem(hd(X1))|~singletonP(X1)|~ssList(X1)), inference(spm,[status(thm)],[c_0_36, c_0_135])).
% 162.80/162.90  cnf(c_0_141, negated_conjecture, (segmentP(esk51_0,app(app(cons(X1,nil),esk52_1(X1)),cons(X1,nil)))|~ssItem(X1)|~memberP(esk50_0,X1)), inference(split_conjunct,[status(thm)],[c_0_131])).
% 162.80/162.90  cnf(c_0_142, negated_conjecture, (esk51_0=nil), inference(rw,[status(thm)],[c_0_136, c_0_137])).
% 162.80/162.90  cnf(c_0_143, plain, (~segmentP(nil,X1)|~rearsegP(X1,X2)|~singletonP(X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_140])).
% 162.80/162.90  cnf(c_0_144, plain, (rearsegP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_74]), c_0_55])).
% 162.80/162.90  cnf(c_0_145, negated_conjecture, (segmentP(nil,app(app(cons(X1,nil),esk52_1(X1)),cons(X1,nil)))|~memberP(esk50_0,X1)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_141, c_0_142])).
% 162.80/162.90  cnf(c_0_146, plain, (~segmentP(nil,app(X1,X2))|~singletonP(X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_55])).
% 162.80/162.90  cnf(c_0_147, negated_conjecture, (segmentP(nil,app(app(X1,esk52_1(esk5_1(X1))),X1))|~singletonP(X1)|~memberP(esk50_0,esk5_1(X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_34]), c_0_36])).
% 162.80/162.90  cnf(c_0_148, negated_conjecture, (~singletonP(X1)|~memberP(esk50_0,esk5_1(X1))|~ssList(app(X1,esk52_1(esk5_1(X1))))|~ssList(X1)), inference(spm,[status(thm)],[c_0_146, c_0_147])).
% 162.80/162.90  fof(c_0_149, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
% 162.80/162.90  cnf(c_0_150, negated_conjecture, (~singletonP(X1)|~memberP(esk50_0,esk5_1(X1))|~ssList(esk52_1(esk5_1(X1)))|~ssList(X1)), inference(spm,[status(thm)],[c_0_148, c_0_55])).
% 162.80/162.90  cnf(c_0_151, negated_conjecture, (ssList(esk52_1(X1))|~ssItem(X1)|~memberP(esk50_0,X1)), inference(split_conjunct,[status(thm)],[c_0_131])).
% 162.80/162.90  cnf(c_0_152, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_149])).
% 162.80/162.90  cnf(c_0_153, negated_conjecture, (~singletonP(X1)|~memberP(esk50_0,esk5_1(X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_36])).
% 162.80/162.90  cnf(c_0_154, plain, (ssList(cons(X1,app(X2,X3)))|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_152]), c_0_58])).
% 162.80/162.90  cnf(c_0_155, negated_conjecture, (~singletonP(X1)|~memberP(esk50_0,hd(X1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_153, c_0_135])).
% 162.80/162.90  cnf(c_0_156, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_80]), c_0_35])])).
% 162.80/162.90  fof(c_0_157, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])).
% 162.80/162.90  cnf(c_0_158, negated_conjecture, (~memberP(esk50_0,hd(cons(X1,nil)))|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_44]), c_0_156])).
% 162.80/162.90  cnf(c_0_159, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_157])).
% 162.80/162.90  cnf(c_0_160, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_157])).
% 162.80/162.90  cnf(c_0_161, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_157])).
% 162.80/162.90  cnf(c_0_162, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_131])).
% 162.80/162.90  cnf(c_0_163, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_131])).
% 162.80/162.90  cnf(c_0_164, negated_conjecture, (nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_131])).
% 162.80/162.90  cnf(c_0_165, negated_conjecture, (~memberP(esk50_0,X1)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_130]), c_0_35])])).
% 162.80/162.90  cnf(c_0_166, plain, (nil=X1|memberP(X1,esk45_1(X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_159]), c_0_160]), c_0_161])).
% 162.80/162.90  cnf(c_0_167, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_162, c_0_163])).
% 162.80/162.90  cnf(c_0_168, negated_conjecture, (esk50_0!=nil), inference(rw,[status(thm)],[c_0_164, c_0_163])).
% 162.80/162.90  cnf(c_0_169, negated_conjecture, (~ssItem(esk45_1(esk50_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_167])]), c_0_168])).
% 162.80/162.90  cnf(c_0_170, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_160]), c_0_167])]), c_0_168]), ['proof']).
% 162.80/162.90  # SZS output end CNFRefutation
% 162.80/162.90  # Training examples: 0 positive, 0 negative
% 162.80/162.93  
% 162.80/162.93  # -------------------------------------------------
% 162.80/162.93  # User time                : 159.072 s
% 162.80/162.93  # System time              : 3.730 s
% 162.80/162.93  # Total time               : 162.801 s
% 162.80/162.93  # Maximum resident set size: 1868 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
