<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC402+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC402+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC402+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC402%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC402+1.p">SWC402+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n153.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:48 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :   45 ( 238 expanded)
%              Number of clauses        :   34 (  86 expanded)
%              Number of leaves         :    5 (  56 expanded)
%              Depth                    :   10
%              Number of atoms          :  214 (3836 expanded)
%              Number of equality atoms :   47 (1129 expanded)
%              Maximal formula depth    :   30 (   5 average)
%              Maximal clause size      :   29 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( rearsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax6)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_2"></A>fof(c_0_2,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ~ ssList(X4)
                  | X2 != X4
                  | X1 != X3
                  | ! [X5] :
                      ( ssList(X5)
                     => ! [X6] :
                          ( ~ ssList(X6)
                          | app(app(X5,X3),X6) != X4
                          | ~ totalorderedP(X3)
                          | ? [X7] :
                              ( ssItem(X7)
                              & ? [X8] :
                                  ( ssList(X8)
                                  & app(X8,cons(X7,nil)) = X5
                                  & ? [X9] :
                                      ( ssItem(X9)
                                      & ? [X10] :
                                          ( ssList(X10)
                                          & app(cons(X9,nil),X10) = X3
                                          & leq(X7,X9) ) ) ) )
                          | ? [X11] :
                              ( ssItem(X11)
                              & ? [X12] :
                                  ( ssList(X12)
                                  & app(cons(X11,nil),X12) = X6
                                  & ? [X13] :
                                      ( ssItem(X13)
                                      & ? [X14] :
                                          ( ssList(X14)
                                          & app(X14,cons(X13,nil)) = X3
                                          & leq(X13,X11) ) ) ) ) ) )
                  | ! [X15] :
                      ( ~ ssItem(X15)
                      | ~ memberP(X1,X15)
                      | memberP(X2,X15) )
                  | ( nil != X4
                    & nil = X3 ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X3,X2) = app(X1,X2)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax79)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(app(X2,X3),X1)
              &lt;=> ( memberP(X2,X1)
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax36)).

<A NAME="c_0_5"></A>fof(c_0_5,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk7_2(X4,X5))
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(esk7_2(X4,X5),X5) = X4
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X7,X5) != X4
        | rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])).

<A NAME="c_0_6"></A>fof(c_0_6,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_7"></A>fof(c_0_7,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ~ ssList(X4)
                    | X2 != X4
                    | X1 != X3
                    | ! [X5] :
                        ( ssList(X5)
                       => ! [X6] :
                            ( ~ ssList(X6)
                            | app(app(X5,X3),X6) != X4
                            | ~ totalorderedP(X3)
                            | ? [X7] :
                                ( ssItem(X7)
                                & ? [X8] :
                                    ( ssList(X8)
                                    & app(X8,cons(X7,nil)) = X5
                                    & ? [X9] :
                                        ( ssItem(X9)
                                        & ? [X10] :
                                            ( ssList(X10)
                                            & app(cons(X9,nil),X10) = X3
                                            & leq(X7,X9) ) ) ) )
                            | ? [X11] :
                                ( ssItem(X11)
                                & ? [X12] :
                                    ( ssList(X12)
                                    & app(cons(X11,nil),X12) = X6
                                    & ? [X13] :
                                        ( ssItem(X13)
                                        & ? [X14] :
                                            ( ssList(X14)
                                            & app(X14,cons(X13,nil)) = X3
                                            & leq(X13,X11) ) ) ) ) ) )
                    | ! [X15] :
                        ( ~ ssItem(X15)
                        | ~ memberP(X1,X15)
                        | memberP(X2,X15) )
                    | ( nil != X4
                      & nil = X3 ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_2])])).

<A NAME="c_0_8"></A>cnf(c_0_8,plain,
    ( rearsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X3,X2) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_9"></A>cnf(c_0_9,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_10"></A>fof(c_0_10,negated_conjecture,(
    ! [X22,X23,X24,X25,X26,X27,X28,X29] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ssList(esk52_0)
      & ssList(esk53_0)
      & app(app(esk52_0,esk50_0),esk53_0) = esk51_0
      & totalorderedP(esk50_0)
      & ( ~ ssItem(X22)
        | ~ ssList(X23)
        | app(X23,cons(X22,nil)) != esk52_0
        | ~ ssItem(X24)
        | ~ ssList(X25)
        | app(cons(X24,nil),X25) != esk50_0
        | ~ leq(X22,X24) )
      & ( ~ ssItem(X26)
        | ~ ssList(X27)
        | app(cons(X26,nil),X27) != esk53_0
        | ~ ssItem(X28)
        | ~ ssList(X29)
        | app(X29,cons(X28,nil)) != esk50_0
        | ~ leq(X28,X26) )
      & ssItem(esk54_0)
      & memberP(esk48_0,esk54_0)
      & ~ memberP(esk49_0,esk54_0)
      & ( nil = esk51_0
        | nil != esk50_0 ) ) ),
    inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).

<A NAME="c_0_11"></A>fof(c_0_11,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X6,X5) != app(X4,X5)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

<A NAME="c_0_12"></A>cnf(c_0_12,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_8]),<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( app(app(esk52_0,esk50_0),esk53_0) = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( ssList(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,plain,
    ( X1 = X2
    | app(X1,X3) != app(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( rearsegP(esk51_0,esk53_0)
    | ~ ssList(app(esk52_0,esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_13]),c_0_14])])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( ssList(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_23"></A>fof(c_0_23,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(app(X5,X6),X4)
        | memberP(X5,X4)
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X5,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( memberP(esk48_0,esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( X1 = app(esk52_0,esk50_0)
    | app(X1,esk53_0) != esk51_0
    | ~ ssList(app(esk52_0,esk50_0))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_13]),c_0_14])])).

<A NAME="c_0_26"></A>cnf(c_0_26,plain,
    ( app(esk7_2(X1,X2),X2) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( rearsegP(esk51_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_9]),c_0_19]),c_0_20])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,plain,
    ( ssList(esk7_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,plain,
    ( memberP(app(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(X3,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( memberP(esk50_0,esk54_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( ssItem(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( X1 = app(esk52_0,esk50_0)
    | app(X1,esk53_0) != esk51_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_9]),c_0_19]),c_0_20])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( app(esk7_2(esk51_0,esk53_0),esk53_0) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_27]),c_0_14]),c_0_28])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( ssList(esk7_2(esk51_0,esk53_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_27]),c_0_14]),c_0_28])])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( memberP(app(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(X2,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( memberP(app(X1,esk50_0),esk54_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31]),c_0_19]),c_0_32])])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( esk7_2(esk51_0,esk53_0) = app(esk52_0,esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_34]),c_0_35])])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( memberP(app(app(X1,esk50_0),X2),esk54_0)
    | ~ ssList(app(X1,esk50_0))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_37]),c_0_32])])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( ssList(app(esk52_0,esk50_0)) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( ~ memberP(esk49_0,esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( memberP(app(app(esk52_0,esk50_0),X1),esk54_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_40]),c_0_20])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( ~ memberP(esk51_0,esk54_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_13]),c_0_14])]),<A HREF="#c_0_43">c_0_43</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC402+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.22  % Computer   : n153.star.cs.uiowa.edu
% 0.03/0.22  % Model      : x86_64 x86_64
% 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.22  % Memory     : 32218.75MB
% 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.22  % CPULimit   : 300
% 0.03/0.22  % DateTime   : Tue May 10 20:43:53 CDT 2016
% 0.03/0.22  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.40  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.40  # and selection function SelectCQPrecWNTNp.
% 0.06/0.40  #
% 0.06/0.40  # Preprocessing time       : 0.014 s
% 0.06/0.40  # Presaturation interreduction done
% 0.06/0.40  
% 0.06/0.40  # Proof found!
% 0.06/0.40  # SZS status Theorem
% 0.06/0.40  # SZS output start CNFRefutation
% 0.06/0.40  fof(c_0_0, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax6)).
% 0.06/0.40  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.06/0.40  fof(c_0_2, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~(ssList(X4))|X2!=X4)|X1!=X3)|![X5]:(ssList(X5)=>![X6]:((((~(ssList(X6))|app(app(X5,X3),X6)!=X4)|~(totalorderedP(X3)))|?[X7]:(ssItem(X7)&?[X8]:((ssList(X8)&app(X8,cons(X7,nil))=X5)&?[X9]:(ssItem(X9)&?[X10]:((ssList(X10)&app(cons(X9,nil),X10)=X3)&leq(X7,X9))))))|?[X11]:(ssItem(X11)&?[X12]:((ssList(X12)&app(cons(X11,nil),X12)=X6)&?[X13]:(ssItem(X13)&?[X14]:((ssList(X14)&app(X14,cons(X13,nil))=X3)&leq(X13,X11))))))))|![X15]:((~(ssItem(X15))|~(memberP(X1,X15)))|memberP(X2,X15)))|(nil!=X4&nil=X3)))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.06/0.40  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax79)).
% 0.06/0.40  fof(c_0_4, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(memberP(app(X2,X3),X1)&lt;=>(memberP(X2,X1)|memberP(X3,X1))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax36)).
% 0.06/0.40  fof(c_0_5, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])])])).
% 0.06/0.40  fof(c_0_6, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.06/0.40  fof(c_0_7, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~ssList(X4)|X2!=X4)|X1!=X3)|![X5]:(ssList(X5)=>![X6]:((((~ssList(X6)|app(app(X5,X3),X6)!=X4)|~totalorderedP(X3))|?[X7]:(ssItem(X7)&?[X8]:((ssList(X8)&app(X8,cons(X7,nil))=X5)&?[X9]:(ssItem(X9)&?[X10]:((ssList(X10)&app(cons(X9,nil),X10)=X3)&leq(X7,X9))))))|?[X11]:(ssItem(X11)&?[X12]:((ssList(X12)&app(cons(X11,nil),X12)=X6)&?[X13]:(ssItem(X13)&?[X14]:((ssList(X14)&app(X14,cons(X13,nil))=X3)&leq(X13,X11))))))))|![X15]:((~ssItem(X15)|~memberP(X1,X15))|memberP(X2,X15)))|(nil!=X4&nil=X3))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_2])])).
% 0.06/0.40  cnf(c_0_8, plain, (rearsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X3,X2)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.06/0.40  cnf(c_0_9, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
% 0.06/0.40  fof(c_0_10, negated_conjecture, (![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:![X29]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(ssList(esk52_0)&((((ssList(esk53_0)&app(app(esk52_0,esk50_0),esk53_0)=esk51_0)&totalorderedP(esk50_0))&(~ssItem(X22)|((~ssList(X23)|app(X23,cons(X22,nil))!=esk52_0)|(~ssItem(X24)|((~ssList(X25)|app(cons(X24,nil),X25)!=esk50_0)|~leq(X22,X24))))))&(~ssItem(X26)|((~ssList(X27)|app(cons(X26,nil),X27)!=esk53_0)|(~ssItem(X28)|((~ssList(X29)|app(X29,cons(X28,nil))!=esk50_0)|~leq(X28,X26))))))))&((ssItem(esk54_0)&memberP(esk48_0,esk54_0))&~memberP(esk49_0,esk54_0)))&(nil=esk51_0|nil!=esk50_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
% 0.06/0.40  fof(c_0_11, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
% 0.06/0.40  cnf(c_0_12, plain, (rearsegP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_8]), c_0_9])).
% 0.06/0.40  cnf(c_0_13, negated_conjecture, (app(app(esk52_0,esk50_0),esk53_0)=esk51_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.06/0.40  cnf(c_0_14, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.06/0.40  cnf(c_0_15, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.06/0.40  cnf(c_0_16, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.06/0.40  cnf(c_0_17, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.06/0.40  cnf(c_0_18, negated_conjecture, (rearsegP(esk51_0,esk53_0)|~ssList(app(esk52_0,esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])])).
% 0.06/0.40  cnf(c_0_19, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
% 0.06/0.40  cnf(c_0_20, negated_conjecture, (ssList(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.06/0.40  cnf(c_0_21, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.06/0.40  cnf(c_0_22, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.06/0.40  fof(c_0_23, plain, (![X4]:![X5]:![X6]:(((((~memberP(app(X5,X6),X4)|(memberP(X5,X4)|memberP(X6,X4)))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&(((((~memberP(X5,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
% 0.06/0.40  cnf(c_0_24, negated_conjecture, (memberP(esk48_0,esk54_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.06/0.40  cnf(c_0_25, negated_conjecture, (X1=app(esk52_0,esk50_0)|app(X1,esk53_0)!=esk51_0|~ssList(app(esk52_0,esk50_0))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_13]), c_0_14])])).
% 0.06/0.40  cnf(c_0_26, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.06/0.40  cnf(c_0_27, negated_conjecture, (rearsegP(esk51_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_9]), c_0_19]), c_0_20])])).
% 0.06/0.40  cnf(c_0_28, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
% 0.06/0.40  cnf(c_0_29, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.06/0.40  cnf(c_0_30, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 0.06/0.40  cnf(c_0_31, negated_conjecture, (memberP(esk50_0,esk54_0)), inference(rw,[status(thm)],[c_0_24, c_0_16])).
% 0.06/0.40  cnf(c_0_32, negated_conjecture, (ssItem(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.06/0.40  cnf(c_0_33, negated_conjecture, (X1=app(esk52_0,esk50_0)|app(X1,esk53_0)!=esk51_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_9]), c_0_19]), c_0_20])])).
% 0.06/0.40  cnf(c_0_34, negated_conjecture, (app(esk7_2(esk51_0,esk53_0),esk53_0)=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_14]), c_0_28])])).
% 0.06/0.40  cnf(c_0_35, negated_conjecture, (ssList(esk7_2(esk51_0,esk53_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_27]), c_0_14]), c_0_28])])).
% 0.06/0.40  cnf(c_0_36, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 0.06/0.40  cnf(c_0_37, negated_conjecture, (memberP(app(X1,esk50_0),esk54_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_19]), c_0_32])])).
% 0.06/0.40  cnf(c_0_38, negated_conjecture, (esk7_2(esk51_0,esk53_0)=app(esk52_0,esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
% 0.06/0.40  cnf(c_0_39, negated_conjecture, (memberP(app(app(X1,esk50_0),X2),esk54_0)|~ssList(app(X1,esk50_0))|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_32])])).
% 0.06/0.40  cnf(c_0_40, negated_conjecture, (ssList(app(esk52_0,esk50_0))), inference(rw,[status(thm)],[c_0_35, c_0_38])).
% 0.06/0.40  cnf(c_0_41, negated_conjecture, (~memberP(esk49_0,esk54_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.06/0.40  cnf(c_0_42, negated_conjecture, (memberP(app(app(esk52_0,esk50_0),X1),esk54_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_20])])).
% 0.06/0.40  cnf(c_0_43, negated_conjecture, (~memberP(esk51_0,esk54_0)), inference(rw,[status(thm)],[c_0_41, c_0_22])).
% 0.06/0.40  cnf(c_0_44, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_13]), c_0_14])]), c_0_43]), ['proof']).
% 0.06/0.40  # SZS output end CNFRefutation
% 0.06/0.40  # Training examples: 0 positive, 0 negative
% 0.06/0.40  
% 0.06/0.40  # -------------------------------------------------
% 0.06/0.40  # User time                : 0.171 s
% 0.06/0.40  # System time              : 0.007 s
% 0.06/0.40  # Total time               : 0.178 s
% 0.06/0.40  # Maximum resident set size: 1892 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
