<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC303+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC303+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC303+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC303%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC303+1.p">SWC303+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n122.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:25 EDT 2016

% Result     : Theorem 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of formulae       :   58 ( 343 expanded)
%              Number of clauses        :   37 ( 118 expanded)
%              Number of leaves         :   10 (  85 expanded)
%              Depth                    :   16
%              Number of atoms          :  258 (4817 expanded)
%              Number of equality atoms :   53 (1464 expanded)
%              Maximal formula depth    :   27 (   5 average)
%              Maximal clause size      :   30 (   2 average)
%              Maximal term depth       :    6 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ~ ssList(X4)
                  | X4 != X2
                  | X3 != X1
                  | ! [X5] :
                      ( ~ ssList(X5)
                      | app(X3,X5) != X4
                      | ~ totalorderedP(X3)
                      | ? [X6] :
                          ( ssItem(X6)
                          & ? [X7] :
                              ( ssList(X7)
                              & app(cons(X6,nil),X7) = X5
                              & ? [X8] :
                                  ( ssItem(X8)
                                  & ? [X9] :
                                      ( ssList(X9)
                                      & app(X9,cons(X8,nil)) = X3
                                      & leq(X8,X6) ) ) ) ) )
                  | ! [X10] :
                      ( ssItem(X10)
                     => ! [X11] :
                          ( ssItem(X11)
                         => ! [X12] :
                              ( ssList(X12)
                             => ! [X13] :
                                  ( ssList(X13)
                                 => ! [X14] :
                                      ( ~ ssList(X14)
                                      | app(app(app(app(X12,cons(X10,nil)),X13),cons(X11,nil)),X14) != X1
                                      | ~ lt(X11,X10) ) ) ) ) )
                  | ( nil != X4
                    & nil = X3 ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( lt(X1,X2)
          &lt;=> ( X1 != X2
              & leq(X1,X2) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax93)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( totalorderedP(X1)
      &lt;=> ! [X2] :
            ( ssItem(X2)
           => ! [X3] :
                ( ssItem(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ! [X5] :
                        ( ssList(X5)
                       => ! [X6] :
                            ( ssList(X6)
                           => ( app(app(X4,cons(X2,X5)),cons(X3,X6)) = X1
                             => leq(X2,X3) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax11)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( lt(X1,X2)
           => ~ lt(X2,X1) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax33)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( ( leq(X1,X2)
              & leq(X2,X1) )
           => X1 = X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax29)).

<A NAME="c_0_10"></A>fof(c_0_10,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ~ ssList(X4)
                    | X4 != X2
                    | X3 != X1
                    | ! [X5] :
                        ( ~ ssList(X5)
                        | app(X3,X5) != X4
                        | ~ totalorderedP(X3)
                        | ? [X6] :
                            ( ssItem(X6)
                            & ? [X7] :
                                ( ssList(X7)
                                & app(cons(X6,nil),X7) = X5
                                & ? [X8] :
                                    ( ssItem(X8)
                                    & ? [X9] :
                                        ( ssList(X9)
                                        & app(X9,cons(X8,nil)) = X3
                                        & leq(X8,X6) ) ) ) ) )
                    | ! [X10] :
                        ( ssItem(X10)
                       => ! [X11] :
                            ( ssItem(X11)
                           => ! [X12] :
                                ( ssList(X12)
                               => ! [X13] :
                                    ( ssList(X13)
                                   => ! [X14] :
                                        ( ~ ssList(X14)
                                        | app(app(app(app(X12,cons(X10,nil)),X13),cons(X11,nil)),X14) != X1
                                        | ~ lt(X11,X10) ) ) ) ) )
                    | ( nil != X4
                      & nil = X3 ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_11"></A>fof(c_0_11,negated_conjecture,(
    ! [X20,X21,X22,X23] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk51_0 = esk49_0
      & esk50_0 = esk48_0
      & ssList(esk52_0)
      & app(esk50_0,esk52_0) = esk51_0
      & totalorderedP(esk50_0)
      & ( ~ ssItem(X20)
        | ~ ssList(X21)
        | app(cons(X20,nil),X21) != esk52_0
        | ~ ssItem(X22)
        | ~ ssList(X23)
        | app(X23,cons(X22,nil)) != esk50_0
        | ~ leq(X22,X20) )
      & ssItem(esk53_0)
      & ssItem(esk54_0)
      & ssList(esk55_0)
      & ssList(esk56_0)
      & ssList(esk57_0)
      & app(app(app(app(esk55_0,cons(esk53_0,nil)),esk56_0),cons(esk54_0,nil)),esk57_0) = esk48_0
      & lt(esk54_0,esk53_0)
      & ( nil = esk51_0
        | nil != esk50_0 ) ) ),
    inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( app(app(app(app(esk55_0,cons(esk53_0,nil)),esk56_0),cons(esk54_0,nil)),esk57_0) = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( esk50_0 = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_14"></A>fof(c_0_14,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( app(app(app(app(esk55_0,cons(esk53_0,nil)),esk56_0),cons(esk54_0,nil)),esk57_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( ssList(esk56_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( ssList(esk55_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_19"></A>fof(c_0_19,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( app(app(app(esk55_0,app(cons(esk53_0,nil),esk56_0)),cons(esk54_0,nil)),esk57_0) = esk50_0
    | ~ ssList(cons(esk53_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_16]),c_0_17]),c_0_18])])).

<A NAME="c_0_21"></A>cnf(c_0_21,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( ssItem(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_24"></A>fof(c_0_24,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( app(app(app(esk55_0,app(cons(esk53_0,nil),esk56_0)),cons(esk54_0,nil)),esk57_0) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_21]),c_0_22]),c_0_23])])).

<A NAME="c_0_26"></A>cnf(c_0_26,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( app(app(app(esk55_0,cons(esk53_0,esk56_0)),cons(esk54_0,nil)),esk57_0) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26]),c_0_17]),c_0_23])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( ssList(esk57_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( app(app(esk55_0,cons(esk53_0,esk56_0)),app(cons(esk54_0,nil),esk57_0)) = esk50_0
    | ~ ssList(app(esk55_0,cons(esk53_0,esk56_0)))
    | ~ ssList(cons(esk54_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_27]),c_0_28])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( ssItem(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_31"></A>fof(c_0_31,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( app(app(esk55_0,cons(esk53_0,esk56_0)),app(cons(esk54_0,nil),esk57_0)) = esk50_0
    | ~ ssList(app(esk55_0,cons(esk53_0,esk56_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_21]),c_0_22]),c_0_30])])).

<A NAME="c_0_33"></A>cnf(c_0_33,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_34"></A>fof(c_0_34,plain,(
    ! [X3,X4] :
      ( ( X3 != X4
        | ~ lt(X3,X4)
        | ~ ssItem(X4)
        | ~ ssItem(X3) )
      & ( leq(X3,X4)
        | ~ lt(X3,X4)
        | ~ ssItem(X4)
        | ~ ssItem(X3) )
      & ( X3 = X4
        | ~ leq(X3,X4)
        | lt(X3,X4)
        | ~ ssItem(X4)
        | ~ ssItem(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).

<A NAME="c_0_35"></A>fof(c_0_35,plain,(
    ! [X7,X8,X9,X10,X11,X12] :
      ( ( ~ totalorderedP(X7)
        | ~ ssItem(X8)
        | ~ ssItem(X9)
        | ~ ssList(X10)
        | ~ ssList(X11)
        | ~ ssList(X12)
        | app(app(X10,cons(X8,X11)),cons(X9,X12)) != X7
        | leq(X8,X9)
        | ~ ssList(X7) )
      & ( ssItem(esk25_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssItem(esk26_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk27_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk28_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk29_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( app(app(esk27_1(X7),cons(esk25_1(X7),esk28_1(X7))),cons(esk26_1(X7),esk29_1(X7))) = X7
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ~ leq(esk25_1(X7),esk26_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( app(app(esk55_0,cons(esk53_0,esk56_0)),app(cons(esk54_0,nil),esk57_0)) = esk50_0
    | ~ ssList(cons(esk53_0,esk56_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_33]),c_0_18])])).

<A NAME="c_0_37"></A>fof(c_0_37,plain,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( lt(X1,X2)
           => ~ lt(X2,X1) ) ) ) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_38"></A>fof(c_0_38,plain,(
    ! [X3,X4] :
      ( ~ ssItem(X3)
      | ~ ssItem(X4)
      | ~ leq(X3,X4)
      | ~ leq(X4,X3)
      | X3 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_39"></A>cnf(c_0_39,plain,
    ( leq(X1,X2)
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ lt(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( lt(esk54_0,esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( leq(X2,X3)
    | ~ ssList(X1)
    | app(app(X4,cons(X2,X5)),cons(X3,X6)) != X1
    | ~ ssList(X6)
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssItem(X3)
    | ~ ssItem(X2)
    | ~ totalorderedP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( app(app(esk55_0,cons(esk53_0,esk56_0)),app(cons(esk54_0,nil),esk57_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_21]),c_0_17]),c_0_23])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_44"></A>fof(c_0_44,plain,(
    ! [X3,X4] :
      ( ~ ssItem(X3)
      | ~ ssItem(X4)
      | ~ lt(X3,X4)
      | ~ lt(X4,X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( X1 = X2
    | ~ leq(X2,X1)
    | ~ leq(X1,X2)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( leq(esk54_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_40]),c_0_23]),c_0_30])])).

<A NAME="c_0_47"></A>cnf(c_0_47,plain,
    ( leq(X1,X2)
    | ~ totalorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( app(app(esk55_0,cons(esk53_0,esk56_0)),cons(esk54_0,esk57_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_26]),c_0_28]),c_0_30])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( totalorderedP(esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,plain,
    ( ~ lt(X1,X2)
    | ~ lt(X2,X1)
    | ~ ssItem(X1)
    | ~ ssItem(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( esk54_0 = esk53_0
    | ~ leq(esk53_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_46]),c_0_30]),c_0_23])])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( leq(esk53_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_48]),c_0_49]),c_0_50]),c_0_28]),c_0_17]),c_0_18]),c_0_30]),c_0_23])])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( ~ lt(esk53_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,c_0_40]),c_0_30]),c_0_23])])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( esk54_0 = esk53_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_53])])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( ~ lt(esk53_0,esk53_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_55]),<A HREF="#c_0_56">c_0_56</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC303+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.22  % Computer   : n122.star.cs.uiowa.edu
% 0.03/0.22  % Model      : x86_64 x86_64
% 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.22  % Memory     : 32218.75MB
% 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.22  % CPULimit   : 300
% 0.03/0.22  % DateTime   : Tue May 10 20:40:37 CDT 2016
% 0.03/0.22  % CPUTime    : 
% 0.03/0.22  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.28  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.03/0.28  # and selection function SelectCQPrecWNTNp.
% 0.03/0.28  #
% 0.03/0.28  # Preprocessing time       : 0.014 s
% 0.03/0.28  # Presaturation interreduction done
% 0.03/0.28  
% 0.03/0.28  # Proof found!
% 0.03/0.28  # SZS status Theorem
% 0.03/0.28  # SZS output start CNFRefutation
% 0.03/0.28  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~(ssList(X4))|X4!=X2)|X3!=X1)|![X5]:(((~(ssList(X5))|app(X3,X5)!=X4)|~(totalorderedP(X3)))|?[X6]:(ssItem(X6)&?[X7]:((ssList(X7)&app(cons(X6,nil),X7)=X5)&?[X8]:(ssItem(X8)&?[X9]:((ssList(X9)&app(X9,cons(X8,nil))=X3)&leq(X8,X6)))))))|![X10]:(ssItem(X10)=>![X11]:(ssItem(X11)=>![X12]:(ssList(X12)=>![X13]:(ssList(X13)=>![X14]:((~(ssList(X14))|app(app(app(app(X12,cons(X10,nil)),X13),cons(X11,nil)),X14)!=X1)|~(lt(X11,X10))))))))|(nil!=X4&nil=X3)))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.03/0.28  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.03/0.28  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.03/0.28  fof(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.03/0.28  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.03/0.28  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.03/0.28  fof(c_0_6, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>(lt(X1,X2)&lt;=>(X1!=X2&leq(X1,X2)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax93)).
% 0.03/0.28  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(totalorderedP(X1)&lt;=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(app(X4,cons(X2,X5)),cons(X3,X6))=X1=>leq(X2,X3)))))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax11)).
% 0.03/0.28  fof(c_0_8, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>(lt(X1,X2)=>~(lt(X2,X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax33)).
% 0.03/0.28  fof(c_0_9, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>((leq(X1,X2)&leq(X2,X1))=>X1=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax29)).
% 0.03/0.28  fof(c_0_10, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~ssList(X4)|X4!=X2)|X3!=X1)|![X5]:(((~ssList(X5)|app(X3,X5)!=X4)|~totalorderedP(X3))|?[X6]:(ssItem(X6)&?[X7]:((ssList(X7)&app(cons(X6,nil),X7)=X5)&?[X8]:(ssItem(X8)&?[X9]:((ssList(X9)&app(X9,cons(X8,nil))=X3)&leq(X8,X6)))))))|![X10]:(ssItem(X10)=>![X11]:(ssItem(X11)=>![X12]:(ssList(X12)=>![X13]:(ssList(X13)=>![X14]:((~ssList(X14)|app(app(app(app(X12,cons(X10,nil)),X13),cons(X11,nil)),X14)!=X1)|~lt(X11,X10)))))))|(nil!=X4&nil=X3))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.03/0.28  fof(c_0_11, negated_conjecture, (![X20]:![X21]:![X22]:![X23]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(((((ssList(esk51_0)&esk51_0=esk49_0)&esk50_0=esk48_0)&(((ssList(esk52_0)&app(esk50_0,esk52_0)=esk51_0)&totalorderedP(esk50_0))&(~ssItem(X20)|((~ssList(X21)|app(cons(X20,nil),X21)!=esk52_0)|(~ssItem(X22)|((~ssList(X23)|app(X23,cons(X22,nil))!=esk50_0)|~leq(X22,X20)))))))&(ssItem(esk53_0)&(ssItem(esk54_0)&(ssList(esk55_0)&(ssList(esk56_0)&((ssList(esk57_0)&app(app(app(app(esk55_0,cons(esk53_0,nil)),esk56_0),cons(esk54_0,nil)),esk57_0)=esk48_0)&lt(esk54_0,esk53_0)))))))&(nil=esk51_0|nil!=esk50_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).
% 0.03/0.28  cnf(c_0_12, negated_conjecture, (app(app(app(app(esk55_0,cons(esk53_0,nil)),esk56_0),cons(esk54_0,nil)),esk57_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.03/0.28  cnf(c_0_13, negated_conjecture, (esk50_0=esk48_0), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.03/0.28  fof(c_0_14, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.03/0.28  cnf(c_0_15, negated_conjecture, (app(app(app(app(esk55_0,cons(esk53_0,nil)),esk56_0),cons(esk54_0,nil)),esk57_0)=esk50_0), inference(rw,[status(thm)],[c_0_12, c_0_13])).
% 0.03/0.28  cnf(c_0_16, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.03/0.28  cnf(c_0_17, negated_conjecture, (ssList(esk56_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.03/0.28  cnf(c_0_18, negated_conjecture, (ssList(esk55_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.03/0.28  fof(c_0_19, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.03/0.28  cnf(c_0_20, negated_conjecture, (app(app(app(esk55_0,app(cons(esk53_0,nil),esk56_0)),cons(esk54_0,nil)),esk57_0)=esk50_0|~ssList(cons(esk53_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])])).
% 0.03/0.28  cnf(c_0_21, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.03/0.28  cnf(c_0_22, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.03/0.28  cnf(c_0_23, negated_conjecture, (ssItem(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.03/0.28  fof(c_0_24, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 0.03/0.28  cnf(c_0_25, negated_conjecture, (app(app(app(esk55_0,app(cons(esk53_0,nil),esk56_0)),cons(esk54_0,nil)),esk57_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])])).
% 0.03/0.28  cnf(c_0_26, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
% 0.03/0.28  cnf(c_0_27, negated_conjecture, (app(app(app(esk55_0,cons(esk53_0,esk56_0)),cons(esk54_0,nil)),esk57_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_17]), c_0_23])])).
% 0.03/0.28  cnf(c_0_28, negated_conjecture, (ssList(esk57_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.03/0.28  cnf(c_0_29, negated_conjecture, (app(app(esk55_0,cons(esk53_0,esk56_0)),app(cons(esk54_0,nil),esk57_0))=esk50_0|~ssList(app(esk55_0,cons(esk53_0,esk56_0)))|~ssList(cons(esk54_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_27]), c_0_28])])).
% 0.03/0.28  cnf(c_0_30, negated_conjecture, (ssItem(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.03/0.28  fof(c_0_31, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 0.03/0.28  cnf(c_0_32, negated_conjecture, (app(app(esk55_0,cons(esk53_0,esk56_0)),app(cons(esk54_0,nil),esk57_0))=esk50_0|~ssList(app(esk55_0,cons(esk53_0,esk56_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_21]), c_0_22]), c_0_30])])).
% 0.03/0.28  cnf(c_0_33, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
% 0.03/0.28  fof(c_0_34, plain, (![X3]:![X4]:(((((X3!=X4|~lt(X3,X4))|~ssItem(X4))|~ssItem(X3))&(((leq(X3,X4)|~lt(X3,X4))|~ssItem(X4))|~ssItem(X3)))&((((X3=X4|~leq(X3,X4))|lt(X3,X4))|~ssItem(X4))|~ssItem(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
% 0.03/0.28  fof(c_0_35, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((~totalorderedP(X7)|(~ssItem(X8)|(~ssItem(X9)|(~ssList(X10)|(~ssList(X11)|(~ssList(X12)|(app(app(X10,cons(X8,X11)),cons(X9,X12))!=X7|leq(X8,X9))))))))|~ssList(X7))&(((ssItem(esk25_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssItem(esk26_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk27_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk28_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk29_1(X7))|totalorderedP(X7))|~ssList(X7))&(((app(app(esk27_1(X7),cons(esk25_1(X7),esk28_1(X7))),cons(esk26_1(X7),esk29_1(X7)))=X7|totalorderedP(X7))|~ssList(X7))&((~leq(esk25_1(X7),esk26_1(X7))|totalorderedP(X7))|~ssList(X7)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
% 0.03/0.28  cnf(c_0_36, negated_conjecture, (app(app(esk55_0,cons(esk53_0,esk56_0)),app(cons(esk54_0,nil),esk57_0))=esk50_0|~ssList(cons(esk53_0,esk56_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_18])])).
% 0.03/0.28  fof(c_0_37, plain, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>(lt(X1,X2)=>~lt(X2,X1))))), inference(fof_simplification,[status(thm)],[c_0_8])).
% 0.03/0.28  fof(c_0_38, plain, (![X3]:![X4]:(~ssItem(X3)|(~ssItem(X4)|((~leq(X3,X4)|~leq(X4,X3))|X3=X4)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 0.03/0.28  cnf(c_0_39, plain, (leq(X1,X2)|~ssItem(X1)|~ssItem(X2)|~lt(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
% 0.03/0.28  cnf(c_0_40, negated_conjecture, (lt(esk54_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.03/0.28  cnf(c_0_41, plain, (leq(X2,X3)|~ssList(X1)|app(app(X4,cons(X2,X5)),cons(X3,X6))!=X1|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssItem(X3)|~ssItem(X2)|~totalorderedP(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
% 0.03/0.28  cnf(c_0_42, negated_conjecture, (app(app(esk55_0,cons(esk53_0,esk56_0)),app(cons(esk54_0,nil),esk57_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_21]), c_0_17]), c_0_23])])).
% 0.03/0.28  cnf(c_0_43, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.03/0.28  fof(c_0_44, plain, (![X3]:![X4]:(~ssItem(X3)|(~ssItem(X4)|(~lt(X3,X4)|~lt(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])).
% 0.03/0.28  cnf(c_0_45, plain, (X1=X2|~leq(X2,X1)|~leq(X1,X2)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
% 0.03/0.28  cnf(c_0_46, negated_conjecture, (leq(esk54_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_23]), c_0_30])])).
% 0.03/0.28  cnf(c_0_47, plain, (leq(X1,X2)|~totalorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_41])).
% 0.03/0.28  cnf(c_0_48, negated_conjecture, (app(app(esk55_0,cons(esk53_0,esk56_0)),cons(esk54_0,esk57_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_26]), c_0_28]), c_0_30])])).
% 0.03/0.28  cnf(c_0_49, negated_conjecture, (totalorderedP(esk50_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.03/0.28  cnf(c_0_50, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_43, c_0_13])).
% 0.03/0.28  cnf(c_0_51, plain, (~lt(X1,X2)|~lt(X2,X1)|~ssItem(X1)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
% 0.03/0.28  cnf(c_0_52, negated_conjecture, (esk54_0=esk53_0|~leq(esk53_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_30]), c_0_23])])).
% 0.03/0.28  cnf(c_0_53, negated_conjecture, (leq(esk53_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50]), c_0_28]), c_0_17]), c_0_18]), c_0_30]), c_0_23])])).
% 0.03/0.28  cnf(c_0_54, negated_conjecture, (~lt(esk53_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_40]), c_0_30]), c_0_23])])).
% 0.03/0.28  cnf(c_0_55, negated_conjecture, (esk54_0=esk53_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_53])])).
% 0.03/0.28  cnf(c_0_56, negated_conjecture, (~lt(esk53_0,esk53_0)), inference(rw,[status(thm)],[c_0_54, c_0_55])).
% 0.03/0.28  cnf(c_0_57, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_55]), c_0_56]), ['proof']).
% 0.03/0.28  # SZS output end CNFRefutation
% 0.03/0.28  # Training examples: 0 positive, 0 negative
% 0.03/0.28  
% 0.03/0.28  # -------------------------------------------------
% 0.03/0.28  # User time                : 0.059 s
% 0.03/0.28  # System time              : 0.001 s
% 0.03/0.28  # Total time               : 0.060 s
% 0.03/0.28  # Maximum resident set size: 1884 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
