<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC022-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC022-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC022-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC022-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC022-1.p">SWC022-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n100.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:20 EDT 2016

% Result     : Unsatisfiable 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of clauses        :   55 ( 206 expanded)
%              Number of leaves         :   22 ( 124 expanded)
%              Depth                    :   20
%              Number of atoms          :  148 ( 417 expanded)
%              Number of equality atoms :   60 ( 151 expanded)
%              Maximal clause size      :    7 (   3 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( neq(sk2,nil)
    | neq(sk2,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( sk3 = X1
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | cons(X2,nil) != X1
    | hd(sk4) != X2
    | ~ neq(nil,sk4)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause76)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause97)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause104)).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( frontsegP(X3,X1)
    | app(X1,X2) != X3
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause144)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause75)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( singletonP(X2)
    | cons(X1,nil) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause116)).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( ~ ssList(X1)
    | ~ neq(X1,nil)
    | ~ frontsegP(sk2,X1)
    | ~ frontsegP(sk1,X1)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( frontsegP(X1,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause61)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause60)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( ~ singletonP(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause11)).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( neq(sk2,nil) ),
    inference(cn,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( neq(sk4,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( cons(hd(sk4),nil) = sk1
    | ~ neq(nil,sk4)
    | ~ ssItem(hd(sk4))
    | ~ ssList(cons(hd(sk4),nil)) ),
    inference(er,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_3]),c_0_23])])])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( cons(hd(sk4),nil) = sk1
    | sk4 = nil
    | ~ ssItem(hd(sk4))
    | ~ ssList(cons(hd(sk4),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_5]),c_0_6]),c_0_25])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( cons(hd(sk4),nil) = sk1
    | sk4 = nil
    | ~ ssList(cons(hd(sk4),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_7]),c_0_25])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( cons(hd(sk4),nil) = sk1
    | sk4 = nil
    | ~ ssItem(hd(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_8]),c_0_6])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( cons(hd(sk4),nil) = sk1
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_7]),c_0_25])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( tl(sk1) = nil
    | sk4 = nil
    | ~ ssItem(hd(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_29]),c_0_6])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( tl(sk1) = nil
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_7]),c_0_25])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( hd(sk1) = hd(sk4)
    | sk4 = nil
    | ~ ssItem(hd(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_29]),c_0_6])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( cons(hd(sk1),nil) = sk1
    | sk4 = nil
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_31]),c_0_12])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( hd(sk1) = hd(sk4)
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_7]),c_0_25])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( cons(hd(sk1),X1) = app(sk1,X1)
    | sk1 = nil
    | sk4 = nil
    | ~ ssItem(hd(sk1))
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,<A HREF="#c_0_33">c_0_33</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( sk4 = nil
    | sk1 = nil
    | ssItem(hd(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_34]),c_0_12])])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( cons(hd(sk4),X1) = app(sk1,X1)
    | sk4 = nil
    | sk1 = nil
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_34]),<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( app(sk1,tl(sk4)) = sk4
    | sk1 = nil
    | sk4 = nil
    | ~ ssList(tl(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_37]),c_0_25])])).

<A NAME="c_0_39"></A>cnf(c_0_39,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_14]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( app(sk1,tl(sk4)) = sk4
    | sk4 = nil
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_16]),c_0_25])])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1)
    | ~ ssList(cons(X1,nil)) ),
    inference(er,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( ~ neq(X1,nil)
    | ~ frontsegP(sk1,X1)
    | ~ frontsegP(sk4,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_1]),c_0_23])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( sk1 = nil
    | sk4 = nil
    | frontsegP(sk4,sk1)
    | ~ ssList(tl(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_40]),c_0_12])])).

<A NAME="c_0_44"></A>cnf(c_0_44,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_8]),c_0_6])])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( ~ neq(sk1,nil)
    | ~ frontsegP(sk4,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_19]),c_0_12])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( sk4 = nil
    | sk1 = nil
    | frontsegP(sk4,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_16]),c_0_25])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( sk4 = nil
    | singletonP(sk1)
    | ~ ssItem(hd(sk4)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( sk1 = nil
    | sk4 = nil
    | ~ neq(sk1,nil) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( sk4 = nil
    | singletonP(sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_7]),c_0_25])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( sk4 = nil
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_5]),c_0_12]),c_0_6])])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( ~ neq(nil,nil)
    | ~ frontsegP(sk4,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_20]),c_0_6]),c_0_12])])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( sk4 = nil ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_50]),<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,c_0_20]),c_0_25])])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_52]),<A HREF="#c_0_53">c_0_53</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC022-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n100.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:34:23 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.28  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.03/0.28  # and selection function SelectCQArNTEqFirst.
% 0.03/0.28  #
% 0.03/0.28  # Preprocessing time       : 0.010 s
% 0.03/0.28  # Presaturation interreduction done
% 0.03/0.28  
% 0.03/0.28  # Proof found!
% 0.03/0.28  # SZS status Unsatisfiable
% 0.03/0.28  # SZS output start CNFRefutation
% 0.03/0.28  cnf(c_0_0, negated_conjecture, (neq(sk2,nil)|neq(sk2,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 0.03/0.28  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 0.03/0.28  cnf(c_0_2, negated_conjecture, (sk3=X1|~ssList(X1)|~ssItem(X2)|cons(X2,nil)!=X1|hd(sk4)!=X2|~neq(nil,sk4)|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_12)).
% 0.03/0.28  cnf(c_0_3, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.03/0.28  cnf(c_0_4, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 0.03/0.28  cnf(c_0_5, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause100)).
% 0.03/0.28  cnf(c_0_6, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.03/0.28  cnf(c_0_7, axiom, (ssItem(hd(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause76)).
% 0.03/0.28  cnf(c_0_8, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.03/0.28  cnf(c_0_9, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause96)).
% 0.03/0.28  cnf(c_0_10, axiom, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause97)).
% 0.03/0.28  cnf(c_0_11, axiom, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause104)).
% 0.03/0.28  cnf(c_0_12, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.03/0.28  cnf(c_0_13, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
% 0.03/0.28  cnf(c_0_14, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause144)).
% 0.03/0.28  cnf(c_0_15, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.03/0.28  cnf(c_0_16, axiom, (ssList(tl(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause75)).
% 0.03/0.28  cnf(c_0_17, axiom, (singletonP(X2)|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause116)).
% 0.03/0.28  cnf(c_0_18, negated_conjecture, (~ssList(X1)|~neq(X1,nil)|~frontsegP(sk2,X1)|~frontsegP(sk1,X1)|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_11)).
% 0.03/0.28  cnf(c_0_19, axiom, (frontsegP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause61)).
% 0.03/0.28  cnf(c_0_20, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause60)).
% 0.03/0.28  cnf(c_0_21, axiom, (~singletonP(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause11)).
% 0.03/0.28  cnf(c_0_22, negated_conjecture, (neq(sk2,nil)), inference(cn,[status(thm)],[c_0_0])).
% 0.03/0.28  cnf(c_0_23, negated_conjecture, (neq(sk4,nil)), inference(rw,[status(thm)],[c_0_22, c_0_1])).
% 0.03/0.28  cnf(c_0_24, negated_conjecture, (cons(hd(sk4),nil)=sk1|~neq(nil,sk4)|~ssItem(hd(sk4))|~ssList(cons(hd(sk4),nil))), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_2, c_0_3]), c_0_23])])])])).
% 0.03/0.28  cnf(c_0_25, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_4, c_0_1])).
% 0.03/0.28  cnf(c_0_26, negated_conjecture, (cons(hd(sk4),nil)=sk1|sk4=nil|~ssItem(hd(sk4))|~ssList(cons(hd(sk4),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_5]), c_0_6]), c_0_25])])).
% 0.03/0.28  cnf(c_0_27, negated_conjecture, (cons(hd(sk4),nil)=sk1|sk4=nil|~ssList(cons(hd(sk4),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_7]), c_0_25])])).
% 0.03/0.28  cnf(c_0_28, negated_conjecture, (cons(hd(sk4),nil)=sk1|sk4=nil|~ssItem(hd(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_8]), c_0_6])])).
% 0.03/0.28  cnf(c_0_29, negated_conjecture, (cons(hd(sk4),nil)=sk1|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_7]), c_0_25])])).
% 0.03/0.28  cnf(c_0_30, negated_conjecture, (tl(sk1)=nil|sk4=nil|~ssItem(hd(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_29]), c_0_6])])).
% 0.03/0.28  cnf(c_0_31, negated_conjecture, (tl(sk1)=nil|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_7]), c_0_25])])).
% 0.03/0.28  cnf(c_0_32, negated_conjecture, (hd(sk1)=hd(sk4)|sk4=nil|~ssItem(hd(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_29]), c_0_6])])).
% 0.03/0.28  cnf(c_0_33, negated_conjecture, (cons(hd(sk1),nil)=sk1|sk4=nil|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_31]), c_0_12])])).
% 0.03/0.28  cnf(c_0_34, negated_conjecture, (hd(sk1)=hd(sk4)|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_7]), c_0_25])])).
% 0.03/0.28  cnf(c_0_35, negated_conjecture, (cons(hd(sk1),X1)=app(sk1,X1)|sk1=nil|sk4=nil|~ssItem(hd(sk1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_13, c_0_33])).
% 0.03/0.28  cnf(c_0_36, negated_conjecture, (sk4=nil|sk1=nil|ssItem(hd(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_34]), c_0_12])])).
% 0.03/0.28  cnf(c_0_37, negated_conjecture, (cons(hd(sk4),X1)=app(sk1,X1)|sk4=nil|sk1=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_34]), c_0_36])).
% 0.03/0.28  cnf(c_0_38, negated_conjecture, (app(sk1,tl(sk4))=sk4|sk1=nil|sk4=nil|~ssList(tl(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_37]), c_0_25])])).
% 0.03/0.28  cnf(c_0_39, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_14]), c_0_15])).
% 0.03/0.28  cnf(c_0_40, negated_conjecture, (app(sk1,tl(sk4))=sk4|sk4=nil|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_16]), c_0_25])])).
% 0.03/0.28  cnf(c_0_41, plain, (singletonP(cons(X1,nil))|~ssItem(X1)|~ssList(cons(X1,nil))), inference(er,[status(thm)],[c_0_17])).
% 0.03/0.28  cnf(c_0_42, negated_conjecture, (~neq(X1,nil)|~frontsegP(sk1,X1)|~frontsegP(sk4,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_1]), c_0_23])])).
% 0.03/0.28  cnf(c_0_43, negated_conjecture, (sk1=nil|sk4=nil|frontsegP(sk4,sk1)|~ssList(tl(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_12])])).
% 0.03/0.28  cnf(c_0_44, plain, (singletonP(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_8]), c_0_6])])).
% 0.03/0.28  cnf(c_0_45, negated_conjecture, (~neq(sk1,nil)|~frontsegP(sk4,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_19]), c_0_12])])).
% 0.03/0.28  cnf(c_0_46, negated_conjecture, (sk4=nil|sk1=nil|frontsegP(sk4,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_16]), c_0_25])])).
% 0.03/0.28  cnf(c_0_47, negated_conjecture, (sk4=nil|singletonP(sk1)|~ssItem(hd(sk4))), inference(spm,[status(thm)],[c_0_44, c_0_29])).
% 0.03/0.28  cnf(c_0_48, negated_conjecture, (sk1=nil|sk4=nil|~neq(sk1,nil)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
% 0.03/0.28  cnf(c_0_49, negated_conjecture, (sk4=nil|singletonP(sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_7]), c_0_25])])).
% 0.03/0.28  cnf(c_0_50, negated_conjecture, (sk4=nil|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_5]), c_0_12]), c_0_6])])).
% 0.03/0.28  cnf(c_0_51, negated_conjecture, (~neq(nil,nil)|~frontsegP(sk4,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_20]), c_0_6]), c_0_12])])).
% 0.03/0.28  cnf(c_0_52, negated_conjecture, (sk4=nil), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_21])).
% 0.03/0.28  cnf(c_0_53, negated_conjecture, (~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_20]), c_0_25])])).
% 0.03/0.28  cnf(c_0_54, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_52]), c_0_53]), ['proof']).
% 0.03/0.28  # SZS output end CNFRefutation
% 0.03/0.28  # Training examples: 0 positive, 0 negative
% 0.03/0.28  
% 0.03/0.28  # -------------------------------------------------
% 0.03/0.28  # User time                : 0.048 s
% 0.03/0.28  # System time              : 0.003 s
% 0.03/0.28  # Total time               : 0.051 s
% 0.03/0.28  # Maximum resident set size: 1864 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
