<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC413+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC413+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC413+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC413%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC413+1.p">SWC413+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n147.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:51 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : -

% Comments   : 
%------------------------------------------------------------------------------
%----WARNING: Could not form TPTP format derivation
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC413+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n147.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:43:38 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.62  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.62  # and selection function SelectCQPrecWNTNp.
% 0.06/0.62  #
% 0.06/0.62  # Preprocessing time       : 0.014 s
% 0.06/0.62  # Presaturation interreduction done
% 0.06/0.62  
% 0.06/0.62  # Proof found!
% 0.06/0.62  # SZS status Theorem
% 0.06/0.62  # SZS output start CNFRefutation
% 0.06/0.62  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((?[X5]:(ssItem(X5)&?[X6]:(ssItem(X6)&?[X7]:((ssList(X7)&app(app(cons(X5,nil),cons(X6,nil)),X7)=X2)&app(app(cons(X6,nil),cons(X5,nil)),X7)=X1)))|![X8]:(ssItem(X8)=>![X9]:(ssItem(X9)=>![X10]:(ssList(X10)=>app(app(cons(X8,nil),cons(X9,nil)),X10)!=X2))))|![X11]:(ssItem(X11)=>![X12]:(ssItem(X12)=>![X13]:(ssList(X13)=>(app(app(cons(X11,nil),cons(X12,nil)),X13)!=X4|app(app(cons(X12,nil),cons(X11,nil)),X13)!=X3)))))&(?[X14]:(ssItem(X14)&?[X15]:(ssItem(X15)&?[X16]:(ssList(X16)&app(app(cons(X14,nil),cons(X15,nil)),X16)=X4)))|![X8]:(ssItem(X8)=>![X9]:(ssItem(X9)=>![X10]:(ssList(X10)=>app(app(cons(X8,nil),cons(X9,nil)),X10)!=X2))))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.06/0.62  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.06/0.62  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.06/0.62  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax27)).
% 0.06/0.62  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.62  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax23)).
% 0.06/0.62  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax25)).
% 0.06/0.62  fof(c_0_7, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((?[X5]:(ssItem(X5)&?[X6]:(ssItem(X6)&?[X7]:((ssList(X7)&app(app(cons(X5,nil),cons(X6,nil)),X7)=X2)&app(app(cons(X6,nil),cons(X5,nil)),X7)=X1)))|![X8]:(ssItem(X8)=>![X9]:(ssItem(X9)=>![X10]:(ssList(X10)=>app(app(cons(X8,nil),cons(X9,nil)),X10)!=X2))))|![X11]:(ssItem(X11)=>![X12]:(ssItem(X12)=>![X13]:(ssList(X13)=>(app(app(cons(X11,nil),cons(X12,nil)),X13)!=X4|app(app(cons(X12,nil),cons(X11,nil)),X13)!=X3)))))&(?[X14]:(ssItem(X14)&?[X15]:(ssItem(X15)&?[X16]:(ssList(X16)&app(app(cons(X14,nil),cons(X15,nil)),X16)=X4)))|![X8]:(ssItem(X8)=>![X9]:(ssItem(X9)=>![X10]:(ssList(X10)=>app(app(cons(X8,nil),cons(X9,nil)),X10)!=X2)))))))))))), inference(assume_negation,[status(cth)],[c_0_0])).
% 0.06/0.62  fof(c_0_8, negated_conjecture, 
% (![X21]:![X22]:![X23]:![X30]:![X31]:![X32]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((esk49_0=esk51_0&esk48_0=esk50_0)&(((((~ssItem(X30)|(~ssItem(X31)|(~ssList(X32)|app(app(cons(X30,nil),cons(X31,nil)),X32)!=esk51_0)))|(~ssItem(X21)|(~ssItem(X22)|((~ssList(X23)|app(app(cons(X21,nil),cons(X22,nil)),X23)!=esk49_0)|app(app(cons(X22,nil),cons(X21,nil)),X23)!=esk48_0))))&((ssItem(esk58_0)|(~ssItem(X21)|(~ssItem(X22)|((~ssList(X23)|app(app(cons(X21,nil),cons(X22,nil)),X23)!=esk49_0)|app(app(cons(X22,nil),cons(X21,nil)),X23)!=esk48_0))))&((ssItem(esk59_0)|(~ssItem(X21)|(~ssItem(X22)|((~ssList(X23)|app(app(cons(X21,nil),cons(X22,nil)),X23)!=esk49_0)|app(app(cons(X22,nil),cons(X21,nil)),X23)!=esk48_0))))&((ssList(esk60_0)|(~ssItem(X21)|(~ssItem(X22)|((~ssList(X23)|app(app(cons(X21,nil),cons(X22,nil)),X23)!=esk49_0)|app(app(cons(X22,nil),cons(X21,nil)),X23)!=esk48_0))))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|(~ssItem(X21)|(~ssItem(X22)|((~ssList(X23)|app(app(cons(X21,n
% il),cons(X22,nil)),X23)!=esk49_0)|app(app(cons(X22,nil),cons(X21,nil)),X23)!=esk48_0))))))))&((((~ssItem(X30)|(~ssItem(X31)|(~ssList(X32)|app(app(cons(X30,nil),cons(X31,nil)),X32)!=esk51_0)))|ssItem(esk52_0))&((ssItem(esk58_0)|ssItem(esk52_0))&((ssItem(esk59_0)|ssItem(esk52_0))&((ssList(esk60_0)|ssItem(esk52_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|ssItem(esk52_0))))))&((((~ssItem(X30)|(~ssItem(X31)|(~ssList(X32)|app(app(cons(X30,nil),cons(X31,nil)),X32)!=esk51_0)))|ssItem(esk53_0))&((ssItem(esk58_0)|ssItem(esk53_0))&((ssItem(esk59_0)|ssItem(esk53_0))&((ssList(esk60_0)|ssItem(esk53_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|ssItem(esk53_0))))))&((((~ssItem(X30)|(~ssItem(X31)|(~ssList(X32)|app(app(cons(X30,nil),cons(X31,nil)),X32)!=esk51_0)))|ssList(esk54_0))&((ssItem(esk58_0)|ssList(esk54_0))&((ssItem(esk59_0)|ssList(esk54_0))&((ssList(esk60_0)|ssList(esk54_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|ssList(esk54_0))))))&(((~ssItem(X30)|
% (~ssItem(X31)|(~ssList(X32)|app(app(cons(X30,nil),cons(X31,nil)),X32)!=esk51_0)))|app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0)&((ssItem(esk58_0)|app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0)&((ssItem(esk59_0)|app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0)&((ssList(esk60_0)|app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0)&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0)))))))))&((((~ssItem(X30)|(~ssItem(X31)|(~ssList(X32)|app(app(cons(X30,nil),cons(X31,nil)),X32)!=esk51_0)))|ssItem(esk55_0))&((ssItem(esk58_0)|ssItem(esk55_0))&((ssItem(esk59_0)|ssItem(esk55_0))&((ssList(esk60_0)|ssItem(esk55_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|ssItem(esk55_0))))))&((((~ssItem(X30)|(~ssItem(X31)|(~ssList(X32)|app(app(cons(X30,nil),cons(X31,nil)),X32)!=esk51_0)))|ssItem(esk56_0))&((ssItem(esk58_0)|ssItem(esk56_0))&((ssItem(esk59_0)|ssItem(esk56_0))&((ssList(es
% k60_0)|ssItem(esk56_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|ssItem(esk56_0))))))&((((~ssItem(X30)|(~ssItem(X31)|(~ssList(X32)|app(app(cons(X30,nil),cons(X31,nil)),X32)!=esk51_0)))|ssList(esk57_0))&((ssItem(esk58_0)|ssList(esk57_0))&((ssItem(esk59_0)|ssList(esk57_0))&((ssList(esk60_0)|ssList(esk57_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|ssList(esk57_0))))))&((((~ssItem(X30)|(~ssItem(X31)|(~ssList(X32)|app(app(cons(X30,nil),cons(X31,nil)),X32)!=esk51_0)))|app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0)&((ssItem(esk58_0)|app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0)&((ssItem(esk59_0)|app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0)&((ssList(esk60_0)|app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0)&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0)))))&(((~ssItem(X30)|(~ssItem(X31)|(~ssList(X32)|app(app(cons(X30,nil),cons(X31,n
% il)),X32)!=esk51_0)))|app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0)&((ssItem(esk58_0)|app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0)&((ssItem(esk59_0)|app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0)&((ssList(esk60_0)|app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0)&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0)))))))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
% 0.06/0.62  cnf(c_0_9, negated_conjecture, (ssList(esk57_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_10, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_11, negated_conjecture, (ssItem(esk55_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_12, negated_conjecture, (ssItem(esk56_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_13, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_14, negated_conjecture, (ssList(esk57_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_9, c_0_10])).
% 0.06/0.62  cnf(c_0_15, negated_conjecture, (ssList(esk57_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_16, negated_conjecture, (ssList(esk57_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_17, negated_conjecture, (ssList(esk57_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_18, negated_conjecture, (ssList(esk57_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_19, negated_conjecture, (ssItem(esk55_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_11, c_0_10])).
% 0.06/0.62  cnf(c_0_20, negated_conjecture, (ssItem(esk55_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_21, negated_conjecture, (ssItem(esk55_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_22, negated_conjecture, (ssItem(esk55_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_23, negated_conjecture, (ssItem(esk55_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_24, negated_conjecture, (ssItem(esk56_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_12, c_0_10])).
% 0.06/0.62  cnf(c_0_25, negated_conjecture, (ssItem(esk56_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_26, negated_conjecture, (ssItem(esk56_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_27, negated_conjecture, (ssItem(esk56_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_28, negated_conjecture, (ssItem(esk56_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_29, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_30, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_31, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_32, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_33, negated_conjecture, (app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk48_0|app(app(cons(X2,nil),cons(X1,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X1)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_34, negated_conjecture, (app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk49_0), inference(rw,[status(thm)],[c_0_13, c_0_10])).
% 0.06/0.62  cnf(c_0_35, negated_conjecture, (ssList(esk57_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17]), c_0_18])).
% 0.06/0.62  cnf(c_0_36, negated_conjecture, (ssItem(esk55_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22]), c_0_23])).
% 0.06/0.62  cnf(c_0_37, negated_conjecture, (ssItem(esk56_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27]), c_0_28])).
% 0.06/0.62  cnf(c_0_38, negated_conjecture, (app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk48_0), inference(rw,[status(thm)],[c_0_29, c_0_30])).
% 0.06/0.62  fof(c_0_39, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.06/0.62  cnf(c_0_40, negated_conjecture, (ssItem(esk58_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk48_0|app(app(cons(X2,nil),cons(X1,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X1)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_41, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk49_0|ssItem(esk58_0)), inference(rw,[status(thm)],[c_0_31, c_0_10])).
% 0.06/0.62  cnf(c_0_42, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk48_0|ssItem(esk58_0)), inference(rw,[status(thm)],[c_0_32, c_0_30])).
% 0.06/0.62  cnf(c_0_43, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_44, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_45, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_46, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_47, negated_conjecture, (app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
% 0.06/0.62  cnf(c_0_48, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.62  cnf(c_0_49, negated_conjecture, (ssItem(esk58_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_35]), c_0_36]), c_0_37])]), c_0_42])).
% 0.06/0.62  fof(c_0_50, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.06/0.62  cnf(c_0_51, negated_conjecture, (ssItem(esk59_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk48_0|app(app(cons(X2,nil),cons(X1,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X1)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_52, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk49_0|ssItem(esk59_0)), inference(rw,[status(thm)],[c_0_43, c_0_10])).
% 0.06/0.62  cnf(c_0_53, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk48_0|ssItem(esk59_0)), inference(rw,[status(thm)],[c_0_44, c_0_30])).
% 0.06/0.62  cnf(c_0_54, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_55, negated_conjecture, (ssList(esk60_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk48_0|app(app(cons(X2,nil),cons(X1,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X1)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_56, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk49_0|ssList(esk60_0)), inference(rw,[status(thm)],[c_0_45, c_0_10])).
% 0.06/0.62  cnf(c_0_57, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk48_0|ssList(esk60_0)), inference(rw,[status(thm)],[c_0_46, c_0_30])).
% 0.06/0.62  cnf(c_0_58, negated_conjecture, (ssItem(esk52_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  fof(c_0_59, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
% 0.06/0.62  cnf(c_0_60, negated_conjecture, (app(cons(esk58_0,cons(esk59_0,nil)),esk60_0)=esk49_0|~ssList(cons(esk59_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])])).
% 0.06/0.62  cnf(c_0_61, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
% 0.06/0.62  cnf(c_0_62, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.06/0.62  cnf(c_0_63, negated_conjecture, (ssItem(esk59_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_35]), c_0_36]), c_0_37])]), c_0_53])).
% 0.06/0.62  cnf(c_0_64, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_54, c_0_10])).
% 0.06/0.62  cnf(c_0_65, negated_conjecture, (ssList(esk60_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_35]), c_0_36]), c_0_37])]), c_0_57])).
% 0.06/0.62  cnf(c_0_66, negated_conjecture, (ssItem(esk52_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_58, c_0_10])).
% 0.06/0.62  cnf(c_0_67, negated_conjecture, (ssItem(esk52_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_68, negated_conjecture, (ssItem(esk52_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_69, negated_conjecture, (ssItem(esk52_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_70, negated_conjecture, (ssItem(esk52_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_71, negated_conjecture, (ssItem(esk53_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_72, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_59])).
% 0.06/0.62  cnf(c_0_73, negated_conjecture, (app(cons(esk58_0,cons(esk59_0,nil)),esk60_0)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63])])).
% 0.06/0.62  cnf(c_0_74, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_47]), c_0_65]), c_0_63]), c_0_49])])).
% 0.06/0.62  cnf(c_0_75, negated_conjecture, (ssItem(esk52_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68]), c_0_69]), c_0_70])).
% 0.06/0.62  cnf(c_0_76, negated_conjecture, (ssItem(esk53_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_71, c_0_10])).
% 0.06/0.62  cnf(c_0_77, negated_conjecture, (ssItem(esk53_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_78, negated_conjecture, (ssItem(esk53_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_79, negated_conjecture, (ssItem(esk53_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_80, negated_conjecture, (ssItem(esk53_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_81, negated_conjecture, (ssList(esk54_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_82, negated_conjecture, (cons(esk58_0,app(cons(esk59_0,nil),esk60_0))=esk49_0|~ssList(cons(esk59_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_65]), c_0_49])])).
% 0.06/0.62  cnf(c_0_83, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_84, negated_conjecture, (app(cons(esk52_0,cons(esk53_0,nil)),esk54_0)=esk49_0|~ssList(cons(esk53_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_48]), c_0_75])])).
% 0.06/0.62  cnf(c_0_85, negated_conjecture, (ssItem(esk53_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78]), c_0_79]), c_0_80])).
% 0.06/0.62  cnf(c_0_86, negated_conjecture, (ssList(esk54_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_81, c_0_10])).
% 0.06/0.62  cnf(c_0_87, negated_conjecture, (ssList(esk54_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_88, negated_conjecture, (ssList(esk54_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_89, negated_conjecture, (ssList(esk54_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_90, negated_conjecture, (ssList(esk54_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  fof(c_0_91, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 0.06/0.62  cnf(c_0_92, negated_conjecture, (cons(esk58_0,app(cons(esk59_0,nil),esk60_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_61]), c_0_62]), c_0_63])])).
% 0.06/0.62  cnf(c_0_93, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk49_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_10]), c_0_10])).
% 0.06/0.62  cnf(c_0_94, negated_conjecture, (app(cons(esk52_0,cons(esk53_0,nil)),esk54_0)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_61]), c_0_62]), c_0_85])])).
% 0.06/0.62  cnf(c_0_95, negated_conjecture, (ssList(esk54_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88]), c_0_89]), c_0_90])).
% 0.06/0.62  cnf(c_0_96, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_91])).
% 0.06/0.62  cnf(c_0_97, negated_conjecture, (cons(esk58_0,cons(esk59_0,esk60_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_48]), c_0_65]), c_0_63])])).
% 0.06/0.62  cnf(c_0_98, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_47]), c_0_65]), c_0_63]), c_0_49])])).
% 0.06/0.62  cnf(c_0_99, negated_conjecture, (cons(esk52_0,app(cons(esk53_0,nil),esk54_0))=esk49_0|~ssList(cons(esk53_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_94]), c_0_95]), c_0_75])])).
% 0.06/0.62  cnf(c_0_100, negated_conjecture, (hd(esk49_0)=esk58_0|~ssList(cons(esk59_0,esk60_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_49])])).
% 0.06/0.62  cnf(c_0_101, negated_conjecture, (app(cons(esk55_0,cons(esk56_0,nil)),esk57_0)=esk49_0|~ssList(cons(esk56_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_48]), c_0_36])])).
% 0.06/0.62  cnf(c_0_102, negated_conjecture, (cons(esk52_0,app(cons(esk53_0,nil),esk54_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_61]), c_0_62]), c_0_85])])).
% 0.06/0.62  cnf(c_0_103, negated_conjecture, (hd(esk49_0)=esk58_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_61]), c_0_65]), c_0_63])])).
% 0.06/0.62  cnf(c_0_104, negated_conjecture, (app(cons(esk55_0,cons(esk56_0,nil)),esk57_0)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_61]), c_0_62]), c_0_37])])).
% 0.06/0.62  cnf(c_0_105, negated_conjecture, (esk52_0=esk58_0|~ssList(app(cons(esk53_0,nil),esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_102]), c_0_103]), c_0_75])])).
% 0.06/0.62  fof(c_0_106, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
% 0.06/0.62  cnf(c_0_107, negated_conjecture, (cons(esk55_0,app(cons(esk56_0,nil),esk57_0))=esk49_0|~ssList(cons(esk56_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_104]), c_0_35]), c_0_36])])).
% 0.06/0.62  cnf(c_0_108, negated_conjecture, (esk52_0=esk58_0|~ssList(cons(esk53_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_48]), c_0_95]), c_0_85])])).
% 0.06/0.62  cnf(c_0_109, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_106])).
% 0.06/0.62  cnf(c_0_110, negated_conjecture, (cons(esk55_0,app(cons(esk56_0,nil),esk57_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_61]), c_0_62]), c_0_37])])).
% 0.06/0.62  cnf(c_0_111, negated_conjecture, (cons(esk52_0,cons(esk53_0,esk54_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_48]), c_0_95]), c_0_85])])).
% 0.06/0.62  cnf(c_0_112, negated_conjecture, (esk52_0=esk58_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_61]), c_0_95]), c_0_85])])).
% 0.06/0.62  cnf(c_0_113, negated_conjecture, (tl(esk49_0)=cons(esk59_0,esk60_0)|~ssList(cons(esk59_0,esk60_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_97]), c_0_49])])).
% 0.06/0.62  cnf(c_0_114, negated_conjecture, (esk55_0=esk58_0|~ssList(app(cons(esk56_0,nil),esk57_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_110]), c_0_103]), c_0_36])])).
% 0.06/0.62  cnf(c_0_115, negated_conjecture, (cons(esk58_0,cons(esk53_0,esk54_0))=esk49_0), inference(rw,[status(thm)],[c_0_111, c_0_112])).
% 0.06/0.62  cnf(c_0_116, negated_conjecture, (tl(esk49_0)=cons(esk59_0,esk60_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_61]), c_0_65]), c_0_63])])).
% 0.06/0.62  cnf(c_0_117, negated_conjecture, (esk55_0=esk58_0|~ssList(cons(esk56_0,esk57_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_48]), c_0_35]), c_0_37])])).
% 0.06/0.62  cnf(c_0_118, negated_conjecture, (cons(esk53_0,esk54_0)=cons(esk59_0,esk60_0)|~ssList(cons(esk53_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_115]), c_0_116]), c_0_49])])).
% 0.06/0.62  cnf(c_0_119, negated_conjecture, (cons(esk55_0,cons(esk56_0,esk57_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_48]), c_0_35]), c_0_37])])).
% 0.06/0.62  cnf(c_0_120, negated_conjecture, (esk55_0=esk58_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_61]), c_0_35]), c_0_37])])).
% 0.06/0.62  cnf(c_0_121, negated_conjecture, (cons(esk53_0,esk54_0)=cons(esk59_0,esk60_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_61]), c_0_95]), c_0_85])])).
% 0.06/0.62  cnf(c_0_122, negated_conjecture, (cons(esk58_0,cons(esk56_0,esk57_0))=esk49_0), inference(rw,[status(thm)],[c_0_119, c_0_120])).
% 0.06/0.62  cnf(c_0_123, negated_conjecture, (hd(cons(esk59_0,esk60_0))=esk53_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_121]), c_0_95]), c_0_85])])).
% 0.06/0.62  cnf(c_0_124, negated_conjecture, (cons(esk56_0,esk57_0)=cons(esk59_0,esk60_0)|~ssList(cons(esk56_0,esk57_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_122]), c_0_116]), c_0_49])])).
% 0.06/0.62  cnf(c_0_125, negated_conjecture, (esk53_0=esk59_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_123]), c_0_65]), c_0_63])])).
% 0.06/0.62  cnf(c_0_126, negated_conjecture, (cons(esk56_0,esk57_0)=cons(esk59_0,esk60_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_61]), c_0_35]), c_0_37])])).
% 0.06/0.62  cnf(c_0_127, negated_conjecture, (hd(cons(esk59_0,esk60_0))=esk59_0), inference(rw,[status(thm)],[c_0_123, c_0_125])).
% 0.06/0.62  cnf(c_0_128, negated_conjecture, (cons(esk59_0,esk54_0)=cons(esk59_0,esk60_0)), inference(rw,[status(thm)],[c_0_121, c_0_125])).
% 0.06/0.62  cnf(c_0_129, negated_conjecture, (esk56_0=esk59_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_126]), c_0_127]), c_0_35]), c_0_37])])).
% 0.06/0.62  cnf(c_0_130, negated_conjecture, (tl(cons(esk59_0,esk60_0))=esk54_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_128]), c_0_95]), c_0_63])])).
% 0.06/0.62  cnf(c_0_131, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_132, negated_conjecture, (cons(esk59_0,esk57_0)=cons(esk59_0,esk60_0)), inference(rw,[status(thm)],[c_0_126, c_0_129])).
% 0.06/0.62  cnf(c_0_133, negated_conjecture, (esk54_0=esk60_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_130]), c_0_65]), c_0_63])])).
% 0.06/0.62  cnf(c_0_134, negated_conjecture, (app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk48_0|app(app(cons(X2,nil),cons(X1,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X1)|~ssItem(X2)|app(app(cons(X4,nil),cons(X5,nil)),X6)!=esk51_0|~ssList(X6)|~ssItem(X5)|~ssItem(X4)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.62  cnf(c_0_135, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk48_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_30]), c_0_10])).
% 0.06/0.62  cnf(c_0_136, negated_conjecture, (esk57_0=esk60_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_132]), c_0_130]), c_0_133]), c_0_35]), c_0_63])])).
% 0.06/0.62  cnf(c_0_137, negated_conjecture, (app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|app(app(cons(X4,nil),cons(X5,nil)),X6)!=esk49_0|app(app(cons(X5,nil),cons(X4,nil)),X6)!=esk48_0|~ssList(X3)|~ssList(X6)|~ssItem(X2)|~ssItem(X1)|~ssItem(X4)|~ssItem(X5)), inference(rw,[status(thm)],[c_0_134, c_0_10])).
% 0.06/0.62  cnf(c_0_138, negated_conjecture, (app(app(cons(esk59_0,nil),cons(esk58_0,nil)),esk60_0)=esk48_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_135, c_0_120]), c_0_129]), c_0_136])).
% 0.06/0.62  cnf(c_0_139, negated_conjecture, (app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk49_0|app(app(cons(X2,nil),cons(X1,nil)),X3)!=esk48_0|~ssList(X3)|~ssItem(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_47]), c_0_65]), c_0_63]), c_0_49])])).
% 0.06/0.62  cnf(c_0_140, negated_conjecture, (app(app(cons(esk59_0,nil),cons(esk58_0,nil)),esk60_0)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_47]), c_0_65]), c_0_63]), c_0_49])])).
% 0.06/0.62  cnf(c_0_141, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_47]), c_0_140]), c_0_65]), c_0_49]), c_0_63])]), ['proof']).
% 0.06/0.62  # SZS output end CNFRefutation
% 0.06/0.62  # Training examples: 0 positive, 0 negative
% 0.06/0.63  
% 0.06/0.63  # -------------------------------------------------
% 0.06/0.63  # User time                : 0.389 s
% 0.06/0.63  # System time              : 0.008 s
% 0.06/0.63  # Total time               : 0.397 s
% 0.06/0.63  # Maximum resident set size: 1892 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
