<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC045-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC045-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC045-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC045-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC045-1.p">SWC045-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n155.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:25 EDT 2016

% Result     : Unsatisfiable 173.73s
% Output     : CNFRefutation 173.73s
% Verified   : 
% Statistics : Number of clauses        :   50 ( 116 expanded)
%              Number of leaves         :   26 (  74 expanded)
%              Depth                    :    9
%              Number of atoms          :  132 ( 281 expanded)
%              Number of equality atoms :   36 (  71 expanded)
%              Maximal clause size      :    6 (   3 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_1"></A>cnf(c_0_1,axiom,
    ( cons(skaf44(X1),nil) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause101)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( ssItem(skaf44(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause47)).

<A NAME="c_0_3"></A>cnf(c_0_3,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_4"></A>cnf(c_0_4,axiom,
    ( singletonP(X2)
    | cons(X1,nil) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause116)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( nil = X2
    | tl(app(X2,X1)) = app(tl(X2),X1)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause133)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ~ singletonP(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause11)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( nil = X1
    | ~ segmentP(nil,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause80)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( nil = sk2 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( segmentP(sk4,app(app(cons(X1,nil),sk5(X1)),cons(X1,nil)))
    | ~ ssItem(X1)
    | ~ memberP(sk3,X1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause157)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( cons(X1,X2) != nil
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause98)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( memberP(cons(X2,X3),X1)
    | X1 != X2
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause138)).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( ssList(sk5(X1))
    | ~ ssItem(X1)
    | ~ memberP(sk3,X1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( cons(skaf83(X1),skaf82(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause109)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( ssItem(skaf83(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause12)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( ssList(skaf82(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause13)).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( nil != sk1 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_26"></A>cnf(c_0_26,plain,
    ( tl(X1) = nil
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,c_0_1]),c_0_2]),c_0_3])])).

<A NAME="c_0_27"></A>cnf(c_0_27,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1)
    | ~ ssList(cons(X1,nil)) ),
    inference(er,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,plain,
    ( app(tl(X1),nil) = tl(X1)
    | nil = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,c_0_6]),c_0_3])])).

<A NAME="c_0_29"></A>cnf(c_0_29,plain,
    ( tl(cons(X1,nil)) = nil
    | ~ ssItem(X1)
    | ~ ssList(cons(X1,nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,plain,
    ( cons(X1,nil) != nil
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_4]),c_0_3])])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( app(nil,nil) = nil
    | ~ ssItem(X1)
    | ~ ssList(cons(X1,nil)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_29]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,plain,
    ( nil = X1
    | app(app(X2,X1),X3) != nil
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_9]),c_0_3])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( sk4 = nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,plain,
    ( app(nil,nil) = nil
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_12]),c_0_3])])).

<A NAME="c_0_35"></A>cnf(c_0_35,plain,
    ( nil = X1
    | app(X2,app(X1,X3)) != nil
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( segmentP(nil,app(app(cons(X1,nil),sk5(X1)),cons(X1,nil)))
    | ~ memberP(sk3,X1)
    | ~ ssItem(X1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,<A HREF="#c_0_33">c_0_33</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,plain,
    ( ssList(cons(X1,app(X2,X3)))
    | ~ ssItem(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_16]),<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,plain,
    ( app(nil,nil) = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_39"></A>cnf(c_0_39,plain,
    ( app(X1,cons(X2,app(X3,X4))) != nil
    | ~ ssItem(X2)
    | ~ ssList(X4)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_16]),c_0_12]),<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( app(app(cons(X1,nil),sk5(X1)),cons(X1,nil)) = nil
    | ~ memberP(sk3,X1)
    | ~ ssItem(X1)
    | ~ ssList(app(app(cons(X1,nil),sk5(X1)),cons(X1,nil))) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_38]),c_0_3])])).

<A NAME="c_0_42"></A>cnf(c_0_42,plain,
    ( app(X1,cons(X2,nil)) != nil
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_38]),c_0_3])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( ~ memberP(sk3,X1)
    | ~ ssItem(X1)
    | ~ ssList(app(cons(X1,nil),sk5(X1))) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_15]),c_0_41]),<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( ~ memberP(sk3,X1)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_15]),c_0_19]),<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,plain,
    ( nil = X1
    | memberP(X1,skaf83(X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_20]),c_0_21]),c_0_22])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( sk3 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_46]),c_0_21]),c_0_47])]),<A HREF="#c_0_48">c_0_48</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC045-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n155.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:34:38 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 151.04/151.28  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 151.04/151.28  # and selection function SelectCQArNTEqFirst.
% 151.04/151.28  #
% 151.04/151.28  # Preprocessing time       : 0.009 s
% 151.04/151.28  # Presaturation interreduction done
% 151.44/151.66  # No success with AutoSched0
% 151.44/151.66  # Trying AutoSched1 for 73 seconds
% 173.73/173.97  # AutoSched1-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S00
% 173.73/173.97  # and selection function NoSelection.
% 173.73/173.97  #
% 173.73/173.97  # Preprocessing time       : 0.010 s
% 173.73/173.97  # Presaturation interreduction done
% 173.73/173.97  
% 173.73/173.97  # Proof found!
% 173.73/173.97  # SZS status Unsatisfiable
% 173.73/173.97  # SZS output start CNFRefutation
% 173.73/173.97  cnf(c_0_0, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause96)).
% 173.73/173.97  cnf(c_0_1, axiom, (cons(skaf44(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause101)).
% 173.73/173.97  cnf(c_0_2, axiom, (ssItem(skaf44(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause47)).
% 173.73/173.97  cnf(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 173.73/173.97  cnf(c_0_4, axiom, (singletonP(X2)|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause116)).
% 173.73/173.97  cnf(c_0_5, axiom, (nil=X2|tl(app(X2,X1))=app(tl(X2),X1)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause133)).
% 173.73/173.97  cnf(c_0_6, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause73)).
% 173.73/173.97  cnf(c_0_7, axiom, (~singletonP(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause11)).
% 173.73/173.97  cnf(c_0_8, axiom, (nil=X1|~segmentP(nil,X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause80)).
% 173.73/173.97  cnf(c_0_9, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause173)).
% 173.73/173.97  cnf(c_0_10, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 173.73/173.97  cnf(c_0_11, negated_conjecture, (nil=sk2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 173.73/173.97  cnf(c_0_12, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause86)).
% 173.73/173.97  cnf(c_0_13, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause149)).
% 173.73/173.97  cnf(c_0_14, negated_conjecture, (segmentP(sk4,app(app(cons(X1,nil),sk5(X1)),cons(X1,nil)))|~ssItem(X1)|~memberP(sk3,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_11)).
% 173.73/173.97  cnf(c_0_15, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
% 173.73/173.97  cnf(c_0_16, axiom, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause157)).
% 173.73/173.97  cnf(c_0_17, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause98)).
% 173.73/173.97  cnf(c_0_18, axiom, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause138)).
% 173.73/173.97  cnf(c_0_19, negated_conjecture, (ssList(sk5(X1))|~ssItem(X1)|~memberP(sk3,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
% 173.73/173.97  cnf(c_0_20, axiom, (cons(skaf83(X1),skaf82(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause109)).
% 173.73/173.97  cnf(c_0_21, axiom, (ssItem(skaf83(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause12)).
% 173.73/173.97  cnf(c_0_22, axiom, (ssList(skaf82(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause13)).
% 173.73/173.97  cnf(c_0_23, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 173.73/173.97  cnf(c_0_24, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 173.73/173.97  cnf(c_0_25, negated_conjecture, (nil!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
% 173.73/173.97  cnf(c_0_26, plain, (tl(X1)=nil|~singletonP(X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2]), c_0_3])])).
% 173.73/173.97  cnf(c_0_27, plain, (singletonP(cons(X1,nil))|~ssItem(X1)|~ssList(cons(X1,nil))), inference(er,[status(thm)],[c_0_4])).
% 173.73/173.97  cnf(c_0_28, plain, (app(tl(X1),nil)=tl(X1)|nil=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_6]), c_0_3])])).
% 173.73/173.97  cnf(c_0_29, plain, (tl(cons(X1,nil))=nil|~ssItem(X1)|~ssList(cons(X1,nil))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
% 173.73/173.97  cnf(c_0_30, plain, (cons(X1,nil)!=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_4]), c_0_3])])).
% 173.73/173.97  cnf(c_0_31, plain, (app(nil,nil)=nil|~ssItem(X1)|~ssList(cons(X1,nil))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
% 173.73/173.97  cnf(c_0_32, plain, (nil=X1|app(app(X2,X1),X3)!=nil|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_3])])).
% 173.73/173.97  cnf(c_0_33, negated_conjecture, (sk4=nil), inference(rw,[status(thm)],[c_0_10, c_0_11])).
% 173.73/173.97  cnf(c_0_34, plain, (app(nil,nil)=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_12]), c_0_3])])).
% 173.73/173.97  cnf(c_0_35, plain, (nil=X1|app(X2,app(X1,X3))!=nil|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(spm,[status(thm)],[c_0_32, c_0_13])).
% 173.73/173.97  cnf(c_0_36, negated_conjecture, (segmentP(nil,app(app(cons(X1,nil),sk5(X1)),cons(X1,nil)))|~memberP(sk3,X1)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_14, c_0_33])).
% 173.73/173.97  cnf(c_0_37, plain, (ssList(cons(X1,app(X2,X3)))|~ssItem(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_12])).
% 173.73/173.97  cnf(c_0_38, plain, (app(nil,nil)=nil), inference(spm,[status(thm)],[c_0_34, c_0_2])).
% 173.73/173.97  cnf(c_0_39, plain, (app(X1,cons(X2,app(X3,X4)))!=nil|~ssItem(X2)|~ssList(X4)|~ssList(X1)|~ssList(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_16]), c_0_12]), c_0_17])).
% 173.73/173.97  cnf(c_0_40, negated_conjecture, (app(app(cons(X1,nil),sk5(X1)),cons(X1,nil))=nil|~memberP(sk3,X1)|~ssItem(X1)|~ssList(app(app(cons(X1,nil),sk5(X1)),cons(X1,nil)))), inference(spm,[status(thm)],[c_0_8, c_0_36])).
% 173.73/173.97  cnf(c_0_41, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_3])])).
% 173.73/173.97  cnf(c_0_42, plain, (app(X1,cons(X2,nil))!=nil|~ssItem(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_38]), c_0_3])])).
% 173.73/173.97  cnf(c_0_43, negated_conjecture, (~memberP(sk3,X1)|~ssItem(X1)|~ssList(app(cons(X1,nil),sk5(X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_15]), c_0_41]), c_0_42])).
% 173.73/173.97  cnf(c_0_44, plain, (memberP(cons(X1,X2),X1)|~ssItem(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_18])).
% 173.73/173.97  cnf(c_0_45, negated_conjecture, (~memberP(sk3,X1)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_15]), c_0_19]), c_0_41])).
% 173.73/173.97  cnf(c_0_46, plain, (nil=X1|memberP(X1,skaf83(X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_20]), c_0_21]), c_0_22])])).
% 173.73/173.97  cnf(c_0_47, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
% 173.73/173.97  cnf(c_0_48, negated_conjecture, (sk3!=nil), inference(rw,[status(thm)],[c_0_25, c_0_24])).
% 173.73/173.97  cnf(c_0_49, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_21]), c_0_47])]), c_0_48]), ['proof']).
% 173.73/173.97  # SZS output end CNFRefutation
% 173.73/173.97  # Training examples: 0 positive, 0 negative
% 173.83/174.03  
% 173.83/174.03  # -------------------------------------------------
% 173.83/174.03  # User time                : 169.259 s
% 173.83/174.03  # System time              : 4.577 s
% 173.83/174.03  # Total time               : 173.837 s
% 173.83/174.03  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
