<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC071+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC071+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC071+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC071%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC071+1.p">SWC071+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n168.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:32 EDT 2016

% Result     : Theorem 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of formulae       :   41 ( 188 expanded)
%              Number of clauses        :   29 (  67 expanded)
%              Number of leaves         :    6 (  46 expanded)
%              Depth                    :    9
%              Number of atoms          :  203 (3515 expanded)
%              Number of equality atoms :   64 (1363 expanded)
%              Maximal formula depth    :   29 (   5 average)
%              Maximal clause size      :   48 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ? [X5] :
                        ( ssList(X5)
                        & neq(X5,nil)
                        & segmentP(X2,X5)
                        & segmentP(X1,X5) )
                    | ( nil = X2
                      & nil = X1 )
                    | ( ! [X6] :
                          ( ssItem(X6)
                         => ! [X7] :
                              ( ssList(X7)
                             => ! [X8] :
                                  ( ssList(X8)
                                 => ( cons(X6,nil) != X3
                                    | app(app(X7,X3),X8) != X4
                                    | ? [X9] :
                                        ( ssItem(X9)
                                        & memberP(X7,X9)
                                        & lt(X6,X9) )
                                    | ? [X10] :
                                        ( ssItem(X10)
                                        & memberP(X8,X10)
                                        & lt(X10,X6) ) ) ) ) )
                      & ( nil != X4
                        | nil != X3 ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) != X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax18)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( segmentP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax7)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,X1) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax55)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_6"></A>fof(c_0_6,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ? [X5] :
                          ( ssList(X5)
                          & neq(X5,nil)
                          & segmentP(X2,X5)
                          & segmentP(X1,X5) )
                      | ( nil = X2
                        & nil = X1 )
                      | ( ! [X6] :
                            ( ssItem(X6)
                           => ! [X7] :
                                ( ssList(X7)
                               => ! [X8] :
                                    ( ssList(X8)
                                   => ( cons(X6,nil) != X3
                                      | app(app(X7,X3),X8) != X4
                                      | ? [X9] :
                                          ( ssItem(X9)
                                          & memberP(X7,X9)
                                          & lt(X6,X9) )
                                      | ? [X10] :
                                          ( ssItem(X10)
                                          & memberP(X8,X10)
                                          & lt(X10,X6) ) ) ) ) )
                        & ( nil != X4
                          | nil != X3 ) ) ) ) ) ) ) ),
    inference(assume_negation,[status(cth)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_7"></A>fof(c_0_7,negated_conjecture,(
    ! [X15,X19,X20] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ( ~ ssList(X15)
        | ~ neq(X15,nil)
        | ~ segmentP(esk49_0,X15)
        | ~ segmentP(esk48_0,X15) )
      & ( nil != esk49_0
        | nil != esk48_0 )
      & ( nil = esk51_0
        | ssItem(esk52_0) )
      & ( nil = esk50_0
        | ssItem(esk52_0) )
      & ( nil = esk51_0
        | ssList(esk53_0) )
      & ( nil = esk50_0
        | ssList(esk53_0) )
      & ( nil = esk51_0
        | ssList(esk54_0) )
      & ( nil = esk50_0
        | ssList(esk54_0) )
      & ( nil = esk51_0
        | cons(esk52_0,nil) = esk50_0 )
      & ( nil = esk50_0
        | cons(esk52_0,nil) = esk50_0 )
      & ( nil = esk51_0
        | app(app(esk53_0,esk50_0),esk54_0) = esk51_0 )
      & ( nil = esk50_0
        | app(app(esk53_0,esk50_0),esk54_0) = esk51_0 )
      & ( nil = esk51_0
        | ~ ssItem(X19)
        | ~ memberP(esk53_0,X19)
        | ~ lt(esk52_0,X19) )
      & ( nil = esk50_0
        | ~ ssItem(X19)
        | ~ memberP(esk53_0,X19)
        | ~ lt(esk52_0,X19) )
      & ( nil = esk51_0
        | ~ ssItem(X20)
        | ~ memberP(esk54_0,X20)
        | ~ lt(X20,esk52_0) )
      & ( nil = esk50_0
        | ~ ssItem(X20)
        | ~ memberP(esk54_0,X20)
        | ~ lt(X20,esk52_0) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).

<A NAME="c_0_8"></A>fof(c_0_8,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) != X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( nil != esk48_0
    | nil != esk49_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_12"></A>fof(c_0_12,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk8_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ssList(esk9_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(app(X9,X6),X10) != X5
        | segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).

<A NAME="c_0_13"></A>cnf(c_0_13,plain,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( cons(esk52_0,nil) = esk50_0
    | nil = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ssItem(esk52_0)
    | nil = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( esk50_0 != nil
    | esk51_0 != nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_10]),<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_19"></A>fof(c_0_19,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,plain,
    ( segmentP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(app(X3,X2),X4) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( app(app(esk53_0,esk50_0),esk54_0) = esk51_0
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( esk50_0 != nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_14]),c_0_15])]),c_0_16]),<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( ssList(esk54_0)
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( ssList(esk53_0)
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( ~ segmentP(esk50_0,X1)
    | ~ segmentP(esk51_0,X1)
    | ~ ssList(X1)
    | ~ neq(X1,nil) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_10]),<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,plain,
    ( segmentP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( app(app(esk53_0,esk50_0),esk54_0) = esk51_0 ),
    inference(sr,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( ssList(esk54_0) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( ssList(esk53_0) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( ~ segmentP(esk51_0,esk50_0)
    | ~ neq(esk50_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_28]),c_0_29])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( segmentP(esk51_0,esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31]),c_0_32]),c_0_33]),c_0_34]),c_0_29])])).

<A NAME="c_0_37"></A>fof(c_0_37,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( ~ neq(esk50_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_36])])).

<A NAME="c_0_39"></A>cnf(c_0_39,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_39]),c_0_15]),c_0_29])]),<A HREF="#c_0_23">c_0_23</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.04  % Problem    : SWC071+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.05  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.24  % Computer   : n168.star.cs.uiowa.edu
% 0.03/0.24  % Model      : x86_64 x86_64
% 0.03/0.24  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.24  % Memory     : 32218.75MB
% 0.03/0.24  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.24  % CPULimit   : 300
% 0.03/0.24  % DateTime   : Tue May 10 20:34:53 CDT 2016
% 0.03/0.24  % CPUTime    : 
% 0.03/0.24  # Version: 1.9.1pre011
% 0.03/0.24  # No SInE strategy applied
% 0.03/0.24  # Trying AutoSched0 for 151 seconds
% 0.03/0.29  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.03/0.29  # and selection function SelectCQPrecWNTNp.
% 0.03/0.29  #
% 0.03/0.29  # Preprocessing time       : 0.014 s
% 0.03/0.29  # Presaturation interreduction done
% 0.03/0.29  
% 0.03/0.29  # Proof found!
% 0.03/0.29  # SZS status Theorem
% 0.03/0.29  # SZS output start CNFRefutation
% 0.03/0.29  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))|(nil=X2&nil=X1))|(![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>(((cons(X6,nil)!=X3|app(app(X7,X3),X8)!=X4)|?[X9]:((ssItem(X9)&memberP(X7,X9))&lt(X6,X9)))|?[X10]:((ssItem(X10)&memberP(X8,X10))&lt(X10,X6))))))&(nil!=X4|nil!=X3)))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.03/0.29  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax18)).
% 0.03/0.29  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax7)).
% 0.03/0.29  fof(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.03/0.29  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>segmentP(X1,X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax55)).
% 0.03/0.29  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.03/0.29  fof(c_0_6, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))|(nil=X2&nil=X1))|(![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>(((cons(X6,nil)!=X3|app(app(X7,X3),X8)!=X4)|?[X9]:((ssItem(X9)&memberP(X7,X9))&lt(X6,X9)))|?[X10]:((ssItem(X10)&memberP(X8,X10))&lt(X10,X6))))))&(nil!=X4|nil!=X3))))))))), inference(assume_negation,[status(cth)],[c_0_0])).
% 0.03/0.29  fof(c_0_7, negated_conjecture, (![X15]:![X19]:![X20]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(((~ssList(X15)|~neq(X15,nil))|~segmentP(esk49_0,X15))|~segmentP(esk48_0,X15)))&(nil!=esk49_0|nil!=esk48_0))&(((nil=esk51_0|ssItem(esk52_0))&(nil=esk50_0|ssItem(esk52_0)))&(((nil=esk51_0|ssList(esk53_0))&(nil=esk50_0|ssList(esk53_0)))&(((nil=esk51_0|ssList(esk54_0))&(nil=esk50_0|ssList(esk54_0)))&(((((nil=esk51_0|cons(esk52_0,nil)=esk50_0)&(nil=esk50_0|cons(esk52_0,nil)=esk50_0))&((nil=esk51_0|app(app(esk53_0,esk50_0),esk54_0)=esk51_0)&(nil=esk50_0|app(app(esk53_0,esk50_0),esk54_0)=esk51_0)))&((nil=esk51_0|((~ssItem(X19)|~memberP(esk53_0,X19))|~lt(esk52_0,X19)))&(nil=esk50_0|((~ssItem(X19)|~memberP(esk53_0,X19))|~lt(esk52_0,X19)))))&((nil=esk51_0|((~ssItem(X20)|~memberP(esk54_0,X20))|~lt(X20,esk52_0)))&(nil=esk50_0|((~ssItem(X20)|~memberP(esk54_0,X20))|~lt(X20,esk52_0)))))))))))))), 
% inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
% 0.03/0.29  fof(c_0_8, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.03/0.29  cnf(c_0_9, negated_conjecture, (nil!=esk48_0|nil!=esk49_0), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.03/0.29  cnf(c_0_10, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.03/0.29  cnf(c_0_11, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.03/0.29  fof(c_0_12, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
% 0.03/0.29  cnf(c_0_13, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.03/0.29  cnf(c_0_14, negated_conjecture, (cons(esk52_0,nil)=esk50_0|nil=esk51_0), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.03/0.29  cnf(c_0_15, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.03/0.29  cnf(c_0_16, negated_conjecture, (ssItem(esk52_0)|nil=esk51_0), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.03/0.29  cnf(c_0_17, negated_conjecture, (esk50_0!=nil|esk51_0!=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_10]), c_0_11])).
% 0.03/0.29  cnf(c_0_18, negated_conjecture, (~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.03/0.29  fof(c_0_19, plain, (![X2]:(~ssList(X2)|segmentP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
% 0.03/0.29  cnf(c_0_20, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.03/0.29  cnf(c_0_21, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.03/0.29  cnf(c_0_22, negated_conjecture, (app(app(esk53_0,esk50_0),esk54_0)=esk51_0|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.03/0.29  cnf(c_0_23, negated_conjecture, (esk50_0!=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])]), c_0_16]), c_0_17])).
% 0.03/0.29  cnf(c_0_24, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.03/0.29  cnf(c_0_25, negated_conjecture, (ssList(esk54_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.03/0.29  cnf(c_0_26, negated_conjecture, (ssList(esk53_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.03/0.29  cnf(c_0_27, negated_conjecture, (~segmentP(esk50_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_10]), c_0_11])).
% 0.03/0.29  cnf(c_0_28, plain, (segmentP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.03/0.29  cnf(c_0_29, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_20, c_0_10])).
% 0.03/0.29  cnf(c_0_30, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_21])).
% 0.03/0.29  cnf(c_0_31, negated_conjecture, (app(app(esk53_0,esk50_0),esk54_0)=esk51_0), inference(sr,[status(thm)],[c_0_22, c_0_23])).
% 0.03/0.29  cnf(c_0_32, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_24, c_0_11])).
% 0.03/0.29  cnf(c_0_33, negated_conjecture, (ssList(esk54_0)), inference(sr,[status(thm)],[c_0_25, c_0_23])).
% 0.03/0.29  cnf(c_0_34, negated_conjecture, (ssList(esk53_0)), inference(sr,[status(thm)],[c_0_26, c_0_23])).
% 0.03/0.29  cnf(c_0_35, negated_conjecture, (~segmentP(esk51_0,esk50_0)|~neq(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])])).
% 0.03/0.29  cnf(c_0_36, negated_conjecture, (segmentP(esk51_0,esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_29])])).
% 0.03/0.29  fof(c_0_37, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
% 0.03/0.29  cnf(c_0_38, negated_conjecture, (~neq(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36])])).
% 0.03/0.29  cnf(c_0_39, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
% 0.03/0.29  cnf(c_0_40, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_15]), c_0_29])]), c_0_23]), ['proof']).
% 0.03/0.29  # SZS output end CNFRefutation
% 0.03/0.29  # Training examples: 0 positive, 0 negative
% 0.03/0.30  
% 0.03/0.30  # -------------------------------------------------
% 0.03/0.30  # User time                : 0.054 s
% 0.03/0.30  # System time              : 0.004 s
% 0.03/0.30  # Total time               : 0.058 s
% 0.03/0.30  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
