<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC021-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC021-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC021-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC021-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC021-1.p">SWC021-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n104.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:19 EDT 2016

% Result     : Unsatisfiable 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of clauses        :  128 (1972 expanded)
%              Number of leaves         :   56 (1279 expanded)
%              Depth                    :   29
%              Number of atoms          :  322 (3388 expanded)
%              Number of equality atoms :   80 ( 868 expanded)
%              Maximal clause size      :    7 (   3 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( ~ ssList(X1)
    | ~ neq(X1,nil)
    | ~ frontsegP(sk2,X1)
    | ~ frontsegP(sk1,X1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_4"></A>cnf(c_0_4,axiom,
    ( frontsegP(X1,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause61)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( frontsegP(sk4,sk3) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( rearsegP(X3,X2)
    | app(X1,X2) != X3
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause143)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( app(app(skaf47(X1,X2),X2),skaf48(X2,X1)) = X1
    | ~ segmentP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause168)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( ssList(skaf48(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause48)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( ssList(skaf47(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause49)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( rearsegP(app(X3,X1),X2)
    | ~ rearsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause136)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( rearsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause58)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( app(skaf46(X1,X2),X2) = X1
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause131)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( ssList(skaf46(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause50)).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( nil != sk2
    | nil != sk1 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_24"></A>cnf(c_0_24,axiom,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause104)).

<A NAME="c_0_25"></A>cnf(c_0_25,axiom,
    ( nil = X2
    | hd(app(X2,X1)) = hd(X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause123)).

<A NAME="c_0_26"></A>cnf(c_0_26,axiom,
    ( nil = X1
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause118)).

<A NAME="c_0_27"></A>cnf(c_0_27,axiom,
    ( X4 = X2
    | cons(X1,X2) != cons(X3,X4)
    | ~ ssItem(X3)
    | ~ ssItem(X1)
    | ~ ssList(X4)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause171)).

<A NAME="c_0_28"></A>cnf(c_0_28,axiom,
    ( cons(skaf83(X1),skaf82(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause109)).

<A NAME="c_0_29"></A>cnf(c_0_29,axiom,
    ( ssItem(skaf83(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause12)).

<A NAME="c_0_30"></A>cnf(c_0_30,axiom,
    ( ssList(skaf82(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause13)).

<A NAME="c_0_31"></A>cnf(c_0_31,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_32"></A>cnf(c_0_32,axiom,
    ( cons(X1,X2) != nil
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause98)).

<A NAME="c_0_33"></A>cnf(c_0_33,axiom,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause76)).

<A NAME="c_0_34"></A>cnf(c_0_34,axiom,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause75)).

<A NAME="c_0_35"></A>cnf(c_0_35,axiom,
    ( memberP(cons(X2,X3),X1)
    | X1 != X2
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause138)).

<A NAME="c_0_36"></A>cnf(c_0_36,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_37"></A>cnf(c_0_37,axiom,
    ( X1 = X3
    | cons(X1,X2) != cons(X3,X4)
    | ~ ssItem(X3)
    | ~ ssItem(X1)
    | ~ ssList(X4)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause170)).

<A NAME="c_0_38"></A>cnf(c_0_38,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_39"></A>cnf(c_0_39,axiom,
    ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
    | ~ memberP(X1,X2)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause169)).

<A NAME="c_0_40"></A>cnf(c_0_40,axiom,
    ( singletonP(X2)
    | cons(X1,nil) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause116)).

<A NAME="c_0_41"></A>cnf(c_0_41,axiom,
    ( frontsegP(cons(X3,X1),cons(X4,X2))
    | ~ frontsegP(X1,X2)
    | X3 != X4
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssItem(X4)
    | ~ ssItem(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause178)).

<A NAME="c_0_42"></A>cnf(c_0_42,axiom,
    ( ssList(skaf43(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause52)).

<A NAME="c_0_43"></A>cnf(c_0_43,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_44"></A>cnf(c_0_44,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_45"></A>cnf(c_0_45,axiom,
    ( cons(skaf44(X1),nil) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause101)).

<A NAME="c_0_46"></A>cnf(c_0_46,axiom,
    ( ssItem(skaf44(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause47)).

<A NAME="c_0_47"></A>cnf(c_0_47,axiom,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause60)).

<A NAME="c_0_48"></A>cnf(c_0_48,axiom,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause157)).

<A NAME="c_0_49"></A>cnf(c_0_49,axiom,
    ( app(X1,X2) = nil
    | nil != X1
    | nil != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause145)).

<A NAME="c_0_50"></A>cnf(c_0_50,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_51"></A>cnf(c_0_51,axiom,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause97)).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( ~ ssList(X1)
    | ~ neq(sk3,X1)
    | ~ frontsegP(sk4,X1)
    | ~ segmentP(X1,sk3)
    | ~ strictorderedP(X1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_53"></A>cnf(c_0_53,axiom,
    ( segmentP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause56)).

<A NAME="c_0_54"></A>cnf(c_0_54,axiom,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause99)).

<A NAME="c_0_55"></A>cnf(c_0_55,axiom,
    ( strictorderedP(cons(X1,nil))
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause66)).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( ~ neq(X1,nil)
    | ~ frontsegP(sk3,X1)
    | ~ frontsegP(sk4,X1)
    | ~ ssList(X1) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,c_0_1]),<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( ~ neq(sk3,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_4]),c_0_5]),c_0_57])])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( sk3 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_6]),c_0_57]),c_0_7])])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( frontsegP(sk4,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,<A HREF="#c_0_59">c_0_59</A>])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( app(nil,skaf45(sk4,nil)) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_60]),c_0_7]),c_0_61])])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( skaf45(sk4,nil) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_62]),c_0_11])])).

<A NAME="c_0_64"></A>cnf(c_0_64,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,negated_conjecture,
    ( app(nil,sk4) = sk4 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,negated_conjecture,
    ( segmentP(app(sk4,X1),sk4)
    | ~ ssList(app(sk4,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,c_0_65]),c_0_61]),c_0_7])])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( segmentP(sk4,sk4) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>,c_0_13]),c_0_61]),c_0_7])])).

<A NAME="c_0_68"></A>cnf(c_0_68,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_14]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_69"></A>cnf(c_0_69,negated_conjecture,
    ( app(app(skaf47(sk4,sk4),sk4),skaf48(sk4,sk4)) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_67]),c_0_61])])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( rearsegP(sk4,skaf48(sk4,sk4))
    | ~ ssList(app(skaf47(sk4,sk4),sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_69]),c_0_17])])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( rearsegP(sk4,skaf48(sk4,sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>,c_0_15]),c_0_18]),c_0_61])])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( rearsegP(app(X1,X2),nil)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_20]),c_0_7])])).

<A NAME="c_0_73"></A>cnf(c_0_73,negated_conjecture,
    ( app(skaf46(sk4,skaf48(sk4,sk4)),skaf48(sk4,sk4)) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_71]),c_0_17]),c_0_61])])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( rearsegP(sk4,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_73]),c_0_22]),c_0_17])])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( app(skaf46(sk4,nil),nil) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_74]),c_0_7]),c_0_61])])).

<A NAME="c_0_76"></A>cnf(c_0_76,negated_conjecture,
    ( skaf46(sk4,nil) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_75]),c_0_22])])).

<A NAME="c_0_77"></A>cnf(c_0_77,negated_conjecture,
    ( sk3 != nil
    | sk4 != nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_1]),<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_78"></A>cnf(c_0_78,plain,
    ( cons(hd(X1),tl(app(X1,X2))) = app(X1,X2)
    | nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_25]),c_0_15]),<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( app(sk4,nil) = sk4 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( sk4 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_59])])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( skaf82(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_28]),c_0_29]),c_0_30])])]),c_0_31]),<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( cons(hd(sk4),tl(sk4)) = sk4 ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_79]),c_0_61]),c_0_7])]),<A HREF="#c_0_80">c_0_80</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( skaf82(sk4) = tl(sk4)
    | ~ ssItem(hd(sk4))
    | ~ ssList(tl(sk4)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,<A HREF="#c_0_82">c_0_82</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( rearsegP(sk4,sk4) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_65]),c_0_61]),c_0_7])])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( skaf82(sk4) = tl(sk4)
    | ~ ssList(tl(sk4)) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_33]),c_0_61])]),<A HREF="#c_0_80">c_0_80</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( rearsegP(app(X1,sk4),sk4)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_84]),c_0_61])])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( skaf82(sk4) = tl(sk4) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_34]),c_0_61])]),<A HREF="#c_0_80">c_0_80</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( rearsegP(app(X1,app(X2,sk4)),sk4)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>,c_0_36]),c_0_61])]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( app(nil,app(sk4,X1)) = app(sk4,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_65]),c_0_7]),c_0_61])])).

<A NAME="c_0_91"></A>cnf(c_0_91,plain,
    ( skaf83(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_28]),c_0_29]),c_0_30])])]),c_0_31]),<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_92"></A>cnf(c_0_92,negated_conjecture,
    ( ssList(tl(sk4)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,<A HREF="#c_0_87">c_0_87</A>])).

<A NAME="c_0_93"></A>cnf(c_0_93,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_94"></A>cnf(c_0_94,plain,
    ( nil = X1
    | memberP(X1,skaf83(X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>,c_0_28]),c_0_29]),c_0_30])])).

<A NAME="c_0_95"></A>cnf(c_0_95,negated_conjecture,
    ( rearsegP(app(sk4,sk4),sk4) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>,c_0_90]),c_0_7]),c_0_61])])).

<A NAME="c_0_96"></A>cnf(c_0_96,negated_conjecture,
    ( skaf83(sk4) = hd(sk4)
    | ~ ssItem(hd(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_82]),c_0_92])])).

<A NAME="c_0_97"></A>cnf(c_0_97,plain,
    ( nil = X1
    | ~ ssList(app(app(X2,cons(skaf83(X1),X3)),X1))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(app(X2,cons(skaf83(X1),X3)),X1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_28]),c_0_29]),c_0_30])])).

<A NAME="c_0_98"></A>cnf(c_0_98,plain,
    ( app(skaf42(X1,skaf83(X1)),cons(skaf83(X1),skaf43(skaf83(X1),X1))) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_94]),c_0_29])])).

<A NAME="c_0_99"></A>cnf(c_0_99,negated_conjecture,
    ( app(skaf46(app(sk4,sk4),sk4),sk4) = app(sk4,sk4)
    | ~ ssList(app(sk4,sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_95]),c_0_61])])).

<A NAME="c_0_100"></A>cnf(c_0_100,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1)
    | ~ ssList(cons(X1,nil)) ),
    inference(er,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_101"></A>cnf(c_0_101,plain,
    ( frontsegP(cons(X1,X2),cons(X1,X3))
    | ~ frontsegP(X2,X3)
    | ~ ssItem(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_102"></A>cnf(c_0_102,negated_conjecture,
    ( skaf83(sk4) = hd(sk4) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,c_0_33]),c_0_61])]),<A HREF="#c_0_80">c_0_80</A>])).

<A NAME="c_0_103"></A>cnf(c_0_103,plain,
    ( nil = X1
    | ~ ssList(app(X1,X1))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(X1,X1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_98]),c_0_42]),c_0_43])])).

<A NAME="c_0_104"></A>cnf(c_0_104,negated_conjecture,
    ( app(skaf46(app(sk4,sk4),sk4),sk4) = app(sk4,sk4) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_15]),c_0_61])])).

<A NAME="c_0_105"></A>cnf(c_0_105,plain,
    ( tl(X1) = nil
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_45]),c_0_46]),c_0_7])])).

<A NAME="c_0_106"></A>cnf(c_0_106,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>,c_0_31]),c_0_7])])).

<A NAME="c_0_107"></A>cnf(c_0_107,plain,
    ( frontsegP(cons(X1,X2),cons(X1,nil))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_101">c_0_101</A>,c_0_47]),c_0_7])])).

<A NAME="c_0_108"></A>cnf(c_0_108,negated_conjecture,
    ( ssItem(hd(sk4)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,<A HREF="#c_0_102">c_0_102</A>])).

<A NAME="c_0_109"></A>cnf(c_0_109,plain,
    ( ssList(cons(X1,app(X2,X3)))
    | ~ ssItem(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_48]),<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_110"></A>cnf(c_0_110,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_49])]),c_0_7])])).

<A NAME="c_0_111"></A>cnf(c_0_111,plain,
    ( nil = X1
    | ssItem(X2)
    | ~ ssList(app(X1,X1))
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_103">c_0_103</A>,<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_112"></A>cnf(c_0_112,negated_conjecture,
    ( ssList(app(sk4,sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_104]),c_0_22]),c_0_61])])).

<A NAME="c_0_113"></A>cnf(c_0_113,plain,
    ( cons(X1,tl(cons(X1,X2))) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_51]),c_0_31]),<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_114"></A>cnf(c_0_114,plain,
    ( tl(cons(X1,nil)) = nil
    | ~ ssItem(X1)
    | ~ ssList(cons(X1,nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_105">c_0_105</A>,<A HREF="#c_0_106">c_0_106</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,negated_conjecture,
    ( ~ neq(nil,X1)
    | ~ frontsegP(sk4,X1)
    | ~ ssList(X1)
    | ~ strictorderedP(X1) ),
    inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_59]),c_0_59]),<A HREF="#c_0_53">c_0_53</A>])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( frontsegP(sk4,cons(hd(sk4),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>,c_0_82]),c_0_108]),c_0_92])])).

<A NAME="c_0_117"></A>cnf(c_0_117,plain,
    ( ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>,c_0_110]),c_0_7])])).

<A NAME="c_0_118"></A>cnf(c_0_118,negated_conjecture,
    ( ssItem(X1) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_111">c_0_111</A>,c_0_112]),c_0_61])]),<A HREF="#c_0_80">c_0_80</A>])).

<A NAME="c_0_119"></A>cnf(c_0_119,plain,
    ( tl(cons(X1,X2)) != cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(tl(cons(X1,X2)))
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,<A HREF="#c_0_113">c_0_113</A>])).

<A NAME="c_0_120"></A>cnf(c_0_120,plain,
    ( tl(cons(X1,nil)) = nil
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_114">c_0_114</A>,c_0_31]),c_0_7])])).

<A NAME="c_0_121"></A>cnf(c_0_121,negated_conjecture,
    ( ~ neq(nil,cons(hd(sk4),nil))
    | ~ ssList(cons(hd(sk4),nil))
    | ~ strictorderedP(cons(hd(sk4),nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_115">c_0_115</A>,<A HREF="#c_0_116">c_0_116</A>])).

<A NAME="c_0_122"></A>cnf(c_0_122,plain,
    ( ssList(cons(X1,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_117">c_0_117</A>,c_0_118])])).

<A NAME="c_0_123"></A>cnf(c_0_123,plain,
    ( strictorderedP(cons(X1,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_118])])).

<A NAME="c_0_124"></A>cnf(c_0_124,plain,
    ( cons(X1,nil) != nil
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>,c_0_120]),c_0_7])])).

<A NAME="c_0_125"></A>cnf(c_0_125,negated_conjecture,
    ( ~ neq(nil,cons(hd(sk4),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_121">c_0_121</A>,c_0_122])]),c_0_123])])).

<A NAME="c_0_126"></A>cnf(c_0_126,plain,
    ( cons(X1,nil) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_124">c_0_124</A>,c_0_118])])).

<A NAME="c_0_127"></A>cnf(c_0_127,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_125">c_0_125</A>,c_0_6]),c_0_7]),c_0_122])]),<A HREF="#c_0_126">c_0_126</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC021-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n104.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:34:23 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.60  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.06/0.60  # and selection function SelectCQArNTEqFirst.
% 0.06/0.60  #
% 0.06/0.60  # Preprocessing time       : 0.009 s
% 0.06/0.60  # Presaturation interreduction done
% 0.06/0.60  
% 0.06/0.60  # Proof found!
% 0.06/0.60  # SZS status Unsatisfiable
% 0.06/0.60  # SZS output start CNFRefutation
% 0.06/0.60  cnf(c_0_0, negated_conjecture, (~ssList(X1)|~neq(X1,nil)|~frontsegP(sk2,X1)|~frontsegP(sk1,X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
% 0.06/0.60  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 0.06/0.60  cnf(c_0_2, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 0.06/0.60  cnf(c_0_3, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 0.06/0.60  cnf(c_0_4, axiom, (frontsegP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause61)).
% 0.06/0.60  cnf(c_0_5, negated_conjecture, (frontsegP(sk4,sk3)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 0.06/0.60  cnf(c_0_6, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause100)).
% 0.06/0.60  cnf(c_0_7, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.06/0.60  cnf(c_0_8, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
% 0.06/0.60  cnf(c_0_9, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause132)).
% 0.06/0.60  cnf(c_0_10, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause74)).
% 0.06/0.60  cnf(c_0_11, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause51)).
% 0.06/0.60  cnf(c_0_12, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause173)).
% 0.06/0.60  cnf(c_0_13, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause73)).
% 0.06/0.60  cnf(c_0_14, axiom, (rearsegP(X3,X2)|app(X1,X2)!=X3|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause143)).
% 0.06/0.60  cnf(c_0_15, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.06/0.60  cnf(c_0_16, axiom, (app(app(skaf47(X1,X2),X2),skaf48(X2,X1))=X1|~segmentP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause168)).
% 0.06/0.60  cnf(c_0_17, axiom, (ssList(skaf48(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause48)).
% 0.06/0.60  cnf(c_0_18, axiom, (ssList(skaf47(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause49)).
% 0.06/0.60  cnf(c_0_19, axiom, (rearsegP(app(X3,X1),X2)|~rearsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause136)).
% 0.06/0.60  cnf(c_0_20, axiom, (rearsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause58)).
% 0.06/0.60  cnf(c_0_21, axiom, (app(skaf46(X1,X2),X2)=X1|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause131)).
% 0.06/0.60  cnf(c_0_22, axiom, (ssList(skaf46(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause50)).
% 0.06/0.60  cnf(c_0_23, negated_conjecture, (nil!=sk2|nil!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_11)).
% 0.06/0.60  cnf(c_0_24, axiom, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause104)).
% 0.06/0.60  cnf(c_0_25, axiom, (nil=X2|hd(app(X2,X1))=hd(X2)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause123)).
% 0.06/0.60  cnf(c_0_26, axiom, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause118)).
% 0.06/0.60  cnf(c_0_27, axiom, (X4=X2|cons(X1,X2)!=cons(X3,X4)|~ssItem(X3)|~ssItem(X1)|~ssList(X4)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause171)).
% 0.06/0.60  cnf(c_0_28, axiom, (cons(skaf83(X1),skaf82(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause109)).
% 0.06/0.60  cnf(c_0_29, axiom, (ssItem(skaf83(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause12)).
% 0.06/0.60  cnf(c_0_30, axiom, (ssList(skaf82(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause13)).
% 0.06/0.60  cnf(c_0_31, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.06/0.60  cnf(c_0_32, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause98)).
% 0.06/0.60  cnf(c_0_33, axiom, (ssItem(hd(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause76)).
% 0.06/0.60  cnf(c_0_34, axiom, (ssList(tl(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause75)).
% 0.06/0.60  cnf(c_0_35, axiom, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause138)).
% 0.06/0.60  cnf(c_0_36, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause149)).
% 0.06/0.60  cnf(c_0_37, axiom, (X1=X3|cons(X1,X2)!=cons(X3,X4)|~ssItem(X3)|~ssItem(X1)|~ssList(X4)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause170)).
% 0.06/0.60  cnf(c_0_38, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause179)).
% 0.06/0.60  cnf(c_0_39, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause169)).
% 0.06/0.60  cnf(c_0_40, axiom, (singletonP(X2)|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause116)).
% 0.06/0.60  cnf(c_0_41, axiom, (frontsegP(cons(X3,X1),cons(X4,X2))|~frontsegP(X1,X2)|X3!=X4|~ssList(X2)|~ssList(X1)|~ssItem(X4)|~ssItem(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause178)).
% 0.06/0.60  cnf(c_0_42, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause52)).
% 0.06/0.60  cnf(c_0_43, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause53)).
% 0.06/0.60  cnf(c_0_44, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause96)).
% 0.06/0.60  cnf(c_0_45, axiom, (cons(skaf44(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause101)).
% 0.06/0.60  cnf(c_0_46, axiom, (ssItem(skaf44(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause47)).
% 0.06/0.60  cnf(c_0_47, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause60)).
% 0.06/0.60  cnf(c_0_48, axiom, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause157)).
% 0.06/0.60  cnf(c_0_49, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause145)).
% 0.06/0.60  cnf(c_0_50, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause72)).
% 0.06/0.60  cnf(c_0_51, axiom, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause97)).
% 0.06/0.60  cnf(c_0_52, negated_conjecture, (~ssList(X1)|~neq(sk3,X1)|~frontsegP(sk4,X1)|~segmentP(X1,sk3)|~strictorderedP(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
% 0.06/0.60  cnf(c_0_53, axiom, (segmentP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause56)).
% 0.06/0.60  cnf(c_0_54, axiom, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause99)).
% 0.06/0.60  cnf(c_0_55, axiom, (strictorderedP(cons(X1,nil))|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause66)).
% 0.06/0.60  cnf(c_0_56, negated_conjecture, (~neq(X1,nil)|~frontsegP(sk3,X1)|~frontsegP(sk4,X1)|~ssList(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
% 0.06/0.60  cnf(c_0_57, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_3, c_0_1])).
% 0.06/0.60  cnf(c_0_58, negated_conjecture, (~neq(sk3,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_4]), c_0_5]), c_0_57])])).
% 0.06/0.60  cnf(c_0_59, negated_conjecture, (sk3=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_6]), c_0_57]), c_0_7])])).
% 0.06/0.60  cnf(c_0_60, negated_conjecture, (frontsegP(sk4,nil)), inference(rw,[status(thm)],[c_0_5, c_0_59])).
% 0.06/0.60  cnf(c_0_61, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_8, c_0_2])).
% 0.06/0.60  cnf(c_0_62, negated_conjecture, (app(nil,skaf45(sk4,nil))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_60]), c_0_7]), c_0_61])])).
% 0.06/0.60  cnf(c_0_63, negated_conjecture, (skaf45(sk4,nil)=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_62]), c_0_11])])).
% 0.06/0.60  cnf(c_0_64, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(er,[status(thm)],[c_0_12])).
% 0.06/0.60  cnf(c_0_65, negated_conjecture, (app(nil,sk4)=sk4), inference(rw,[status(thm)],[c_0_62, c_0_63])).
% 0.06/0.60  cnf(c_0_66, negated_conjecture, (segmentP(app(sk4,X1),sk4)|~ssList(app(sk4,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_61]), c_0_7])])).
% 0.06/0.60  cnf(c_0_67, negated_conjecture, (segmentP(sk4,sk4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_13]), c_0_61]), c_0_7])])).
% 0.06/0.60  cnf(c_0_68, plain, (rearsegP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_14]), c_0_15])).
% 0.06/0.60  cnf(c_0_69, negated_conjecture, (app(app(skaf47(sk4,sk4),sk4),skaf48(sk4,sk4))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_67]), c_0_61])])).
% 0.06/0.60  cnf(c_0_70, negated_conjecture, (rearsegP(sk4,skaf48(sk4,sk4))|~ssList(app(skaf47(sk4,sk4),sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_17])])).
% 0.06/0.60  cnf(c_0_71, negated_conjecture, (rearsegP(sk4,skaf48(sk4,sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_15]), c_0_18]), c_0_61])])).
% 0.06/0.60  cnf(c_0_72, plain, (rearsegP(app(X1,X2),nil)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_7])])).
% 0.06/0.60  cnf(c_0_73, negated_conjecture, (app(skaf46(sk4,skaf48(sk4,sk4)),skaf48(sk4,sk4))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_71]), c_0_17]), c_0_61])])).
% 0.06/0.60  cnf(c_0_74, negated_conjecture, (rearsegP(sk4,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_22]), c_0_17])])).
% 0.06/0.60  cnf(c_0_75, negated_conjecture, (app(skaf46(sk4,nil),nil)=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_74]), c_0_7]), c_0_61])])).
% 0.06/0.60  cnf(c_0_76, negated_conjecture, (skaf46(sk4,nil)=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_75]), c_0_22])])).
% 0.06/0.60  cnf(c_0_77, negated_conjecture, (sk3!=nil|sk4!=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_1]), c_0_2])).
% 0.06/0.60  cnf(c_0_78, plain, (cons(hd(X1),tl(app(X1,X2)))=app(X1,X2)|nil=X1|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_15]), c_0_26])).
% 0.06/0.60  cnf(c_0_79, negated_conjecture, (app(sk4,nil)=sk4), inference(rw,[status(thm)],[c_0_75, c_0_76])).
% 0.06/0.60  cnf(c_0_80, negated_conjecture, (sk4!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_59])])).
% 0.06/0.60  cnf(c_0_81, plain, (skaf82(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])])]), c_0_31]), c_0_32])).
% 0.06/0.60  cnf(c_0_82, negated_conjecture, (cons(hd(sk4),tl(sk4))=sk4), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_61]), c_0_7])]), c_0_80])).
% 0.06/0.60  cnf(c_0_83, negated_conjecture, (skaf82(sk4)=tl(sk4)|~ssItem(hd(sk4))|~ssList(tl(sk4))), inference(spm,[status(thm)],[c_0_81, c_0_82])).
% 0.06/0.60  cnf(c_0_84, negated_conjecture, (rearsegP(sk4,sk4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_65]), c_0_61]), c_0_7])])).
% 0.06/0.60  cnf(c_0_85, negated_conjecture, (skaf82(sk4)=tl(sk4)|~ssList(tl(sk4))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_33]), c_0_61])]), c_0_80])).
% 0.06/0.60  cnf(c_0_86, negated_conjecture, (rearsegP(app(X1,sk4),sk4)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_84]), c_0_61])])).
% 0.06/0.60  cnf(c_0_87, negated_conjecture, (skaf82(sk4)=tl(sk4)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_34]), c_0_61])]), c_0_80])).
% 0.06/0.60  cnf(c_0_88, plain, (memberP(cons(X1,X2),X1)|~ssItem(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_35])).
% 0.06/0.60  cnf(c_0_89, negated_conjecture, (rearsegP(app(X1,app(X2,sk4)),sk4)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_36]), c_0_61])]), c_0_15])).
% 0.06/0.60  cnf(c_0_90, negated_conjecture, (app(nil,app(sk4,X1))=app(sk4,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_65]), c_0_7]), c_0_61])])).
% 0.06/0.60  cnf(c_0_91, plain, (skaf83(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_28]), c_0_29]), c_0_30])])]), c_0_31]), c_0_32])).
% 0.06/0.60  cnf(c_0_92, negated_conjecture, (ssList(tl(sk4))), inference(spm,[status(thm)],[c_0_30, c_0_87])).
% 0.06/0.60  cnf(c_0_93, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_38])).
% 0.06/0.60  cnf(c_0_94, plain, (nil=X1|memberP(X1,skaf83(X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_28]), c_0_29]), c_0_30])])).
% 0.06/0.60  cnf(c_0_95, negated_conjecture, (rearsegP(app(sk4,sk4),sk4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_7]), c_0_61])])).
% 0.06/0.60  cnf(c_0_96, negated_conjecture, (skaf83(sk4)=hd(sk4)|~ssItem(hd(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_82]), c_0_92])])).
% 0.06/0.60  cnf(c_0_97, plain, (nil=X1|~ssList(app(app(X2,cons(skaf83(X1),X3)),X1))|~ssList(X3)|~ssList(X2)|~ssList(X1)|~duplicatefreeP(app(app(X2,cons(skaf83(X1),X3)),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_28]), c_0_29]), c_0_30])])).
% 0.06/0.60  cnf(c_0_98, plain, (app(skaf42(X1,skaf83(X1)),cons(skaf83(X1),skaf43(skaf83(X1),X1)))=X1|nil=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_94]), c_0_29])])).
% 0.06/0.60  cnf(c_0_99, negated_conjecture, (app(skaf46(app(sk4,sk4),sk4),sk4)=app(sk4,sk4)|~ssList(app(sk4,sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_95]), c_0_61])])).
% 0.06/0.60  cnf(c_0_100, plain, (singletonP(cons(X1,nil))|~ssItem(X1)|~ssList(cons(X1,nil))), inference(er,[status(thm)],[c_0_40])).
% 0.06/0.60  cnf(c_0_101, plain, (frontsegP(cons(X1,X2),cons(X1,X3))|~frontsegP(X2,X3)|~ssItem(X1)|~ssList(X3)|~ssList(X2)), inference(er,[status(thm)],[c_0_41])).
% 0.06/0.60  cnf(c_0_102, negated_conjecture, (skaf83(sk4)=hd(sk4)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_33]), c_0_61])]), c_0_80])).
% 0.06/0.60  cnf(c_0_103, plain, (nil=X1|~ssList(app(X1,X1))|~ssList(X1)|~duplicatefreeP(app(X1,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_42]), c_0_43])])).
% 0.06/0.60  cnf(c_0_104, negated_conjecture, (app(skaf46(app(sk4,sk4),sk4),sk4)=app(sk4,sk4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_15]), c_0_61])])).
% 0.06/0.60  cnf(c_0_105, plain, (tl(X1)=nil|~singletonP(X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46]), c_0_7])])).
% 0.06/0.60  cnf(c_0_106, plain, (singletonP(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_31]), c_0_7])])).
% 0.06/0.60  cnf(c_0_107, plain, (frontsegP(cons(X1,X2),cons(X1,nil))|~ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_47]), c_0_7])])).
% 0.06/0.60  cnf(c_0_108, negated_conjecture, (ssItem(hd(sk4))), inference(spm,[status(thm)],[c_0_29, c_0_102])).
% 0.06/0.60  cnf(c_0_109, plain, (ssList(cons(X1,app(X2,X3)))|~ssItem(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_48]), c_0_31])).
% 0.06/0.60  cnf(c_0_110, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_49])]), c_0_7])])).
% 0.06/0.60  cnf(c_0_111, plain, (nil=X1|ssItem(X2)|~ssList(app(X1,X1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_103, c_0_50])).
% 0.06/0.60  cnf(c_0_112, negated_conjecture, (ssList(app(sk4,sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_104]), c_0_22]), c_0_61])])).
% 0.06/0.60  cnf(c_0_113, plain, (cons(X1,tl(cons(X1,X2)))=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_51]), c_0_31]), c_0_32])).
% 0.06/0.60  cnf(c_0_114, plain, (tl(cons(X1,nil))=nil|~ssItem(X1)|~ssList(cons(X1,nil))), inference(spm,[status(thm)],[c_0_105, c_0_106])).
% 0.06/0.60  cnf(c_0_115, negated_conjecture, (~neq(nil,X1)|~frontsegP(sk4,X1)|~ssList(X1)|~strictorderedP(X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_59]), c_0_59]), c_0_53])).
% 0.06/0.60  cnf(c_0_116, negated_conjecture, (frontsegP(sk4,cons(hd(sk4),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_82]), c_0_108]), c_0_92])])).
% 0.06/0.60  cnf(c_0_117, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_7])])).
% 0.06/0.60  cnf(c_0_118, negated_conjecture, (ssItem(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_61])]), c_0_80])).
% 0.06/0.60  cnf(c_0_119, plain, (tl(cons(X1,X2))!=cons(X1,X2)|~ssItem(X1)|~ssList(tl(cons(X1,X2)))|~ssList(X2)), inference(spm,[status(thm)],[c_0_54, c_0_113])).
% 0.06/0.60  cnf(c_0_120, plain, (tl(cons(X1,nil))=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_31]), c_0_7])])).
% 0.06/0.60  cnf(c_0_121, negated_conjecture, (~neq(nil,cons(hd(sk4),nil))|~ssList(cons(hd(sk4),nil))|~strictorderedP(cons(hd(sk4),nil))), inference(spm,[status(thm)],[c_0_115, c_0_116])).
% 0.06/0.60  cnf(c_0_122, plain, (ssList(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_117, c_0_118])])).
% 0.06/0.60  cnf(c_0_123, plain, (strictorderedP(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_118])])).
% 0.06/0.60  cnf(c_0_124, plain, (cons(X1,nil)!=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_7])])).
% 0.06/0.60  cnf(c_0_125, negated_conjecture, (~neq(nil,cons(hd(sk4),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_121, c_0_122])]), c_0_123])])).
% 0.06/0.60  cnf(c_0_126, plain, (cons(X1,nil)!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_118])])).
% 0.06/0.60  cnf(c_0_127, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_6]), c_0_7]), c_0_122])]), c_0_126]), ['proof']).
% 0.06/0.60  # SZS output end CNFRefutation
% 0.06/0.60  # Training examples: 0 positive, 0 negative
% 0.06/0.61  
% 0.06/0.61  # -------------------------------------------------
% 0.06/0.61  # User time                : 0.366 s
% 0.06/0.61  # System time              : 0.011 s
% 0.06/0.61  # Total time               : 0.377 s
% 0.06/0.61  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
