<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC179+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC179+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC179+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC179%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC179+1.p">SWC179+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n180.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:57 EDT 2016

% Result     : Theorem 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of formulae       :   15 (  33 expanded)
%              Number of clauses        :    9 (  12 expanded)
%              Number of leaves         :    3 (   8 expanded)
%              Depth                    :    6
%              Number of atoms          :  100 ( 518 expanded)
%              Number of equality atoms :   37 ( 198 expanded)
%              Maximal formula depth    :   27 (   7 average)
%              Maximal clause size      :   43 (   1 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | duplicatefreeP(X1)
                    | ( ! [X5] :
                          ( ssItem(X5)
                         => ! [X6] :
                              ( ssList(X6)
                             => ! [X7] :
                                  ( ssList(X7)
                                 => ( cons(X5,nil) != X3
                                    | app(app(X6,X3),X7) != X4
                                    | ? [X8] :
                                        ( ssItem(X8)
                                        & memberP(X6,X8)
                                        & lt(X5,X8) )
                                    | ? [X9] :
                                        ( ssItem(X9)
                                        & memberP(X7,X9)
                                        & lt(X9,X5) ) ) ) ) )
                      & ( nil != X4
                        | nil != X3 ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => duplicatefreeP(cons(X1,nil)) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax71)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    duplicatefreeP(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax72)).

<A NAME="c_0_3"></A>fof(c_0_3,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | duplicatefreeP(X1)
                      | ( ! [X5] :
                            ( ssItem(X5)
                           => ! [X6] :
                                ( ssList(X6)
                               => ! [X7] :
                                    ( ssList(X7)
                                   => ( cons(X5,nil) != X3
                                      | app(app(X6,X3),X7) != X4
                                      | ? [X8] :
                                          ( ssItem(X8)
                                          & memberP(X6,X8)
                                          & lt(X5,X8) )
                                      | ? [X9] :
                                          ( ssItem(X9)
                                          & memberP(X7,X9)
                                          & lt(X9,X5) ) ) ) ) )
                        & ( nil != X4
                          | nil != X3 ) ) ) ) ) ) ) ),
    inference(assume_negation,[status(cth)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_4"></A>fof(c_0_4,negated_conjecture,(
    ! [X17,X18] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ~ duplicatefreeP(esk48_0)
      & ( nil = esk51_0
        | ssItem(esk52_0) )
      & ( nil = esk50_0
        | ssItem(esk52_0) )
      & ( nil = esk51_0
        | ssList(esk53_0) )
      & ( nil = esk50_0
        | ssList(esk53_0) )
      & ( nil = esk51_0
        | ssList(esk54_0) )
      & ( nil = esk50_0
        | ssList(esk54_0) )
      & ( nil = esk51_0
        | cons(esk52_0,nil) = esk50_0 )
      & ( nil = esk50_0
        | cons(esk52_0,nil) = esk50_0 )
      & ( nil = esk51_0
        | app(app(esk53_0,esk50_0),esk54_0) = esk51_0 )
      & ( nil = esk50_0
        | app(app(esk53_0,esk50_0),esk54_0) = esk51_0 )
      & ( nil = esk51_0
        | ~ ssItem(X17)
        | ~ memberP(esk53_0,X17)
        | ~ lt(esk52_0,X17) )
      & ( nil = esk50_0
        | ~ ssItem(X17)
        | ~ memberP(esk53_0,X17)
        | ~ lt(esk52_0,X17) )
      & ( nil = esk51_0
        | ~ ssItem(X18)
        | ~ memberP(esk54_0,X18)
        | ~ lt(X18,esk52_0) )
      & ( nil = esk50_0
        | ~ ssItem(X18)
        | ~ memberP(esk54_0,X18)
        | ~ lt(X18,esk52_0) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).

<A NAME="c_0_5"></A>fof(c_0_5,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | duplicatefreeP(cons(X2,nil)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ~ duplicatefreeP(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_8"></A>cnf(c_0_8,plain,
    ( duplicatefreeP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( cons(esk52_0,nil) = esk50_0
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( ~ duplicatefreeP(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>,<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ssItem(esk52_0)
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_9]),c_0_10]),<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,plain,
    ( duplicatefreeP(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_12]),c_0_13])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC179+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.22  % Computer   : n180.star.cs.uiowa.edu
% 0.03/0.22  % Model      : x86_64 x86_64
% 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.22  % Memory     : 32218.75MB
% 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.22  % CPULimit   : 300
% 0.03/0.22  % DateTime   : Tue May 10 20:36:52 CDT 2016
% 0.03/0.22  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.03/0.25  # and selection function SelectCQPrecWNTNp.
% 0.03/0.25  #
% 0.03/0.25  # Preprocessing time       : 0.014 s
% 0.03/0.25  # Presaturation interreduction done
% 0.03/0.25  
% 0.03/0.25  # Proof found!
% 0.03/0.25  # SZS status Theorem
% 0.03/0.25  # SZS output start CNFRefutation
% 0.03/0.25  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|duplicatefreeP(X1))|(![X5]:(ssItem(X5)=>![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>(((cons(X5,nil)!=X3|app(app(X6,X3),X7)!=X4)|?[X8]:((ssItem(X8)&memberP(X6,X8))&lt(X5,X8)))|?[X9]:((ssItem(X9)&memberP(X7,X9))&lt(X9,X5))))))&(nil!=X4|nil!=X3)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.03/0.25  fof(c_0_1, axiom, (![X1]:(ssItem(X1)=>duplicatefreeP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax71)).
% 0.03/0.25  fof(c_0_2, axiom, (duplicatefreeP(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax72)).
% 0.03/0.25  fof(c_0_3, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|duplicatefreeP(X1))|(![X5]:(ssItem(X5)=>![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>(((cons(X5,nil)!=X3|app(app(X6,X3),X7)!=X4)|?[X8]:((ssItem(X8)&memberP(X6,X8))&lt(X5,X8)))|?[X9]:((ssItem(X9)&memberP(X7,X9))&lt(X9,X5))))))&(nil!=X4|nil!=X3))))))))), inference(assume_negation,[status(cth)],[c_0_0])).
% 0.03/0.25  fof(c_0_4, negated_conjecture, (![X17]:![X18]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((esk49_0=esk51_0&esk48_0=esk50_0)&~duplicatefreeP(esk48_0))&(((nil=esk51_0|ssItem(esk52_0))&(nil=esk50_0|ssItem(esk52_0)))&(((nil=esk51_0|ssList(esk53_0))&(nil=esk50_0|ssList(esk53_0)))&(((nil=esk51_0|ssList(esk54_0))&(nil=esk50_0|ssList(esk54_0)))&(((((nil=esk51_0|cons(esk52_0,nil)=esk50_0)&(nil=esk50_0|cons(esk52_0,nil)=esk50_0))&((nil=esk51_0|app(app(esk53_0,esk50_0),esk54_0)=esk51_0)&(nil=esk50_0|app(app(esk53_0,esk50_0),esk54_0)=esk51_0)))&((nil=esk51_0|((~ssItem(X17)|~memberP(esk53_0,X17))|~lt(esk52_0,X17)))&(nil=esk50_0|((~ssItem(X17)|~memberP(esk53_0,X17))|~lt(esk52_0,X17)))))&((nil=esk51_0|((~ssItem(X18)|~memberP(esk54_0,X18))|~lt(X18,esk52_0)))&(nil=esk50_0|((~ssItem(X18)|~memberP(esk54_0,X18))|~lt(X18,esk52_0)))))))))))))), 
% inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
% 0.03/0.25  fof(c_0_5, plain, (![X2]:(~ssItem(X2)|duplicatefreeP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
% 0.03/0.25  cnf(c_0_6, negated_conjecture, (~duplicatefreeP(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.03/0.25  cnf(c_0_7, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.03/0.25  cnf(c_0_8, plain, (duplicatefreeP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.03/0.25  cnf(c_0_9, negated_conjecture, (cons(esk52_0,nil)=esk50_0|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.03/0.25  cnf(c_0_10, negated_conjecture, (~duplicatefreeP(esk50_0)), inference(rw,[status(thm)],[c_0_6, c_0_7])).
% 0.03/0.25  cnf(c_0_11, negated_conjecture, (ssItem(esk52_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.03/0.25  cnf(c_0_12, negated_conjecture, (esk50_0=nil), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10]), c_0_11])).
% 0.03/0.25  cnf(c_0_13, plain, (duplicatefreeP(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.03/0.25  cnf(c_0_14, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_12]), c_0_13])]), ['proof']).
% 0.03/0.25  # SZS output end CNFRefutation
% 0.03/0.25  # Training examples: 0 positive, 0 negative
% 0.03/0.25  
% 0.03/0.25  # -------------------------------------------------
% 0.03/0.25  # User time                : 0.022 s
% 0.03/0.25  # System time              : 0.006 s
% 0.03/0.25  # Total time               : 0.028 s
% 0.03/0.25  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
