<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC229+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC229+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC229+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC229%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC229+1.p">SWC229+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n166.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:08 EDT 2016

% Result     : Theorem 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of formulae       :   51 ( 177 expanded)
%              Number of clauses        :   32 (  70 expanded)
%              Number of leaves         :    9 (  45 expanded)
%              Depth                    :   12
%              Number of atoms          :  197 (1947 expanded)
%              Number of equality atoms :   45 ( 429 expanded)
%              Maximal formula depth    :   28 (   5 average)
%              Maximal clause size      :   35 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax38)).

<A NAME="c_0_1"></A>fof(c_0_1,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ~ segmentP(X4,X3)
                    | nil = X1
                    | ? [X5] :
                        ( ssItem(X5)
                        & ? [X6] :
                            ( ssList(X6)
                            & ? [X7] :
                                ( ssList(X7)
                                & app(app(X6,cons(X5,nil)),X7) = X1
                                & ! [X8] :
                                    ( ssItem(X8)
                                   => ( ~ memberP(X6,X8)
                                      | ~ memberP(X7,X8)
                                      | ~ leq(X5,X8)
                                      | leq(X8,X5) ) ) ) ) )
                    | ( ~ singletonP(X3)
                      & neq(X4,nil) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( singletonP(X1)
      &lt;=> ? [X2] :
            ( ssItem(X2)
            & cons(X2,nil) = X1 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax4)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( ( segmentP(X1,X2)
              & segmentP(X2,X1) )
           => X1 = X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax54)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,nil) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax57)).

<A NAME="c_0_9"></A>fof(c_0_9,plain,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_10"></A>fof(c_0_10,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ~ segmentP(X4,X3)
                      | nil = X1
                      | ? [X5] :
                          ( ssItem(X5)
                          & ? [X6] :
                              ( ssList(X6)
                              & ? [X7] :
                                  ( ssList(X7)
                                  & app(app(X6,cons(X5,nil)),X7) = X1
                                  & ! [X8] :
                                      ( ssItem(X8)
                                     => ( ~ memberP(X6,X8)
                                        | ~ memberP(X7,X8)
                                        | ~ leq(X5,X8)
                                        | leq(X8,X5) ) ) ) ) )
                      | ( ~ singletonP(X3)
                        & neq(X4,nil) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_1])])).

<A NAME="c_0_11"></A>fof(c_0_11,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | ~ memberP(nil,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).

<A NAME="c_0_12"></A>fof(c_0_12,negated_conjecture,(
    ! [X13,X14,X15] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & segmentP(esk51_0,esk50_0)
      & nil != esk48_0
      & ( ssItem(esk52_3(X13,X14,X15))
        | ~ ssList(X15)
        | app(app(X14,cons(X13,nil)),X15) != esk48_0
        | ~ ssList(X14)
        | ~ ssItem(X13) )
      & ( memberP(X14,esk52_3(X13,X14,X15))
        | ~ ssList(X15)
        | app(app(X14,cons(X13,nil)),X15) != esk48_0
        | ~ ssList(X14)
        | ~ ssItem(X13) )
      & ( memberP(X15,esk52_3(X13,X14,X15))
        | ~ ssList(X15)
        | app(app(X14,cons(X13,nil)),X15) != esk48_0
        | ~ ssList(X14)
        | ~ ssItem(X13) )
      & ( leq(X13,esk52_3(X13,X14,X15))
        | ~ ssList(X15)
        | app(app(X14,cons(X13,nil)),X15) != esk48_0
        | ~ ssList(X14)
        | ~ ssItem(X13) )
      & ( ~ leq(esk52_3(X13,X14,X15),X13)
        | ~ ssList(X15)
        | app(app(X14,cons(X13,nil)),X15) != esk48_0
        | ~ ssList(X14)
        | ~ ssItem(X13) )
      & ( singletonP(esk50_0)
        | ~ neq(esk51_0,nil) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).

<A NAME="c_0_13"></A>cnf(c_0_13,plain,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( memberP(X3,esk52_3(X1,X2,X3))
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | app(app(X2,cons(X1,nil)),X3) != esk48_0
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( app(app(X1,cons(X2,nil)),nil) != esk48_0
    | ~ ssList(X1)
    | ~ ssItem(esk52_3(X2,X1,nil))
    | ~ ssItem(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_14]),c_0_15])])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( ssItem(esk52_3(X1,X2,X3))
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | app(app(X2,cons(X1,nil)),X3) != esk48_0
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_18"></A>fof(c_0_18,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( app(app(X1,cons(X2,nil)),nil) != esk48_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_17]),c_0_15])])).

<A NAME="c_0_20"></A>cnf(c_0_20,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_21"></A>fof(c_0_21,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( app(cons(X1,nil),nil) != esk48_0
    | ~ ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_20]),c_0_15])])).

<A NAME="c_0_23"></A>cnf(c_0_23,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_24"></A>fof(c_0_24,plain,(
    ! [X3,X5] :
      ( ( ssItem(esk5_1(X3))
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( cons(esk5_1(X3),nil) = X3
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( ~ ssItem(X5)
        | cons(X5,nil) != X3
        | singletonP(X3)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( singletonP(esk50_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_27"></A>fof(c_0_27,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( cons(X1,nil) != esk48_0
    | ~ ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( cons(esk5_1(X1),nil) = X1
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_33"></A>fof(c_0_33,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ~ segmentP(X3,X4)
      | ~ segmentP(X4,X3)
      | X3 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( segmentP(esk51_0,esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( singletonP(esk48_0)
    | ~ neq(esk51_0,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( ~ singletonP(esk48_0)
    | ~ ssItem(esk5_1(esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31])]),c_0_32])])).

<A NAME="c_0_39"></A>cnf(c_0_39,plain,
    ( ssItem(esk5_1(X1))
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,plain,
    ( X1 = X2
    | ~ segmentP(X2,X1)
    | ~ segmentP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( segmentP(esk51_0,esk48_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( esk51_0 = nil
    | singletonP(esk48_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_36]),c_0_15]),c_0_37])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( ~ singletonP(esk48_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_39]),c_0_32])])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( esk51_0 = esk48_0
    | ~ segmentP(esk48_0,esk51_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_41]),c_0_37]),c_0_32])])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( esk51_0 = nil ),
    inference(sr,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,<A HREF="#c_0_43">c_0_43</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( nil != esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_47"></A>fof(c_0_47,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( ~ segmentP(esk48_0,nil) ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_45]),c_0_45]),<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_49"></A>cnf(c_0_49,plain,
    ( segmentP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_49]),c_0_32])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC229+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n166.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:36:53 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.27  # AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
% 0.03/0.27  # and selection function SelectMaxLComplexAvoidPosPred.
% 0.03/0.27  #
% 0.03/0.27  # Preprocessing time       : 0.014 s
% 0.03/0.27  # Presaturation interreduction done
% 0.03/0.27  
% 0.03/0.27  # Proof found!
% 0.03/0.27  # SZS status Theorem
% 0.03/0.27  # SZS output start CNFRefutation
% 0.03/0.27  fof(c_0_0, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax38)).
% 0.03/0.27  fof(c_0_1, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~(segmentP(X4,X3)))|nil=X1)|?[X5]:(ssItem(X5)&?[X6]:(ssList(X6)&?[X7]:((ssList(X7)&app(app(X6,cons(X5,nil)),X7)=X1)&![X8]:(ssItem(X8)=>(((~(memberP(X6,X8))|~(memberP(X7,X8)))|~(leq(X5,X8)))|leq(X8,X5)))))))|(~(singletonP(X3))&neq(X4,nil)))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.03/0.27  fof(c_0_2, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.03/0.27  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
% 0.03/0.27  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax84)).
% 0.03/0.27  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)&lt;=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax4)).
% 0.03/0.27  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.03/0.27  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>((segmentP(X1,X2)&segmentP(X2,X1))=>X1=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax54)).
% 0.03/0.27  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>segmentP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax57)).
% 0.03/0.27  fof(c_0_9, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_0])).
% 0.03/0.27  fof(c_0_10, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~segmentP(X4,X3))|nil=X1)|?[X5]:(ssItem(X5)&?[X6]:(ssList(X6)&?[X7]:((ssList(X7)&app(app(X6,cons(X5,nil)),X7)=X1)&![X8]:(ssItem(X8)=>(((~memberP(X6,X8)|~memberP(X7,X8))|~leq(X5,X8))|leq(X8,X5)))))))|(~singletonP(X3)&neq(X4,nil))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_1])])).
% 0.03/0.27  fof(c_0_11, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).
% 0.03/0.27  fof(c_0_12, negated_conjecture, (![X13]:![X14]:![X15]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((((esk49_0=esk51_0&esk48_0=esk50_0)&segmentP(esk51_0,esk50_0))&nil!=esk48_0)&((((ssItem(esk52_3(X13,X14,X15))|(~ssList(X15)|app(app(X14,cons(X13,nil)),X15)!=esk48_0))|~ssList(X14))|~ssItem(X13))&((((((memberP(X14,esk52_3(X13,X14,X15))|(~ssList(X15)|app(app(X14,cons(X13,nil)),X15)!=esk48_0))|~ssList(X14))|~ssItem(X13))&(((memberP(X15,esk52_3(X13,X14,X15))|(~ssList(X15)|app(app(X14,cons(X13,nil)),X15)!=esk48_0))|~ssList(X14))|~ssItem(X13)))&(((leq(X13,esk52_3(X13,X14,X15))|(~ssList(X15)|app(app(X14,cons(X13,nil)),X15)!=esk48_0))|~ssList(X14))|~ssItem(X13)))&(((~leq(esk52_3(X13,X14,X15),X13)|(~ssList(X15)|app(app(X14,cons(X13,nil)),X15)!=esk48_0))|~ssList(X14))|~ssItem(X13)))))&(singletonP(esk50_0)|~neq(esk51_0,nil)))))))), 
% inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).
% 0.03/0.27  cnf(c_0_13, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.03/0.27  cnf(c_0_14, negated_conjecture, (memberP(X3,esk52_3(X1,X2,X3))|~ssItem(X1)|~ssList(X2)|app(app(X2,cons(X1,nil)),X3)!=esk48_0|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.03/0.27  cnf(c_0_15, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.03/0.27  cnf(c_0_16, negated_conjecture, (app(app(X1,cons(X2,nil)),nil)!=esk48_0|~ssList(X1)|~ssItem(esk52_3(X2,X1,nil))|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
% 0.03/0.27  cnf(c_0_17, negated_conjecture, (ssItem(esk52_3(X1,X2,X3))|~ssItem(X1)|~ssList(X2)|app(app(X2,cons(X1,nil)),X3)!=esk48_0|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.03/0.27  fof(c_0_18, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
% 0.03/0.27  cnf(c_0_19, negated_conjecture, (app(app(X1,cons(X2,nil)),nil)!=esk48_0|~ssList(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_15])])).
% 0.03/0.27  cnf(c_0_20, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
% 0.03/0.27  fof(c_0_21, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
% 0.03/0.27  cnf(c_0_22, negated_conjecture, (app(cons(X1,nil),nil)!=esk48_0|~ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_15])])).
% 0.03/0.27  cnf(c_0_23, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
% 0.03/0.27  fof(c_0_24, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
% 0.03/0.27  cnf(c_0_25, negated_conjecture, (singletonP(esk50_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.03/0.27  cnf(c_0_26, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.03/0.27  fof(c_0_27, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
% 0.03/0.27  cnf(c_0_28, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.03/0.27  cnf(c_0_29, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.03/0.27  cnf(c_0_30, negated_conjecture, (cons(X1,nil)!=esk48_0|~ssList(cons(X1,nil))|~ssItem(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
% 0.03/0.27  cnf(c_0_31, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
% 0.03/0.27  cnf(c_0_32, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.03/0.27  fof(c_0_33, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|((~segmentP(X3,X4)|~segmentP(X4,X3))|X3=X4)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
% 0.03/0.27  cnf(c_0_34, negated_conjecture, (segmentP(esk51_0,esk50_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.03/0.27  cnf(c_0_35, negated_conjecture, (singletonP(esk48_0)|~neq(esk51_0,nil)), inference(rw,[status(thm)],[c_0_25, c_0_26])).
% 0.03/0.27  cnf(c_0_36, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
% 0.03/0.27  cnf(c_0_37, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
% 0.03/0.27  cnf(c_0_38, negated_conjecture, (~singletonP(esk48_0)|~ssItem(esk5_1(esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31])]), c_0_32])])).
% 0.03/0.27  cnf(c_0_39, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
% 0.03/0.27  cnf(c_0_40, plain, (X1=X2|~segmentP(X2,X1)|~segmentP(X1,X2)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
% 0.03/0.27  cnf(c_0_41, negated_conjecture, (segmentP(esk51_0,esk48_0)), inference(rw,[status(thm)],[c_0_34, c_0_26])).
% 0.03/0.27  cnf(c_0_42, negated_conjecture, (esk51_0=nil|singletonP(esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_15]), c_0_37])])).
% 0.03/0.27  cnf(c_0_43, negated_conjecture, (~singletonP(esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_32])])).
% 0.03/0.27  cnf(c_0_44, negated_conjecture, (esk51_0=esk48_0|~segmentP(esk48_0,esk51_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_37]), c_0_32])])).
% 0.03/0.27  cnf(c_0_45, negated_conjecture, (esk51_0=nil), inference(sr,[status(thm)],[c_0_42, c_0_43])).
% 0.03/0.27  cnf(c_0_46, negated_conjecture, (nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.03/0.27  fof(c_0_47, plain, (![X2]:(~ssList(X2)|segmentP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
% 0.03/0.27  cnf(c_0_48, negated_conjecture, (~segmentP(esk48_0,nil)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45]), c_0_45]), c_0_46])).
% 0.03/0.27  cnf(c_0_49, plain, (segmentP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
% 0.03/0.27  cnf(c_0_50, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_32])]), ['proof']).
% 0.03/0.27  # SZS output end CNFRefutation
% 0.03/0.27  # Training examples: 0 positive, 0 negative
% 0.03/0.27  
% 0.03/0.27  # -------------------------------------------------
% 0.03/0.27  # User time                : 0.040 s
% 0.03/0.27  # System time              : 0.005 s
% 0.03/0.27  # Total time               : 0.045 s
% 0.03/0.27  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
