<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC264-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC264-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC264-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC264-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC264-1.p">SWC264-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n141.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:16 EDT 2016

% Result     : Unsatisfiable 151.56s
% Output     : CNFRefutation 151.56s
% Verified   : 
% Statistics : Number of clauses        :   17 (  19 expanded)
%              Number of leaves         :   13 (  15 expanded)
%              Depth                    :    3
%              Number of atoms          :   37 (  39 expanded)
%              Number of equality atoms :    3 (   3 expanded)
%              Maximal clause size      :    9 (   2 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,axiom,
    ( lt(X2,X4)
    | app(app(X1,cons(X2,X3)),cons(X4,X5)) != X6
    | ~ ssList(X5)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X4)
    | ~ ssItem(X2)
    | ~ strictorderedP(X6)
    | ~ ssList(X6) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause181)).

<A NAME="c_0_1"></A>cnf(c_0_1,axiom,
    ( totalorderedP(X1)
    | app(app(skaf66(X1),cons(skaf64(X1),skaf67(X1))),cons(skaf65(X1),skaf68(X1))) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause164)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( ssItem(skaf65(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause30)).

<A NAME="c_0_3"></A>cnf(c_0_3,axiom,
    ( ssItem(skaf64(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause31)).

<A NAME="c_0_4"></A>cnf(c_0_4,axiom,
    ( ssList(skaf68(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause27)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( ssList(skaf67(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause28)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( ssList(skaf66(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause29)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( leq(X1,X2)
    | ~ lt(X1,X2)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause103)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( totalorderedP(X1)
    | ~ leq(skaf64(X1),skaf65(X1))
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause91)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( strictorderedP(sk3) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ~ totalorderedP(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_13"></A>cnf(c_0_13,plain,
    ( lt(skaf64(X1),skaf65(X1))
    | totalorderedP(X1)
    | ~ ssList(X1)
    | ~ strictorderedP(X1) ),
    inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,c_0_1]),c_0_2]),c_0_3]),c_0_4]),c_0_5]),c_0_6])])])).

<A NAME="c_0_14"></A>cnf(c_0_14,plain,
    ( totalorderedP(X1)
    | ~ ssList(X1)
    | ~ strictorderedP(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_13]),c_0_2]),c_0_3])]),<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( strictorderedP(sk1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_14]),c_0_12]),c_0_15])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC264-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n141.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:38:23 CDT 2016
% 0.02/0.22  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 151.15/151.30  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 151.15/151.30  # and selection function SelectCQArNTEqFirst.
% 151.15/151.30  #
% 151.15/151.30  # Preprocessing time       : 0.010 s
% 151.15/151.30  # Presaturation interreduction done
% 151.55/151.71  # No success with AutoSched0
% 151.55/151.71  # Trying AutoSched1 for 73 seconds
% 151.56/151.90  # AutoSched1-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S00
% 151.56/151.90  # and selection function NoSelection.
% 151.56/151.90  #
% 151.56/151.90  # Preprocessing time       : 0.010 s
% 151.56/151.90  # Presaturation interreduction done
% 151.56/151.90  
% 151.56/151.90  # Proof found!
% 151.56/151.90  # SZS status Unsatisfiable
% 151.56/151.90  # SZS output start CNFRefutation
% 151.56/151.90  cnf(c_0_0, axiom, (lt(X2,X4)|app(app(X1,cons(X2,X3)),cons(X4,X5))!=X6|~ssList(X5)|~ssList(X3)|~ssList(X1)|~ssItem(X4)|~ssItem(X2)|~strictorderedP(X6)|~ssList(X6)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause181)).
% 151.56/151.90  cnf(c_0_1, axiom, (totalorderedP(X1)|app(app(skaf66(X1),cons(skaf64(X1),skaf67(X1))),cons(skaf65(X1),skaf68(X1)))=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause164)).
% 151.56/151.90  cnf(c_0_2, axiom, (ssItem(skaf65(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause30)).
% 151.56/151.90  cnf(c_0_3, axiom, (ssItem(skaf64(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause31)).
% 151.56/151.90  cnf(c_0_4, axiom, (ssList(skaf68(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause27)).
% 151.56/151.90  cnf(c_0_5, axiom, (ssList(skaf67(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause28)).
% 151.56/151.90  cnf(c_0_6, axiom, (ssList(skaf66(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause29)).
% 151.56/151.90  cnf(c_0_7, axiom, (leq(X1,X2)|~lt(X1,X2)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause103)).
% 151.56/151.90  cnf(c_0_8, axiom, (totalorderedP(X1)|~leq(skaf64(X1),skaf65(X1))|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause91)).
% 151.56/151.90  cnf(c_0_9, negated_conjecture, (strictorderedP(sk3)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 151.56/151.90  cnf(c_0_10, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 151.56/151.90  cnf(c_0_11, negated_conjecture, (~totalorderedP(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
% 151.56/151.90  cnf(c_0_12, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 151.56/151.90  cnf(c_0_13, plain, (lt(skaf64(X1),skaf65(X1))|totalorderedP(X1)|~ssList(X1)|~strictorderedP(X1)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2]), c_0_3]), c_0_4]), c_0_5]), c_0_6])])])).
% 151.56/151.90  cnf(c_0_14, plain, (totalorderedP(X1)|~ssList(X1)|~strictorderedP(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_13]), c_0_2]), c_0_3])]), c_0_8])).
% 151.56/151.90  cnf(c_0_15, negated_conjecture, (strictorderedP(sk1)), inference(rw,[status(thm)],[c_0_9, c_0_10])).
% 151.56/151.90  cnf(c_0_16, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_14]), c_0_12]), c_0_15])]), ['proof']).
% 151.56/151.90  # SZS output end CNFRefutation
% 151.56/151.90  # Training examples: 0 positive, 0 negative
% 151.56/151.90  
% 151.56/151.90  # -------------------------------------------------
% 151.56/151.90  # User time                : 147.472 s
% 151.56/151.90  # System time              : 4.227 s
% 151.56/151.90  # Total time               : 151.699 s
% 151.56/151.90  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
