<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC169+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC169+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC169+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC169%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC169+1.p">SWC169+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n175.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:55 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :  140 (43872 expanded)
%              Number of clauses        :  100 (16231 expanded)
%              Number of leaves         :   20 (11312 expanded)
%              Depth                    :   47
%              Number of atoms          :  412 (418748 expanded)
%              Number of equality atoms :  189 (154758 expanded)
%              Maximal formula depth    :   23 (   4 average)
%              Maximal clause size      :   24 (   2 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ~ ssList(X4)
                  | X2 != X4
                  | X1 != X3
                  | ! [X5] :
                      ( ssItem(X5)
                     => ! [X6] :
                          ( ssItem(X6)
                         => ! [X7] :
                              ( ssList(X7)
                             => ! [X8] :
                                  ( ~ ssList(X8)
                                  | app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8) != X1
                                  | neq(X5,X6) ) ) ) )
                  | ( ! [X9] :
                        ( ~ ssItem(X9)
                        | cons(X9,nil) != X3
                        | ~ memberP(X4,X9) )
                    & ( nil != X4
                      | nil != X3 ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax1)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( frontsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X2,X3) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax5)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax86)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => hd(app(X1,X2)) = hd(X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax85)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => cons(hd(X1),tl(X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax78)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => hd(cons(X2,X1)) = X2 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax23)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssList(tl(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X2,X3) = app(X2,X1)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax80)).

<A NAME="c_0_17"></A>fof(c_0_17,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_18"></A>fof(c_0_18,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => nil != cons(X2,X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax21)).

<A NAME="c_0_19"></A>fof(c_0_19,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_20"></A>fof(c_0_20,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ~ ssList(X4)
                    | X2 != X4
                    | X1 != X3
                    | ! [X5] :
                        ( ssItem(X5)
                       => ! [X6] :
                            ( ssItem(X6)
                           => ! [X7] :
                                ( ssList(X7)
                               => ! [X8] :
                                    ( ~ ssList(X8)
                                    | app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8) != X1
                                    | neq(X5,X6) ) ) ) )
                    | ( ! [X9] :
                          ( ~ ssItem(X9)
                          | cons(X9,nil) != X3
                          | ~ memberP(X4,X9) )
                      & ( nil != X4
                        | nil != X3 ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_21"></A>fof(c_0_21,negated_conjecture,
    ( ssList(esk48_0)
    & ssList(esk49_0)
    & ssList(esk50_0)
    & ssList(esk51_0)
    & esk49_0 = esk51_0
    & esk48_0 = esk50_0
    & ssItem(esk52_0)
    & ssItem(esk53_0)
    & ssList(esk54_0)
    & ssList(esk55_0)
    & app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0) = esk48_0
    & ~ neq(esk52_0,esk53_0)
    & ( nil = esk51_0
      | ssItem(esk56_0) )
    & ( nil = esk50_0
      | ssItem(esk56_0) )
    & ( nil = esk51_0
      | cons(esk56_0,nil) = esk50_0 )
    & ( nil = esk50_0
      | cons(esk56_0,nil) = esk50_0 )
    & ( nil = esk51_0
      | memberP(esk51_0,esk56_0) )
    & ( nil = esk50_0
      | memberP(esk51_0,esk56_0) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])])).

<A NAME="c_0_22"></A>fof(c_0_22,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssItem(X4)
        | ~ ssItem(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssItem(X4)
        | ~ ssItem(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0) = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( ~ neq(esk52_0,esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssItem(X1)
    | ~ ssItem(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( ssItem(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( ssItem(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( esk53_0 = esk52_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26]),c_0_27]),c_0_28])])).

<A NAME="c_0_31"></A>fof(c_0_31,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk52_0,nil)),esk55_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( ssList(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_35"></A>fof(c_0_35,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( app(app(esk54_0,app(cons(esk52_0,nil),cons(esk52_0,nil))),esk55_0) = esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_33]),c_0_34])])).

<A NAME="c_0_37"></A>cnf(c_0_37,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_39"></A>fof(c_0_39,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( app(app(esk54_0,app(cons(esk52_0,nil),cons(esk52_0,nil))),esk55_0) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_37]),c_0_38]),c_0_28])])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,cons(esk52_0,nil))),esk55_0) = esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_41]),c_0_28])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,cons(esk52_0,nil))),esk55_0) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_37]),c_0_38]),c_0_28])])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( ssList(esk55_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( app(esk54_0,app(cons(esk52_0,cons(esk52_0,nil)),esk55_0)) = esk50_0
    | ~ ssList(cons(esk52_0,cons(esk52_0,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_43]),c_0_44]),c_0_34])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( app(esk54_0,app(cons(esk52_0,cons(esk52_0,nil)),esk55_0)) = esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_37]),c_0_28])])).

<A NAME="c_0_47"></A>fof(c_0_47,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssItem(X6)
      | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( app(esk54_0,app(cons(esk52_0,cons(esk52_0,nil)),esk55_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_37]),c_0_38]),c_0_28])])).

<A NAME="c_0_49"></A>cnf(c_0_49,plain,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_50"></A>fof(c_0_50,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk6_2(X4,X5))
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(X5,esk6_2(X4,X5)) = X4
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X5,X7) != X4
        | frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).

<A NAME="c_0_51"></A>fof(c_0_51,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( app(esk54_0,cons(esk52_0,app(cons(esk52_0,nil),esk55_0))) = esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_49]),c_0_44]),c_0_28])])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( frontsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X2,X3) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_54"></A>cnf(c_0_54,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( app(esk54_0,cons(esk52_0,app(cons(esk52_0,nil),esk55_0))) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_37]),c_0_38]),c_0_28])])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_53]),<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( app(esk54_0,cons(esk52_0,cons(esk52_0,esk55_0))) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_41]),c_0_44]),c_0_28])])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( frontsegP(esk50_0,esk54_0)
    | ~ ssList(cons(esk52_0,cons(esk52_0,esk55_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_57]),c_0_34])])).

<A NAME="c_0_59"></A>fof(c_0_59,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( frontsegP(esk50_0,esk54_0)
    | ~ ssList(cons(esk52_0,esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_37]),c_0_28])])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_62"></A>cnf(c_0_62,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>])).

<A NAME="c_0_63"></A>fof(c_0_63,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).

<A NAME="c_0_64"></A>cnf(c_0_64,plain,
    ( app(X2,esk6_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,negated_conjecture,
    ( frontsegP(esk50_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>,c_0_37]),c_0_44]),c_0_28])])).

<A NAME="c_0_66"></A>cnf(c_0_66,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_67"></A>cnf(c_0_67,plain,
    ( ssList(esk6_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_68"></A>fof(c_0_68,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | hd(app(X3,X4)) = hd(X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).

<A NAME="c_0_69"></A>cnf(c_0_69,negated_conjecture,
    ( esk54_0 = nil
    | esk50_0 != nil
    | ~ ssList(app(cons(esk52_0,cons(esk52_0,nil)),esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_48]),c_0_34])])).

<A NAME="c_0_70"></A>cnf(c_0_70,plain,
    ( tl(app(X1,X2)) = app(tl(X1),X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( app(esk54_0,esk6_2(esk50_0,esk54_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,c_0_65]),c_0_34]),c_0_66])])).

<A NAME="c_0_72"></A>cnf(c_0_72,negated_conjecture,
    ( ssList(esk6_2(esk50_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,c_0_65]),c_0_34]),c_0_66])])).

<A NAME="c_0_73"></A>fof(c_0_73,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | cons(hd(X2),tl(X2)) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).

<A NAME="c_0_74"></A>cnf(c_0_74,plain,
    ( hd(app(X1,X2)) = hd(X1)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_75"></A>fof(c_0_75,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | hd(cons(X4,X3)) = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).

<A NAME="c_0_76"></A>cnf(c_0_76,negated_conjecture,
    ( esk54_0 = nil
    | esk50_0 != nil
    | ~ ssList(cons(esk52_0,cons(esk52_0,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>,c_0_54]),c_0_44])])).

<A NAME="c_0_77"></A>cnf(c_0_77,negated_conjecture,
    ( app(tl(esk54_0),esk6_2(esk50_0,esk54_0)) = tl(esk50_0)
    | esk54_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>,c_0_71]),c_0_72]),c_0_34])])).

<A NAME="c_0_78"></A>fof(c_0_78,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssList(tl(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).

<A NAME="c_0_79"></A>fof(c_0_79,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).

<A NAME="c_0_80"></A>cnf(c_0_80,plain,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>])).

<A NAME="c_0_81"></A>cnf(c_0_81,negated_conjecture,
    ( hd(esk54_0) = hd(esk50_0)
    | esk54_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,c_0_71]),c_0_72]),c_0_34])])).

<A NAME="c_0_82"></A>cnf(c_0_82,plain,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( cons(esk56_0,nil) = esk50_0
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( ssItem(esk56_0)
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( esk54_0 = nil
    | esk50_0 != nil
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>,c_0_37]),c_0_28])])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( tl(esk54_0) = nil
    | esk54_0 = nil
    | tl(esk50_0) != nil
    | ~ ssList(tl(esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_77]),c_0_72])])).

<A NAME="c_0_87"></A>cnf(c_0_87,plain,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_79">c_0_79</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( cons(hd(esk50_0),tl(esk54_0)) = esk54_0
    | esk54_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,c_0_81]),c_0_34])])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( hd(esk50_0) = esk56_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,c_0_83]),c_0_38])]),<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_91"></A>cnf(c_0_91,negated_conjecture,
    ( esk54_0 = nil
    | esk50_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_37]),c_0_38]),c_0_28])])).

<A NAME="c_0_92"></A>cnf(c_0_92,negated_conjecture,
    ( tl(esk54_0) = nil
    | esk54_0 = nil
    | tl(esk50_0) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>,c_0_87]),c_0_34])])).

<A NAME="c_0_93"></A>cnf(c_0_93,negated_conjecture,
    ( tl(esk50_0) = nil
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>,c_0_83]),c_0_38])]),<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_94"></A>fof(c_0_94,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X5,X6) != app(X5,X4)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).

<A NAME="c_0_95"></A>cnf(c_0_95,negated_conjecture,
    ( cons(esk56_0,tl(esk54_0)) = esk54_0
    | esk54_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>,c_0_90]),<A HREF="#c_0_91">c_0_91</A>])).

<A NAME="c_0_96"></A>cnf(c_0_96,negated_conjecture,
    ( tl(esk54_0) = nil
    | esk54_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_93]),<A HREF="#c_0_91">c_0_91</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,plain,
    ( X1 = X2
    | app(X3,X1) != app(X3,X2)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>])).

<A NAME="c_0_98"></A>fof(c_0_98,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).

<A NAME="c_0_99"></A>cnf(c_0_99,negated_conjecture,
    ( cons(esk56_0,nil) = esk54_0
    | esk54_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,<A HREF="#c_0_96">c_0_96</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,negated_conjecture,
    ( X1 = esk6_2(esk50_0,esk54_0)
    | app(esk54_0,X1) != esk50_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_71]),c_0_34]),c_0_72])])).

<A NAME="c_0_101"></A>cnf(c_0_101,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_98">c_0_98</A>])).

<A NAME="c_0_102"></A>cnf(c_0_102,negated_conjecture,
    ( esk54_0 = nil
    | esk54_0 = esk50_0 ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_99]),<A HREF="#c_0_91">c_0_91</A>])).

<A NAME="c_0_103"></A>cnf(c_0_103,negated_conjecture,
    ( esk6_2(esk50_0,esk54_0) = nil
    | esk54_0 != esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>,c_0_101]),c_0_38]),c_0_34])])).

<A NAME="c_0_104"></A>cnf(c_0_104,negated_conjecture,
    ( app(esk50_0,esk6_2(esk50_0,esk50_0)) = esk50_0
    | esk54_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>,<A HREF="#c_0_102">c_0_102</A>])).

<A NAME="c_0_105"></A>cnf(c_0_105,negated_conjecture,
    ( esk6_2(esk50_0,esk50_0) = nil
    | esk54_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_103">c_0_103</A>,<A HREF="#c_0_102">c_0_102</A>])).

<A NAME="c_0_106"></A>cnf(c_0_106,negated_conjecture,
    ( app(esk50_0,nil) = esk50_0
    | esk54_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>,<A HREF="#c_0_105">c_0_105</A>])).

<A NAME="c_0_107"></A>cnf(c_0_107,negated_conjecture,
    ( esk54_0 = nil
    | X1 = nil
    | app(esk50_0,X1) != esk50_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_106]),c_0_66]),c_0_38])])).

<A NAME="c_0_108"></A>cnf(c_0_108,negated_conjecture,
    ( app(esk50_0,cons(esk52_0,cons(esk52_0,esk55_0))) = esk50_0
    | esk54_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,<A HREF="#c_0_102">c_0_102</A>])).

<A NAME="c_0_109"></A>cnf(c_0_109,negated_conjecture,
    ( cons(esk52_0,cons(esk52_0,esk55_0)) = nil
    | esk54_0 = nil
    | ~ ssList(cons(esk52_0,cons(esk52_0,esk55_0))) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>,<A HREF="#c_0_108">c_0_108</A>])).

<A NAME="c_0_110"></A>fof(c_0_110,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | nil != cons(X4,X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).

<A NAME="c_0_111"></A>cnf(c_0_111,negated_conjecture,
    ( cons(esk52_0,cons(esk52_0,esk55_0)) = nil
    | esk54_0 = nil
    | ~ ssList(cons(esk52_0,esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>,c_0_37]),c_0_28])])).

<A NAME="c_0_112"></A>cnf(c_0_112,plain,
    ( nil != cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_113"></A>cnf(c_0_113,negated_conjecture,
    ( cons(esk52_0,cons(esk52_0,esk55_0)) = nil
    | esk54_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_111">c_0_111</A>,c_0_37]),c_0_44]),c_0_28])])).

<A NAME="c_0_114"></A>cnf(c_0_114,negated_conjecture,
    ( esk54_0 = nil
    | ~ ssList(cons(esk52_0,esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_113]),c_0_28])])).

<A NAME="c_0_115"></A>fof(c_0_115,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( esk54_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_114">c_0_114</A>,c_0_37]),c_0_44]),c_0_28])])).

<A NAME="c_0_117"></A>cnf(c_0_117,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_115">c_0_115</A>])).

<A NAME="c_0_118"></A>cnf(c_0_118,negated_conjecture,
    ( app(nil,cons(esk52_0,cons(esk52_0,esk55_0))) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,<A HREF="#c_0_116">c_0_116</A>])).

<A NAME="c_0_119"></A>cnf(c_0_119,negated_conjecture,
    ( app(app(nil,cons(esk52_0,cons(esk52_0,nil))),esk55_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,<A HREF="#c_0_116">c_0_116</A>])).

<A NAME="c_0_120"></A>cnf(c_0_120,negated_conjecture,
    ( cons(esk52_0,cons(esk52_0,esk55_0)) = esk50_0
    | ~ ssList(cons(esk52_0,cons(esk52_0,esk55_0))) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_117">c_0_117</A>,<A HREF="#c_0_118">c_0_118</A>])).

<A NAME="c_0_121"></A>cnf(c_0_121,negated_conjecture,
    ( app(cons(esk52_0,cons(esk52_0,nil)),esk55_0) = esk50_0
    | ~ ssList(cons(esk52_0,cons(esk52_0,nil))) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>,<A HREF="#c_0_117">c_0_117</A>])).

<A NAME="c_0_122"></A>cnf(c_0_122,negated_conjecture,
    ( cons(esk52_0,cons(esk52_0,esk55_0)) = esk50_0
    | ~ ssList(cons(esk52_0,esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_120">c_0_120</A>,c_0_37]),c_0_28])])).

<A NAME="c_0_123"></A>cnf(c_0_123,negated_conjecture,
    ( app(cons(esk52_0,cons(esk52_0,nil)),esk55_0) = esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_121">c_0_121</A>,c_0_37]),c_0_28])])).

<A NAME="c_0_124"></A>cnf(c_0_124,negated_conjecture,
    ( cons(esk52_0,cons(esk52_0,esk55_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_122">c_0_122</A>,c_0_37]),c_0_44]),c_0_28])])).

<A NAME="c_0_125"></A>cnf(c_0_125,negated_conjecture,
    ( app(cons(esk52_0,cons(esk52_0,nil)),esk55_0) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>,c_0_37]),c_0_38]),c_0_28])])).

<A NAME="c_0_126"></A>cnf(c_0_126,negated_conjecture,
    ( esk50_0 != nil
    | ~ ssList(cons(esk52_0,esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_124]),c_0_28])])).

<A NAME="c_0_127"></A>cnf(c_0_127,negated_conjecture,
    ( cons(esk52_0,app(cons(esk52_0,nil),esk55_0)) = esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_125]),c_0_44]),c_0_28])])).

<A NAME="c_0_128"></A>cnf(c_0_128,negated_conjecture,
    ( esk50_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_126">c_0_126</A>,c_0_37]),c_0_44]),c_0_28])])).

<A NAME="c_0_129"></A>cnf(c_0_129,negated_conjecture,
    ( cons(esk52_0,app(cons(esk52_0,nil),esk55_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_127">c_0_127</A>,c_0_37]),c_0_38]),c_0_28])])).

<A NAME="c_0_130"></A>cnf(c_0_130,negated_conjecture,
    ( hd(esk50_0) = esk56_0 ),
    inference(sr,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,<A HREF="#c_0_128">c_0_128</A>])).

<A NAME="c_0_131"></A>cnf(c_0_131,negated_conjecture,
    ( esk52_0 = esk56_0
    | ~ ssList(app(cons(esk52_0,nil),esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,c_0_129]),c_0_130]),c_0_28])])).

<A NAME="c_0_132"></A>cnf(c_0_132,negated_conjecture,
    ( esk52_0 = esk56_0
    | ~ ssList(cons(esk52_0,esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>,c_0_41]),c_0_44]),c_0_28])])).

<A NAME="c_0_133"></A>cnf(c_0_133,negated_conjecture,
    ( esk52_0 = esk56_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_132">c_0_132</A>,c_0_37]),c_0_44]),c_0_28])])).

<A NAME="c_0_134"></A>cnf(c_0_134,negated_conjecture,
    ( cons(esk56_0,cons(esk56_0,esk55_0)) = esk50_0 ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_124">c_0_124</A>,c_0_133]),<A HREF="#c_0_133">c_0_133</A>])).

<A NAME="c_0_135"></A>cnf(c_0_135,negated_conjecture,
    ( tl(esk50_0) = nil ),
    inference(sr,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,<A HREF="#c_0_128">c_0_128</A>])).

<A NAME="c_0_136"></A>cnf(c_0_136,negated_conjecture,
    ( ssItem(esk56_0) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,<A HREF="#c_0_128">c_0_128</A>])).

<A NAME="c_0_137"></A>cnf(c_0_137,negated_conjecture,
    ( cons(esk56_0,esk55_0) = nil
    | ~ ssList(cons(esk56_0,esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>,c_0_134]),c_0_135]),c_0_136])])).

<A NAME="c_0_138"></A>cnf(c_0_138,negated_conjecture,
    ( cons(esk56_0,esk55_0) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_137">c_0_137</A>,c_0_37]),c_0_44]),c_0_136])])).

<A NAME="c_0_139"></A>cnf(c_0_139,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_138]),c_0_44]),c_0_136])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC169+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.24  % Computer   : n175.star.cs.uiowa.edu
% 0.02/0.24  % Model      : x86_64 x86_64
% 0.02/0.24  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.24  % Memory     : 32218.75MB
% 0.02/0.24  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.24  % CPULimit   : 300
% 0.02/0.24  % DateTime   : Tue May 10 20:36:38 CDT 2016
% 0.02/0.24  % CPUTime    : 
% 0.02/0.24  # Version: 1.9.1pre011
% 0.02/0.24  # No SInE strategy applied
% 0.02/0.24  # Trying AutoSched0 for 151 seconds
% 0.06/0.33  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.33  # and selection function SelectCQPrecWNTNp.
% 0.06/0.33  #
% 0.06/0.33  # Preprocessing time       : 0.014 s
% 0.06/0.33  # Presaturation interreduction done
% 0.06/0.33  
% 0.06/0.33  # Proof found!
% 0.06/0.33  # SZS status Theorem
% 0.06/0.33  # SZS output start CNFRefutation
% 0.06/0.33  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~(ssList(X4))|X2!=X4)|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:((~(ssList(X8))|app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8)!=X1)|neq(X5,X6))))))|(![X9]:((~(ssItem(X9))|cons(X9,nil)!=X3)|~(memberP(X4,X9)))&(nil!=X4|nil!=X3))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.06/0.33  fof(c_0_1, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax1)).
% 0.06/0.33  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.06/0.33  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.06/0.33  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.33  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.06/0.33  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
% 0.06/0.33  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax5)).
% 0.06/0.33  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.06/0.33  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.06/0.33  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax86)).
% 0.06/0.33  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>hd(app(X1,X2))=hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax85)).
% 0.06/0.33  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax78)).
% 0.06/0.33  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax23)).
% 0.06/0.33  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
% 0.06/0.33  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
% 0.06/0.33  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X2,X3)=app(X2,X1)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax80)).
% 0.06/0.33  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
% 0.06/0.33  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>nil!=cons(X2,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax21)).
% 0.06/0.33  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
% 0.06/0.33  fof(c_0_20, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~ssList(X4)|X2!=X4)|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:((~ssList(X8)|app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8)!=X1)|neq(X5,X6))))))|(![X9]:((~ssItem(X9)|cons(X9,nil)!=X3)|~memberP(X4,X9))&(nil!=X4|nil!=X3)))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.33  fof(c_0_21, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(ssItem(esk52_0)&(ssItem(esk53_0)&(ssList(esk54_0)&((ssList(esk55_0)&app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0)=esk48_0)&~neq(esk52_0,esk53_0))))))&((((nil=esk51_0|ssItem(esk56_0))&(nil=esk50_0|ssItem(esk56_0)))&((nil=esk51_0|cons(esk56_0,nil)=esk50_0)&(nil=esk50_0|cons(esk56_0,nil)=esk50_0)))&((nil=esk51_0|memberP(esk51_0,esk56_0))&(nil=esk50_0|memberP(esk51_0,esk56_0))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])])).
% 0.06/0.33  fof(c_0_22, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssItem(X4))|~ssItem(X3))&(((X3=X4|neq(X3,X4))|~ssItem(X4))|~ssItem(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
% 0.06/0.33  cnf(c_0_23, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_21])).
% 0.06/0.33  cnf(c_0_24, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_21])).
% 0.06/0.33  cnf(c_0_25, negated_conjecture, (~neq(esk52_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
% 0.06/0.33  cnf(c_0_26, plain, (neq(X1,X2)|X1=X2|~ssItem(X1)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.33  cnf(c_0_27, negated_conjecture, (ssItem(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
% 0.06/0.33  cnf(c_0_28, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
% 0.06/0.33  cnf(c_0_29, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0)=esk50_0), inference(rw,[status(thm)],[c_0_23, c_0_24])).
% 0.06/0.33  cnf(c_0_30, negated_conjecture, (esk53_0=esk52_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28])])).
% 0.06/0.33  fof(c_0_31, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.06/0.33  cnf(c_0_32, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk52_0,nil)),esk55_0)=esk50_0), inference(rw,[status(thm)],[c_0_29, c_0_30])).
% 0.06/0.33  cnf(c_0_33, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
% 0.06/0.33  cnf(c_0_34, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
% 0.06/0.33  fof(c_0_35, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
% 0.06/0.33  cnf(c_0_36, negated_conjecture, (app(app(esk54_0,app(cons(esk52_0,nil),cons(esk52_0,nil))),esk55_0)=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])])).
% 0.06/0.33  cnf(c_0_37, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
% 0.06/0.33  cnf(c_0_38, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.06/0.33  fof(c_0_39, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 0.06/0.33  cnf(c_0_40, negated_conjecture, (app(app(esk54_0,app(cons(esk52_0,nil),cons(esk52_0,nil))),esk55_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_28])])).
% 0.06/0.33  cnf(c_0_41, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.33  cnf(c_0_42, negated_conjecture, (app(app(esk54_0,cons(esk52_0,cons(esk52_0,nil))),esk55_0)=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_28])])).
% 0.06/0.33  cnf(c_0_43, negated_conjecture, (app(app(esk54_0,cons(esk52_0,cons(esk52_0,nil))),esk55_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_37]), c_0_38]), c_0_28])])).
% 0.06/0.33  cnf(c_0_44, negated_conjecture, (ssList(esk55_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
% 0.06/0.33  cnf(c_0_45, negated_conjecture, (app(esk54_0,app(cons(esk52_0,cons(esk52_0,nil)),esk55_0))=esk50_0|~ssList(cons(esk52_0,cons(esk52_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_43]), c_0_44]), c_0_34])])).
% 0.06/0.33  cnf(c_0_46, negated_conjecture, (app(esk54_0,app(cons(esk52_0,cons(esk52_0,nil)),esk55_0))=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_37]), c_0_28])])).
% 0.06/0.33  fof(c_0_47, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
% 0.06/0.33  cnf(c_0_48, negated_conjecture, (app(esk54_0,app(cons(esk52_0,cons(esk52_0,nil)),esk55_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_37]), c_0_38]), c_0_28])])).
% 0.06/0.33  cnf(c_0_49, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
% 0.06/0.33  fof(c_0_50, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
% 0.06/0.33  fof(c_0_51, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
% 0.06/0.33  cnf(c_0_52, negated_conjecture, (app(esk54_0,cons(esk52_0,app(cons(esk52_0,nil),esk55_0)))=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_44]), c_0_28])])).
% 0.06/0.33  cnf(c_0_53, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_50])).
% 0.06/0.33  cnf(c_0_54, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
% 0.06/0.33  cnf(c_0_55, negated_conjecture, (app(esk54_0,cons(esk52_0,app(cons(esk52_0,nil),esk55_0)))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_37]), c_0_38]), c_0_28])])).
% 0.06/0.33  cnf(c_0_56, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_53]), c_0_54])).
% 0.06/0.33  cnf(c_0_57, negated_conjecture, (app(esk54_0,cons(esk52_0,cons(esk52_0,esk55_0)))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_41]), c_0_44]), c_0_28])])).
% 0.06/0.33  cnf(c_0_58, negated_conjecture, (frontsegP(esk50_0,esk54_0)|~ssList(cons(esk52_0,cons(esk52_0,esk55_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_34])])).
% 0.06/0.33  fof(c_0_59, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
% 0.06/0.33  cnf(c_0_60, negated_conjecture, (frontsegP(esk50_0,esk54_0)|~ssList(cons(esk52_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_37]), c_0_28])])).
% 0.06/0.33  cnf(c_0_61, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
% 0.06/0.33  cnf(c_0_62, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
% 0.06/0.33  fof(c_0_63, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
% 0.06/0.33  cnf(c_0_64, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
% 0.06/0.33  cnf(c_0_65, negated_conjecture, (frontsegP(esk50_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_37]), c_0_44]), c_0_28])])).
% 0.06/0.33  cnf(c_0_66, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_61, c_0_24])).
% 0.06/0.33  cnf(c_0_67, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
% 0.06/0.33  fof(c_0_68, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|hd(app(X3,X4))=hd(X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
% 0.06/0.33  cnf(c_0_69, negated_conjecture, (esk54_0=nil|esk50_0!=nil|~ssList(app(cons(esk52_0,cons(esk52_0,nil)),esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_48]), c_0_34])])).
% 0.06/0.33  cnf(c_0_70, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
% 0.06/0.33  cnf(c_0_71, negated_conjecture, (app(esk54_0,esk6_2(esk50_0,esk54_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_34]), c_0_66])])).
% 0.06/0.33  cnf(c_0_72, negated_conjecture, (ssList(esk6_2(esk50_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_65]), c_0_34]), c_0_66])])).
% 0.06/0.33  fof(c_0_73, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).
% 0.06/0.33  cnf(c_0_74, plain, (hd(app(X1,X2))=hd(X1)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
% 0.06/0.33  fof(c_0_75, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
% 0.06/0.33  cnf(c_0_76, negated_conjecture, (esk54_0=nil|esk50_0!=nil|~ssList(cons(esk52_0,cons(esk52_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_54]), c_0_44])])).
% 0.06/0.33  cnf(c_0_77, negated_conjecture, (app(tl(esk54_0),esk6_2(esk50_0,esk54_0))=tl(esk50_0)|esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72]), c_0_34])])).
% 0.06/0.33  fof(c_0_78, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).
% 0.06/0.33  fof(c_0_79, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
% 0.06/0.33  cnf(c_0_80, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
% 0.06/0.33  cnf(c_0_81, negated_conjecture, (hd(esk54_0)=hd(esk50_0)|esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_71]), c_0_72]), c_0_34])])).
% 0.06/0.33  cnf(c_0_82, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
% 0.06/0.33  cnf(c_0_83, negated_conjecture, (cons(esk56_0,nil)=esk50_0|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_21])).
% 0.06/0.33  cnf(c_0_84, negated_conjecture, (ssItem(esk56_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_21])).
% 0.06/0.33  cnf(c_0_85, negated_conjecture, (esk54_0=nil|esk50_0!=nil|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_37]), c_0_28])])).
% 0.06/0.33  cnf(c_0_86, negated_conjecture, (tl(esk54_0)=nil|esk54_0=nil|tl(esk50_0)!=nil|~ssList(tl(esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_77]), c_0_72])])).
% 0.06/0.33  cnf(c_0_87, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
% 0.06/0.33  cnf(c_0_88, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_79])).
% 0.06/0.33  cnf(c_0_89, negated_conjecture, (cons(hd(esk50_0),tl(esk54_0))=esk54_0|esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_34])])).
% 0.06/0.33  cnf(c_0_90, negated_conjecture, (hd(esk50_0)=esk56_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_38])]), c_0_84])).
% 0.06/0.33  cnf(c_0_91, negated_conjecture, (esk54_0=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_37]), c_0_38]), c_0_28])])).
% 0.06/0.33  cnf(c_0_92, negated_conjecture, (tl(esk54_0)=nil|esk54_0=nil|tl(esk50_0)!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_34])])).
% 0.06/0.33  cnf(c_0_93, negated_conjecture, (tl(esk50_0)=nil|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_83]), c_0_38])]), c_0_84])).
% 0.06/0.33  fof(c_0_94, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X5,X6)!=app(X5,X4)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
% 0.06/0.33  cnf(c_0_95, negated_conjecture, (cons(esk56_0,tl(esk54_0))=esk54_0|esk54_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_91])).
% 0.06/0.33  cnf(c_0_96, negated_conjecture, (tl(esk54_0)=nil|esk54_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_91])).
% 0.06/0.33  cnf(c_0_97, plain, (X1=X2|app(X3,X1)!=app(X3,X2)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_94])).
% 0.06/0.33  fof(c_0_98, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).
% 0.06/0.33  cnf(c_0_99, negated_conjecture, (cons(esk56_0,nil)=esk54_0|esk54_0=nil), inference(spm,[status(thm)],[c_0_95, c_0_96])).
% 0.06/0.33  cnf(c_0_100, negated_conjecture, (X1=esk6_2(esk50_0,esk54_0)|app(esk54_0,X1)!=esk50_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_71]), c_0_34]), c_0_72])])).
% 0.06/0.33  cnf(c_0_101, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_98])).
% 0.06/0.33  cnf(c_0_102, negated_conjecture, (esk54_0=nil|esk54_0=esk50_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_99]), c_0_91])).
% 0.06/0.33  cnf(c_0_103, negated_conjecture, (esk6_2(esk50_0,esk54_0)=nil|esk54_0!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_38]), c_0_34])])).
% 0.06/0.33  cnf(c_0_104, negated_conjecture, (app(esk50_0,esk6_2(esk50_0,esk50_0))=esk50_0|esk54_0=nil), inference(spm,[status(thm)],[c_0_71, c_0_102])).
% 0.06/0.33  cnf(c_0_105, negated_conjecture, (esk6_2(esk50_0,esk50_0)=nil|esk54_0=nil), inference(spm,[status(thm)],[c_0_103, c_0_102])).
% 0.06/0.33  cnf(c_0_106, negated_conjecture, (app(esk50_0,nil)=esk50_0|esk54_0=nil), inference(spm,[status(thm)],[c_0_104, c_0_105])).
% 0.06/0.33  cnf(c_0_107, negated_conjecture, (esk54_0=nil|X1=nil|app(esk50_0,X1)!=esk50_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_106]), c_0_66]), c_0_38])])).
% 0.06/0.33  cnf(c_0_108, negated_conjecture, (app(esk50_0,cons(esk52_0,cons(esk52_0,esk55_0)))=esk50_0|esk54_0=nil), inference(spm,[status(thm)],[c_0_57, c_0_102])).
% 0.06/0.33  cnf(c_0_109, negated_conjecture, (cons(esk52_0,cons(esk52_0,esk55_0))=nil|esk54_0=nil|~ssList(cons(esk52_0,cons(esk52_0,esk55_0)))), inference(spm,[status(thm)],[c_0_107, c_0_108])).
% 0.06/0.33  fof(c_0_110, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|nil!=cons(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
% 0.06/0.33  cnf(c_0_111, negated_conjecture, (cons(esk52_0,cons(esk52_0,esk55_0))=nil|esk54_0=nil|~ssList(cons(esk52_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_37]), c_0_28])])).
% 0.06/0.33  cnf(c_0_112, plain, (nil!=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_110])).
% 0.06/0.33  cnf(c_0_113, negated_conjecture, (cons(esk52_0,cons(esk52_0,esk55_0))=nil|esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_37]), c_0_44]), c_0_28])])).
% 0.06/0.33  cnf(c_0_114, negated_conjecture, (esk54_0=nil|~ssList(cons(esk52_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_28])])).
% 0.06/0.33  fof(c_0_115, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).
% 0.06/0.33  cnf(c_0_116, negated_conjecture, (esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_37]), c_0_44]), c_0_28])])).
% 0.06/0.33  cnf(c_0_117, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_115])).
% 0.06/0.33  cnf(c_0_118, negated_conjecture, (app(nil,cons(esk52_0,cons(esk52_0,esk55_0)))=esk50_0), inference(rw,[status(thm)],[c_0_57, c_0_116])).
% 0.06/0.33  cnf(c_0_119, negated_conjecture, (app(app(nil,cons(esk52_0,cons(esk52_0,nil))),esk55_0)=esk50_0), inference(rw,[status(thm)],[c_0_43, c_0_116])).
% 0.06/0.33  cnf(c_0_120, negated_conjecture, (cons(esk52_0,cons(esk52_0,esk55_0))=esk50_0|~ssList(cons(esk52_0,cons(esk52_0,esk55_0)))), inference(spm,[status(thm)],[c_0_117, c_0_118])).
% 0.06/0.33  cnf(c_0_121, negated_conjecture, (app(cons(esk52_0,cons(esk52_0,nil)),esk55_0)=esk50_0|~ssList(cons(esk52_0,cons(esk52_0,nil)))), inference(spm,[status(thm)],[c_0_119, c_0_117])).
% 0.06/0.33  cnf(c_0_122, negated_conjecture, (cons(esk52_0,cons(esk52_0,esk55_0))=esk50_0|~ssList(cons(esk52_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_37]), c_0_28])])).
% 0.06/0.33  cnf(c_0_123, negated_conjecture, (app(cons(esk52_0,cons(esk52_0,nil)),esk55_0)=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_37]), c_0_28])])).
% 0.06/0.33  cnf(c_0_124, negated_conjecture, (cons(esk52_0,cons(esk52_0,esk55_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_37]), c_0_44]), c_0_28])])).
% 0.06/0.33  cnf(c_0_125, negated_conjecture, (app(cons(esk52_0,cons(esk52_0,nil)),esk55_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_37]), c_0_38]), c_0_28])])).
% 0.06/0.33  cnf(c_0_126, negated_conjecture, (esk50_0!=nil|~ssList(cons(esk52_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_124]), c_0_28])])).
% 0.06/0.33  cnf(c_0_127, negated_conjecture, (cons(esk52_0,app(cons(esk52_0,nil),esk55_0))=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_125]), c_0_44]), c_0_28])])).
% 0.06/0.33  cnf(c_0_128, negated_conjecture, (esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_37]), c_0_44]), c_0_28])])).
% 0.06/0.33  cnf(c_0_129, negated_conjecture, (cons(esk52_0,app(cons(esk52_0,nil),esk55_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_37]), c_0_38]), c_0_28])])).
% 0.06/0.33  cnf(c_0_130, negated_conjecture, (hd(esk50_0)=esk56_0), inference(sr,[status(thm)],[c_0_90, c_0_128])).
% 0.06/0.33  cnf(c_0_131, negated_conjecture, (esk52_0=esk56_0|~ssList(app(cons(esk52_0,nil),esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_129]), c_0_130]), c_0_28])])).
% 0.06/0.33  cnf(c_0_132, negated_conjecture, (esk52_0=esk56_0|~ssList(cons(esk52_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_41]), c_0_44]), c_0_28])])).
% 0.06/0.33  cnf(c_0_133, negated_conjecture, (esk52_0=esk56_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_37]), c_0_44]), c_0_28])])).
% 0.06/0.33  cnf(c_0_134, negated_conjecture, (cons(esk56_0,cons(esk56_0,esk55_0))=esk50_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_133]), c_0_133])).
% 0.06/0.33  cnf(c_0_135, negated_conjecture, (tl(esk50_0)=nil), inference(sr,[status(thm)],[c_0_93, c_0_128])).
% 0.06/0.33  cnf(c_0_136, negated_conjecture, (ssItem(esk56_0)), inference(sr,[status(thm)],[c_0_84, c_0_128])).
% 0.06/0.33  cnf(c_0_137, negated_conjecture, (cons(esk56_0,esk55_0)=nil|~ssList(cons(esk56_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_134]), c_0_135]), c_0_136])])).
% 0.06/0.33  cnf(c_0_138, negated_conjecture, (cons(esk56_0,esk55_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_37]), c_0_44]), c_0_136])])).
% 0.06/0.33  cnf(c_0_139, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_138]), c_0_44]), c_0_136])]), ['proof']).
% 0.06/0.33  # SZS output end CNFRefutation
% 0.06/0.33  # Training examples: 0 positive, 0 negative
% 0.06/0.33  
% 0.06/0.33  # -------------------------------------------------
% 0.06/0.33  # User time                : 0.083 s
% 0.06/0.33  # System time              : 0.008 s
% 0.06/0.33  # Total time               : 0.090 s
% 0.06/0.33  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
