<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC161-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC161-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC161-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC161-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC161-1.p">SWC161-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n131.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:53 EDT 2016

% Result     : Unsatisfiable 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of clauses        :   36 ( 104 expanded)
%              Number of leaves         :   19 (  75 expanded)
%              Depth                    :   13
%              Number of atoms          :   75 ( 190 expanded)
%              Number of equality atoms :   18 (  38 expanded)
%              Maximal clause size      :    9 (   2 average)
%              Maximal term depth       :    6 (   2 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( app(app(app(app(sk9,cons(sk7,nil)),sk10),cons(sk8,nil)),sk11) = sk1 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_18)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( ssList(sk9) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_15)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssList(sk10) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_16)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ssItem(sk7) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ssList(sk11) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_17)).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( ssItem(sk8) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_14)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( lt(X1,X2)
    | X1 = X2
    | ~ leq(X1,X2)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause121)).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( leq(sk8,sk7) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_19)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( lt(X2,X4)
    | app(app(X1,cons(X2,X3)),cons(X4,X5)) != X6
    | ~ ssList(X5)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X4)
    | ~ ssItem(X2)
    | ~ strictorderedP(X6)
    | ~ ssList(X6) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause181)).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( ~ lt(X1,X2)
    | ~ lt(X2,X1)
    | ~ ssItem(X1)
    | ~ ssItem(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause114)).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( strictorderedP(sk3) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( ~ lt(X1,X1)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause63)).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( app(app(app(app(sk9,cons(sk7,nil)),sk10),cons(sk8,nil)),sk11) = sk3 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( app(app(app(sk9,app(cons(sk7,nil),sk10)),cons(sk8,nil)),sk11) = sk3
    | ~ ssList(cons(sk7,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_2]),c_0_3]),c_0_4])])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( app(app(app(sk9,app(cons(sk7,nil),sk10)),cons(sk8,nil)),sk11) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_5]),c_0_6]),c_0_7])])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( app(app(app(sk9,cons(sk7,sk10)),cons(sk8,nil)),sk11) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_8]),c_0_6]),c_0_4])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( app(app(sk9,cons(sk7,sk10)),app(cons(sk8,nil),sk11)) = sk3
    | ~ ssList(app(sk9,cons(sk7,sk10)))
    | ~ ssList(cons(sk8,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_22]),c_0_9])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( app(app(sk9,cons(sk7,sk10)),app(cons(sk8,nil),sk11)) = sk3
    | ~ ssList(app(sk9,cons(sk7,sk10))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_5]),c_0_10]),c_0_7])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( app(app(sk9,cons(sk7,sk10)),app(cons(sk8,nil),sk11)) = sk3
    | ~ ssList(cons(sk7,sk10)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_11]),c_0_3])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( app(app(sk9,cons(sk7,sk10)),app(cons(sk8,nil),sk11)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_5]),c_0_6]),c_0_4])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( sk8 = sk7
    | lt(sk8,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_13]),c_0_6]),c_0_10])])).

<A NAME="c_0_28"></A>cnf(c_0_28,plain,
    ( lt(X1,X2)
    | ~ ssItem(X2)
    | ~ ssItem(X1)
    | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( app(app(sk9,cons(sk7,sk10)),cons(sk8,sk11)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_8]),c_0_10]),c_0_9])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( sk8 = sk7
    | ~ lt(sk7,sk8) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_27]),c_0_10]),c_0_6])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( lt(sk7,sk8) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_29]),c_0_10]),c_0_6]),c_0_30]),c_0_9]),c_0_4]),c_0_3]),c_0_17])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( sk8 = sk7 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_32])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( lt(sk7,sk7) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,<A HREF="#c_0_33">c_0_33</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_34]),c_0_6])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC161-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n131.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:36:23 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.28  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.03/0.28  # and selection function SelectCQArNTEqFirst.
% 0.03/0.28  #
% 0.03/0.28  # Preprocessing time       : 0.009 s
% 0.03/0.28  # Presaturation interreduction done
% 0.03/0.28  
% 0.03/0.28  # Proof found!
% 0.03/0.28  # SZS status Unsatisfiable
% 0.03/0.28  # SZS output start CNFRefutation
% 0.03/0.28  cnf(c_0_0, negated_conjecture, (app(app(app(app(sk9,cons(sk7,nil)),sk10),cons(sk8,nil)),sk11)=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_18)).
% 0.03/0.28  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.03/0.28  cnf(c_0_2, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
% 0.03/0.28  cnf(c_0_3, negated_conjecture, (ssList(sk9)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_15)).
% 0.03/0.28  cnf(c_0_4, negated_conjecture, (ssList(sk10)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_16)).
% 0.03/0.28  cnf(c_0_5, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.03/0.28  cnf(c_0_6, negated_conjecture, (ssItem(sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
% 0.03/0.28  cnf(c_0_7, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.03/0.28  cnf(c_0_8, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
% 0.03/0.28  cnf(c_0_9, negated_conjecture, (ssList(sk11)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_17)).
% 0.03/0.28  cnf(c_0_10, negated_conjecture, (ssItem(sk8)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_14)).
% 0.03/0.28  cnf(c_0_11, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.03/0.28  cnf(c_0_12, axiom, (lt(X1,X2)|X1=X2|~leq(X1,X2)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause121)).
% 0.03/0.28  cnf(c_0_13, negated_conjecture, (leq(sk8,sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_19)).
% 0.03/0.28  cnf(c_0_14, axiom, (lt(X2,X4)|app(app(X1,cons(X2,X3)),cons(X4,X5))!=X6|~ssList(X5)|~ssList(X3)|~ssList(X1)|~ssItem(X4)|~ssItem(X2)|~strictorderedP(X6)|~ssList(X6)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause181)).
% 0.03/0.28  cnf(c_0_15, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.03/0.28  cnf(c_0_16, axiom, (~lt(X1,X2)|~lt(X2,X1)|~ssItem(X1)|~ssItem(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause114)).
% 0.03/0.28  cnf(c_0_17, negated_conjecture, (strictorderedP(sk3)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_10)).
% 0.03/0.28  cnf(c_0_18, axiom, (~lt(X1,X1)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause63)).
% 0.03/0.28  cnf(c_0_19, negated_conjecture, (app(app(app(app(sk9,cons(sk7,nil)),sk10),cons(sk8,nil)),sk11)=sk3), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 0.03/0.28  cnf(c_0_20, negated_conjecture, (app(app(app(sk9,app(cons(sk7,nil),sk10)),cons(sk8,nil)),sk11)=sk3|~ssList(cons(sk7,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_2]), c_0_3]), c_0_4])])).
% 0.03/0.28  cnf(c_0_21, negated_conjecture, (app(app(app(sk9,app(cons(sk7,nil),sk10)),cons(sk8,nil)),sk11)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_5]), c_0_6]), c_0_7])])).
% 0.03/0.28  cnf(c_0_22, negated_conjecture, (app(app(app(sk9,cons(sk7,sk10)),cons(sk8,nil)),sk11)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_8]), c_0_6]), c_0_4])])).
% 0.03/0.28  cnf(c_0_23, negated_conjecture, (app(app(sk9,cons(sk7,sk10)),app(cons(sk8,nil),sk11))=sk3|~ssList(app(sk9,cons(sk7,sk10)))|~ssList(cons(sk8,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_22]), c_0_9])])).
% 0.03/0.28  cnf(c_0_24, negated_conjecture, (app(app(sk9,cons(sk7,sk10)),app(cons(sk8,nil),sk11))=sk3|~ssList(app(sk9,cons(sk7,sk10)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_5]), c_0_10]), c_0_7])])).
% 0.03/0.28  cnf(c_0_25, negated_conjecture, (app(app(sk9,cons(sk7,sk10)),app(cons(sk8,nil),sk11))=sk3|~ssList(cons(sk7,sk10))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_11]), c_0_3])])).
% 0.03/0.28  cnf(c_0_26, negated_conjecture, (app(app(sk9,cons(sk7,sk10)),app(cons(sk8,nil),sk11))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_5]), c_0_6]), c_0_4])])).
% 0.03/0.28  cnf(c_0_27, negated_conjecture, (sk8=sk7|lt(sk8,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_6]), c_0_10])])).
% 0.03/0.28  cnf(c_0_28, plain, (lt(X1,X2)|~ssItem(X2)|~ssItem(X1)|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))), inference(er,[status(thm)],[c_0_14])).
% 0.03/0.28  cnf(c_0_29, negated_conjecture, (app(app(sk9,cons(sk7,sk10)),cons(sk8,sk11))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_8]), c_0_10]), c_0_9])])).
% 0.03/0.28  cnf(c_0_30, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_15, c_0_1])).
% 0.03/0.28  cnf(c_0_31, negated_conjecture, (sk8=sk7|~lt(sk7,sk8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_27]), c_0_10]), c_0_6])])).
% 0.03/0.28  cnf(c_0_32, negated_conjecture, (lt(sk7,sk8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_10]), c_0_6]), c_0_30]), c_0_9]), c_0_4]), c_0_3]), c_0_17])])).
% 0.03/0.28  cnf(c_0_33, negated_conjecture, (sk8=sk7), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_32])])).
% 0.03/0.28  cnf(c_0_34, negated_conjecture, (lt(sk7,sk7)), inference(rw,[status(thm)],[c_0_32, c_0_33])).
% 0.03/0.28  cnf(c_0_35, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_34]), c_0_6])]), ['proof']).
% 0.03/0.28  # SZS output end CNFRefutation
% 0.03/0.28  # Training examples: 0 positive, 0 negative
% 0.03/0.28  
% 0.03/0.28  # -------------------------------------------------
% 0.03/0.28  # User time                : 0.053 s
% 0.03/0.28  # System time              : 0.004 s
% 0.03/0.28  # Total time               : 0.057 s
% 0.03/0.28  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
