<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC031+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC031+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC031+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC031%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC031+1.p">SWC031+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n134.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:22 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :  141 (143800 expanded)
%              Number of clauses        :   91 (55679 expanded)
%              Number of leaves         :   25 (34191 expanded)
%              Depth                    :   32
%              Number of atoms          :  487 (1466060 expanded)
%              Number of equality atoms :  130 (562906 expanded)
%              Maximal formula depth    :   20 (   4 average)
%              Maximal clause size      :   21 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ~ frontsegP(X4,X3)
                    | ~ totalorderedP(X3)
                    | ? [X5] :
                        ( ssList(X5)
                        & neq(X3,X5)
                        & frontsegP(X4,X5)
                        & segmentP(X5,X3)
                        & totalorderedP(X5) )
                    | ( ( nil != X2
                        | nil = X1 )
                      & ( nil != X1
                        | nil = X2 ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( frontsegP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax46)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( frontsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X2,X3) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax5)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( segmentP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax7)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( rearsegP(X1,X2)
               => rearsegP(app(X3,X1),X2) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax50)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => rearsegP(X1,nil) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax51)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( rearsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax6)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => hd(app(X1,X2)) = hd(X1) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax85)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( frontsegP(cons(X1,X3),cons(X2,X4))
                  &lt;=> ( X1 = X2
                      & frontsegP(X3,X4) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax44)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => cons(hd(X1),tl(X1)) = X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax78)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssList(tl(X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax24)).

<A NAME="c_0_17"></A>fof(c_0_17,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax86)).

<A NAME="c_0_18"></A>fof(c_0_18,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,nil) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax45)).

<A NAME="c_0_19"></A>fof(c_0_19,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssItem(hd(X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax22)).

<A NAME="c_0_20"></A>fof(c_0_20,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,nil) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax57)).

<A NAME="c_0_21"></A>fof(c_0_21,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_22"></A>fof(c_0_22,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => totalorderedP(cons(X1,nil)) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax65)).

<A NAME="c_0_23"></A>fof(c_0_23,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_24"></A>fof(c_0_24,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) != X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax18)).

<A NAME="c_0_25"></A>fof(c_0_25,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ~ frontsegP(X4,X3)
                      | ~ totalorderedP(X3)
                      | ? [X5] :
                          ( ssList(X5)
                          & neq(X3,X5)
                          & frontsegP(X4,X5)
                          & segmentP(X5,X3)
                          & totalorderedP(X5) )
                      | ( ( nil != X2
                          | nil = X1 )
                        & ( nil != X1
                          | nil = X2 ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_26"></A>fof(c_0_26,negated_conjecture,(
    ! [X10] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & frontsegP(esk51_0,esk50_0)
      & totalorderedP(esk50_0)
      & ( ~ ssList(X10)
        | ~ neq(esk50_0,X10)
        | ~ frontsegP(esk51_0,X10)
        | ~ segmentP(X10,esk50_0)
        | ~ totalorderedP(X10) )
      & ( nil = esk48_0
        | nil = esk49_0 )
      & ( nil != esk49_0
        | nil = esk49_0 )
      & ( nil = esk48_0
        | nil != esk48_0 )
      & ( nil != esk49_0
        | nil != esk48_0 ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( nil = esk49_0
    | nil = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_30"></A>fof(c_0_30,plain,(
    ! [X2] :
      ( ( ~ frontsegP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | frontsegP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( frontsegP(esk51_0,esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( esk51_0 = nil
    | esk50_0 = nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_28]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ frontsegP(nil,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( esk50_0 = nil
    | frontsegP(nil,esk50_0) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_37"></A>fof(c_0_37,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk6_2(X4,X5))
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(X5,esk6_2(X4,X5)) = X4
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X5,X7) != X4
        | frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_35]),c_0_36])])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_40"></A>fof(c_0_40,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( app(X2,esk6_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( frontsegP(esk51_0,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( ssList(esk6_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>])).

<A NAME="c_0_46"></A>fof(c_0_46,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk8_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ssList(esk9_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(app(X9,X6),X10) != X5
        | segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).

<A NAME="c_0_47"></A>cnf(c_0_47,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( app(nil,esk6_2(esk51_0,nil)) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_42]),c_0_43]),c_0_44])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( ssList(esk6_2(esk51_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_42]),c_0_43]),c_0_44])])).

<A NAME="c_0_50"></A>cnf(c_0_50,plain,
    ( segmentP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(app(X3,X2),X4) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( esk6_2(esk51_0,nil) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_48]),c_0_49])])).

<A NAME="c_0_52"></A>cnf(c_0_52,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( app(nil,esk51_0) = esk51_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,<A HREF="#c_0_51">c_0_51</A>])).

<A NAME="c_0_54"></A>fof(c_0_54,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( segmentP(app(esk51_0,X1),esk51_0)
    | ~ ssList(app(esk51_0,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_53]),c_0_43]),c_0_44])])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_57"></A>fof(c_0_57,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).

<A NAME="c_0_58"></A>fof(c_0_58,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

<A NAME="c_0_59"></A>cnf(c_0_59,plain,
    ( app(app(esk8_2(X1,X2),X2),esk9_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ segmentP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( segmentP(esk51_0,esk51_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_56]),c_0_44]),c_0_43])])).

<A NAME="c_0_61"></A>cnf(c_0_61,plain,
    ( ssList(esk9_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ segmentP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_62"></A>cnf(c_0_62,plain,
    ( ssList(esk8_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ segmentP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_63"></A>cnf(c_0_63,plain,
    ( frontsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X2,X3) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>])).

<A NAME="c_0_64"></A>cnf(c_0_64,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,negated_conjecture,
    ( app(app(esk8_2(esk51_0,esk51_0),esk51_0),esk9_2(esk51_0,esk51_0)) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,c_0_60]),c_0_44])])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( ssList(esk9_2(esk51_0,esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_60]),c_0_44])])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( ssList(esk8_2(esk51_0,esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_60]),c_0_44])])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_63]),<A HREF="#c_0_64">c_0_64</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( app(esk8_2(esk51_0,esk51_0),app(esk51_0,esk9_2(esk51_0,esk51_0))) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_66]),c_0_67]),c_0_44]),c_0_68])])).

<A NAME="c_0_71"></A>fof(c_0_71,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ rearsegP(X4,X5)
      | rearsegP(app(X6,X4),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_72"></A>fof(c_0_72,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | rearsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).

<A NAME="c_0_73"></A>cnf(c_0_73,negated_conjecture,
    ( frontsegP(esk51_0,esk8_2(esk51_0,esk51_0))
    | ~ ssList(app(esk51_0,esk9_2(esk51_0,esk51_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>,c_0_70]),c_0_68])])).

<A NAME="c_0_74"></A>cnf(c_0_74,plain,
    ( rearsegP(app(X1,X2),X3)
    | ~ rearsegP(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>])).

<A NAME="c_0_75"></A>cnf(c_0_75,plain,
    ( rearsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,negated_conjecture,
    ( frontsegP(esk51_0,esk8_2(esk51_0,esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>,c_0_64]),c_0_67]),c_0_44])])).

<A NAME="c_0_77"></A>fof(c_0_77,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk7_2(X4,X5))
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(esk7_2(X4,X5),X5) = X4
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X7,X5) != X4
        | rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).

<A NAME="c_0_78"></A>cnf(c_0_78,plain,
    ( rearsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,c_0_75]),c_0_43])])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( app(esk8_2(esk51_0,esk51_0),esk6_2(esk51_0,esk8_2(esk51_0,esk51_0))) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_76]),c_0_68]),c_0_44])])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( ssList(esk6_2(esk51_0,esk8_2(esk51_0,esk51_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_76]),c_0_68]),c_0_44])])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( app(esk7_2(X1,X2),X2) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( rearsegP(esk51_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_79]),c_0_80]),c_0_68])])).

<A NAME="c_0_83"></A>cnf(c_0_83,plain,
    ( ssList(esk7_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( app(esk7_2(esk51_0,nil),nil) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_82]),c_0_43]),c_0_44])])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( ssList(esk7_2(esk51_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_82]),c_0_43]),c_0_44])])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( esk7_2(esk51_0,nil) = esk51_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_84]),c_0_85])])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( nil != esk48_0
    | nil != esk49_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_88"></A>fof(c_0_88,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | hd(app(X3,X4)) = hd(X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( app(esk51_0,nil) = esk51_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,<A HREF="#c_0_86">c_0_86</A>])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( esk50_0 != nil
    | esk51_0 != nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,c_0_28]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_91"></A>fof(c_0_91,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).

<A NAME="c_0_92"></A>fof(c_0_92,plain,(
    ! [X5,X6,X7,X8] :
      ( ( X5 = X6
        | ~ frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) )
      & ( frontsegP(X7,X8)
        | ~ frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) )
      & ( X5 != X6
        | ~ frontsegP(X7,X8)
        | frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])).

<A NAME="c_0_93"></A>fof(c_0_93,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | cons(hd(X2),tl(X2)) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).

<A NAME="c_0_94"></A>fof(c_0_94,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssList(tl(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).

<A NAME="c_0_95"></A>fof(c_0_95,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).

<A NAME="c_0_96"></A>cnf(c_0_96,plain,
    ( hd(app(X1,X2)) = hd(X1)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,negated_conjecture,
    ( app(esk51_0,app(nil,X1)) = app(esk51_0,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_89]),c_0_43]),c_0_44])])).

<A NAME="c_0_98"></A>cnf(c_0_98,negated_conjecture,
    ( esk51_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,c_0_38])])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,plain,
    ( frontsegP(cons(X1,X3),cons(X2,X4))
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ ssList(X4)
    | ~ frontsegP(X3,X4)
    | X1 != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>])).

<A NAME="c_0_101"></A>fof(c_0_101,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])).

<A NAME="c_0_102"></A>cnf(c_0_102,plain,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>])).

<A NAME="c_0_103"></A>cnf(c_0_103,plain,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>])).

<A NAME="c_0_104"></A>cnf(c_0_104,plain,
    ( tl(app(X1,X2)) = app(tl(X1),X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>])).

<A NAME="c_0_105"></A>fof(c_0_105,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssItem(hd(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).

<A NAME="c_0_106"></A>cnf(c_0_106,negated_conjecture,
    ( hd(app(esk51_0,X1)) = hd(esk51_0)
    | ~ ssList(app(nil,X1))
    | ~ ssList(X1) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,c_0_97]),c_0_44])]),<A HREF="#c_0_98">c_0_98</A>])).

<A NAME="c_0_107"></A>cnf(c_0_107,negated_conjecture,
    ( ssList(app(esk51_0,X1))
    | ~ ssList(app(nil,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,c_0_97]),c_0_44])])).

<A NAME="c_0_108"></A>cnf(c_0_108,negated_conjecture,
    ( app(esk51_0,X1) != nil
    | ~ ssList(app(nil,X1))
    | ~ ssList(X1) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_97]),c_0_44])]),<A HREF="#c_0_98">c_0_98</A>])).

<A NAME="c_0_109"></A>fof(c_0_109,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])).

<A NAME="c_0_110"></A>cnf(c_0_110,plain,
    ( frontsegP(cons(X1,X2),cons(X1,X3))
    | ~ frontsegP(X2,X3)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>])).

<A NAME="c_0_111"></A>cnf(c_0_111,plain,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_101">c_0_101</A>])).

<A NAME="c_0_112"></A>cnf(c_0_112,plain,
    ( cons(hd(X1),tl(app(X1,X2))) = app(X1,X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>,c_0_96]),c_0_64]),<A HREF="#c_0_99">c_0_99</A>])).

<A NAME="c_0_113"></A>cnf(c_0_113,plain,
    ( nil = X1
    | ssList(app(tl(X1),X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_103">c_0_103</A>,c_0_104]),c_0_64]),<A HREF="#c_0_99">c_0_99</A>])).

<A NAME="c_0_114"></A>cnf(c_0_114,negated_conjecture,
    ( app(tl(esk51_0),nil) = tl(esk51_0) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>,c_0_89]),c_0_43]),c_0_44])]),<A HREF="#c_0_98">c_0_98</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,plain,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_105">c_0_105</A>])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( hd(app(esk51_0,esk51_0)) = hd(esk51_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>,c_0_53]),c_0_44])])).

<A NAME="c_0_117"></A>cnf(c_0_117,negated_conjecture,
    ( ssList(app(esk51_0,esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>,c_0_53]),c_0_44])])).

<A NAME="c_0_118"></A>cnf(c_0_118,negated_conjecture,
    ( app(esk51_0,esk51_0) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>,c_0_53]),c_0_44])])).

<A NAME="c_0_119"></A>cnf(c_0_119,negated_conjecture,
    ( ~ totalorderedP(X1)
    | ~ segmentP(X1,esk50_0)
    | ~ frontsegP(esk51_0,X1)
    | ~ neq(esk50_0,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_120"></A>cnf(c_0_120,plain,
    ( segmentP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>])).

<A NAME="c_0_121"></A>cnf(c_0_121,plain,
    ( frontsegP(cons(X1,X2),cons(X1,nil))
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_110">c_0_110</A>,c_0_111]),c_0_43])])).

<A NAME="c_0_122"></A>cnf(c_0_122,negated_conjecture,
    ( cons(hd(esk51_0),tl(esk51_0)) = esk51_0 ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_89]),c_0_43]),c_0_44])]),<A HREF="#c_0_98">c_0_98</A>])).

<A NAME="c_0_123"></A>cnf(c_0_123,negated_conjecture,
    ( ssList(tl(esk51_0)) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_113">c_0_113</A>,c_0_114]),c_0_43]),c_0_44])]),<A HREF="#c_0_98">c_0_98</A>])).

<A NAME="c_0_124"></A>cnf(c_0_124,negated_conjecture,
    ( ssItem(hd(esk51_0)) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_115">c_0_115</A>,c_0_116]),c_0_117])]),<A HREF="#c_0_118">c_0_118</A>])).

<A NAME="c_0_125"></A>cnf(c_0_125,negated_conjecture,
    ( ~ totalorderedP(X1)
    | ~ frontsegP(esk51_0,X1)
    | ~ ssList(X1)
    | ~ neq(nil,X1) ),
    inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>,c_0_38]),c_0_38]),<A HREF="#c_0_120">c_0_120</A>])).

<A NAME="c_0_126"></A>cnf(c_0_126,negated_conjecture,
    ( frontsegP(esk51_0,cons(hd(esk51_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_121">c_0_121</A>,c_0_122]),c_0_123]),c_0_124])])).

<A NAME="c_0_127"></A>fof(c_0_127,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).

<A NAME="c_0_128"></A>cnf(c_0_128,negated_conjecture,
    ( ~ totalorderedP(cons(hd(esk51_0),nil))
    | ~ ssList(cons(hd(esk51_0),nil))
    | ~ neq(nil,cons(hd(esk51_0),nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_125">c_0_125</A>,<A HREF="#c_0_126">c_0_126</A>])).

<A NAME="c_0_129"></A>cnf(c_0_129,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_127">c_0_127</A>])).

<A NAME="c_0_130"></A>fof(c_0_130,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | totalorderedP(cons(X2,nil)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])).

<A NAME="c_0_131"></A>cnf(c_0_131,negated_conjecture,
    ( ~ totalorderedP(cons(hd(esk51_0),nil))
    | ~ neq(nil,cons(hd(esk51_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_128">c_0_128</A>,c_0_129]),c_0_43]),c_0_124])])).

<A NAME="c_0_132"></A>cnf(c_0_132,plain,
    ( totalorderedP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_130">c_0_130</A>])).

<A NAME="c_0_133"></A>fof(c_0_133,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).

<A NAME="c_0_134"></A>cnf(c_0_134,negated_conjecture,
    ( ~ neq(nil,cons(hd(esk51_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>,c_0_132]),c_0_124])])).

<A NAME="c_0_135"></A>cnf(c_0_135,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_133">c_0_133</A>])).

<A NAME="c_0_136"></A>fof(c_0_136,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) != X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).

<A NAME="c_0_137"></A>cnf(c_0_137,negated_conjecture,
    ( cons(hd(esk51_0),nil) = nil
    | ~ ssList(cons(hd(esk51_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_134">c_0_134</A>,c_0_135]),c_0_43])])).

<A NAME="c_0_138"></A>cnf(c_0_138,plain,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_136">c_0_136</A>])).

<A NAME="c_0_139"></A>cnf(c_0_139,negated_conjecture,
    ( cons(hd(esk51_0),nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_137">c_0_137</A>,c_0_129]),c_0_43]),c_0_124])])).

<A NAME="c_0_140"></A>cnf(c_0_140,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_138">c_0_138</A>,c_0_139]),c_0_43]),c_0_124])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC031+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n134.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:34:38 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.47  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.47  # and selection function SelectCQPrecWNTNp.
% 0.06/0.47  #
% 0.06/0.47  # Preprocessing time       : 0.014 s
% 0.06/0.47  # Presaturation interreduction done
% 0.06/0.47  
% 0.06/0.47  # Proof found!
% 0.06/0.47  # SZS status Theorem
% 0.06/0.47  # SZS output start CNFRefutation
% 0.06/0.47  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~(frontsegP(X4,X3)))|~(totalorderedP(X3)))|?[X5]:((((ssList(X5)&neq(X3,X5))&frontsegP(X4,X5))&segmentP(X5,X3))&totalorderedP(X5)))|((nil!=X2|nil=X1)&(nil!=X1|nil=X2)))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.06/0.47  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>(frontsegP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax46)).
% 0.06/0.47  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax5)).
% 0.06/0.47  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
% 0.06/0.47  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.47  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax7)).
% 0.06/0.47  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax84)).
% 0.06/0.47  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.06/0.47  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.06/0.47  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(rearsegP(X1,X2)=>rearsegP(app(X3,X1),X2)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax50)).
% 0.06/0.47  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>rearsegP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax51)).
% 0.06/0.47  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax6)).
% 0.06/0.47  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>hd(app(X1,X2))=hd(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax85)).
% 0.06/0.47  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.06/0.47  fof(c_0_14, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(frontsegP(cons(X1,X3),cons(X2,X4))&lt;=>(X1=X2&frontsegP(X3,X4)))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax44)).
% 0.06/0.47  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax78)).
% 0.06/0.47  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax24)).
% 0.06/0.47  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax86)).
% 0.06/0.47  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax45)).
% 0.06/0.47  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax22)).
% 0.06/0.47  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>segmentP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax57)).
% 0.06/0.47  fof(c_0_21, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.06/0.47  fof(c_0_22, axiom, (![X1]:(ssItem(X1)=>totalorderedP(cons(X1,nil)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax65)).
% 0.06/0.47  fof(c_0_23, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.06/0.47  fof(c_0_24, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax18)).
% 0.06/0.47  fof(c_0_25, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~frontsegP(X4,X3))|~totalorderedP(X3))|?[X5]:((((ssList(X5)&neq(X3,X5))&frontsegP(X4,X5))&segmentP(X5,X3))&totalorderedP(X5)))|((nil!=X2|nil=X1)&(nil!=X1|nil=X2))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.47  fof(c_0_26, negated_conjecture, (![X10]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((((esk49_0=esk51_0&esk48_0=esk50_0)&frontsegP(esk51_0,esk50_0))&totalorderedP(esk50_0))&((((~ssList(X10)|~neq(esk50_0,X10))|~frontsegP(esk51_0,X10))|~segmentP(X10,esk50_0))|~totalorderedP(X10)))&(((nil=esk48_0|nil=esk49_0)&(nil!=esk49_0|nil=esk49_0))&((nil=esk48_0|nil!=esk48_0)&(nil!=esk49_0|nil!=esk48_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])])).
% 0.06/0.47  cnf(c_0_27, negated_conjecture, (nil=esk49_0|nil=esk48_0), inference(split_conjunct,[status(thm)],[c_0_26])).
% 0.06/0.47  cnf(c_0_28, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_26])).
% 0.06/0.47  cnf(c_0_29, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_26])).
% 0.06/0.47  fof(c_0_30, plain, (![X2]:(((~frontsegP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|frontsegP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.06/0.47  cnf(c_0_31, negated_conjecture, (frontsegP(esk51_0,esk50_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
% 0.06/0.47  cnf(c_0_32, negated_conjecture, (esk51_0=nil|esk50_0=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
% 0.06/0.47  cnf(c_0_33, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
% 0.06/0.47  cnf(c_0_34, plain, (nil=X1|~ssList(X1)|~frontsegP(nil,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
% 0.06/0.47  cnf(c_0_35, negated_conjecture, (esk50_0=nil|frontsegP(nil,esk50_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
% 0.06/0.47  cnf(c_0_36, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_33, c_0_28])).
% 0.06/0.47  fof(c_0_37, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
% 0.06/0.47  cnf(c_0_38, negated_conjecture, (esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
% 0.06/0.47  cnf(c_0_39, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
% 0.06/0.47  fof(c_0_40, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
% 0.06/0.47  cnf(c_0_41, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
% 0.06/0.47  cnf(c_0_42, negated_conjecture, (frontsegP(esk51_0,nil)), inference(rw,[status(thm)],[c_0_31, c_0_38])).
% 0.06/0.47  cnf(c_0_43, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.06/0.47  cnf(c_0_44, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_39, c_0_29])).
% 0.06/0.47  cnf(c_0_45, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
% 0.06/0.47  fof(c_0_46, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
% 0.06/0.47  cnf(c_0_47, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
% 0.06/0.47  cnf(c_0_48, negated_conjecture, (app(nil,esk6_2(esk51_0,nil))=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44])])).
% 0.06/0.47  cnf(c_0_49, negated_conjecture, (ssList(esk6_2(esk51_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_42]), c_0_43]), c_0_44])])).
% 0.06/0.47  cnf(c_0_50, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_46])).
% 0.06/0.47  cnf(c_0_51, negated_conjecture, (esk6_2(esk51_0,nil)=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])])).
% 0.06/0.47  cnf(c_0_52, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_50])).
% 0.06/0.47  cnf(c_0_53, negated_conjecture, (app(nil,esk51_0)=esk51_0), inference(rw,[status(thm)],[c_0_48, c_0_51])).
% 0.06/0.47  fof(c_0_54, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
% 0.06/0.47  cnf(c_0_55, negated_conjecture, (segmentP(app(esk51_0,X1),esk51_0)|~ssList(app(esk51_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_43]), c_0_44])])).
% 0.06/0.47  cnf(c_0_56, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
% 0.06/0.47  fof(c_0_57, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
% 0.06/0.47  fof(c_0_58, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
% 0.06/0.47  cnf(c_0_59, plain, (app(app(esk8_2(X1,X2),X2),esk9_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
% 0.06/0.47  cnf(c_0_60, negated_conjecture, (segmentP(esk51_0,esk51_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_44]), c_0_43])])).
% 0.06/0.47  cnf(c_0_61, plain, (ssList(esk9_2(X1,X2))|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
% 0.06/0.47  cnf(c_0_62, plain, (ssList(esk8_2(X1,X2))|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
% 0.06/0.47  cnf(c_0_63, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
% 0.06/0.47  cnf(c_0_64, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
% 0.06/0.47  cnf(c_0_65, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
% 0.06/0.47  cnf(c_0_66, negated_conjecture, (app(app(esk8_2(esk51_0,esk51_0),esk51_0),esk9_2(esk51_0,esk51_0))=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_44])])).
% 0.06/0.47  cnf(c_0_67, negated_conjecture, (ssList(esk9_2(esk51_0,esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_60]), c_0_44])])).
% 0.06/0.47  cnf(c_0_68, negated_conjecture, (ssList(esk8_2(esk51_0,esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_60]), c_0_44])])).
% 0.06/0.47  cnf(c_0_69, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_63]), c_0_64])).
% 0.06/0.47  cnf(c_0_70, negated_conjecture, (app(esk8_2(esk51_0,esk51_0),app(esk51_0,esk9_2(esk51_0,esk51_0)))=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67]), c_0_44]), c_0_68])])).
% 0.06/0.47  fof(c_0_71, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~rearsegP(X4,X5)|rearsegP(app(X6,X4),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 0.06/0.47  fof(c_0_72, plain, (![X2]:(~ssList(X2)|rearsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
% 0.06/0.47  cnf(c_0_73, negated_conjecture, (frontsegP(esk51_0,esk8_2(esk51_0,esk51_0))|~ssList(app(esk51_0,esk9_2(esk51_0,esk51_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_68])])).
% 0.06/0.47  cnf(c_0_74, plain, (rearsegP(app(X1,X2),X3)|~rearsegP(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
% 0.06/0.47  cnf(c_0_75, plain, (rearsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
% 0.06/0.47  cnf(c_0_76, negated_conjecture, (frontsegP(esk51_0,esk8_2(esk51_0,esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_64]), c_0_67]), c_0_44])])).
% 0.06/0.47  fof(c_0_77, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
% 0.06/0.47  cnf(c_0_78, plain, (rearsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_43])])).
% 0.06/0.47  cnf(c_0_79, negated_conjecture, (app(esk8_2(esk51_0,esk51_0),esk6_2(esk51_0,esk8_2(esk51_0,esk51_0)))=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_76]), c_0_68]), c_0_44])])).
% 0.06/0.47  cnf(c_0_80, negated_conjecture, (ssList(esk6_2(esk51_0,esk8_2(esk51_0,esk51_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_76]), c_0_68]), c_0_44])])).
% 0.06/0.47  cnf(c_0_81, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
% 0.06/0.47  cnf(c_0_82, negated_conjecture, (rearsegP(esk51_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80]), c_0_68])])).
% 0.06/0.47  cnf(c_0_83, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
% 0.06/0.47  cnf(c_0_84, negated_conjecture, (app(esk7_2(esk51_0,nil),nil)=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_43]), c_0_44])])).
% 0.06/0.47  cnf(c_0_85, negated_conjecture, (ssList(esk7_2(esk51_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_82]), c_0_43]), c_0_44])])).
% 0.06/0.47  cnf(c_0_86, negated_conjecture, (esk7_2(esk51_0,nil)=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_84]), c_0_85])])).
% 0.06/0.47  cnf(c_0_87, negated_conjecture, (nil!=esk48_0|nil!=esk49_0), inference(split_conjunct,[status(thm)],[c_0_26])).
% 0.06/0.47  fof(c_0_88, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|hd(app(X3,X4))=hd(X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
% 0.06/0.47  cnf(c_0_89, negated_conjecture, (app(esk51_0,nil)=esk51_0), inference(rw,[status(thm)],[c_0_84, c_0_86])).
% 0.06/0.47  cnf(c_0_90, negated_conjecture, (esk50_0!=nil|esk51_0!=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_28]), c_0_29])).
% 0.06/0.47  fof(c_0_91, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).
% 0.06/0.47  fof(c_0_92, plain, (![X5]:![X6]:![X7]:![X8]:(((((((X5=X6|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5))&(((((frontsegP(X7,X8)|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))&((((((X5!=X6|~frontsegP(X7,X8))|frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])).
% 0.06/0.47  fof(c_0_93, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
% 0.06/0.47  fof(c_0_94, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).
% 0.06/0.47  fof(c_0_95, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
% 0.06/0.47  cnf(c_0_96, plain, (hd(app(X1,X2))=hd(X1)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
% 0.06/0.47  cnf(c_0_97, negated_conjecture, (app(esk51_0,app(nil,X1))=app(esk51_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_89]), c_0_43]), c_0_44])])).
% 0.06/0.47  cnf(c_0_98, negated_conjecture, (esk51_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_38])])).
% 0.06/0.47  cnf(c_0_99, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_91])).
% 0.06/0.47  cnf(c_0_100, plain, (frontsegP(cons(X1,X3),cons(X2,X4))|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~ssList(X4)|~frontsegP(X3,X4)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_92])).
% 0.06/0.47  fof(c_0_101, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])).
% 0.06/0.47  cnf(c_0_102, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_93])).
% 0.06/0.47  cnf(c_0_103, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
% 0.06/0.47  cnf(c_0_104, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
% 0.06/0.47  fof(c_0_105, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).
% 0.06/0.47  cnf(c_0_106, negated_conjecture, (hd(app(esk51_0,X1))=hd(esk51_0)|~ssList(app(nil,X1))|~ssList(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_44])]), c_0_98])).
% 0.06/0.47  cnf(c_0_107, negated_conjecture, (ssList(app(esk51_0,X1))|~ssList(app(nil,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_97]), c_0_44])])).
% 0.06/0.47  cnf(c_0_108, negated_conjecture, (app(esk51_0,X1)!=nil|~ssList(app(nil,X1))|~ssList(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_97]), c_0_44])]), c_0_98])).
% 0.06/0.47  fof(c_0_109, plain, (![X2]:(~ssList(X2)|segmentP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])).
% 0.06/0.47  cnf(c_0_110, plain, (frontsegP(cons(X1,X2),cons(X1,X3))|~frontsegP(X2,X3)|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_100])).
% 0.06/0.47  cnf(c_0_111, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
% 0.06/0.47  cnf(c_0_112, plain, (cons(hd(X1),tl(app(X1,X2)))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_96]), c_0_64]), c_0_99])).
% 0.06/0.47  cnf(c_0_113, plain, (nil=X1|ssList(app(tl(X1),X2))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_64]), c_0_99])).
% 0.06/0.47  cnf(c_0_114, negated_conjecture, (app(tl(esk51_0),nil)=tl(esk51_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_89]), c_0_43]), c_0_44])]), c_0_98])).
% 0.06/0.47  cnf(c_0_115, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_105])).
% 0.06/0.47  cnf(c_0_116, negated_conjecture, (hd(app(esk51_0,esk51_0))=hd(esk51_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_53]), c_0_44])])).
% 0.06/0.47  cnf(c_0_117, negated_conjecture, (ssList(app(esk51_0,esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_53]), c_0_44])])).
% 0.06/0.47  cnf(c_0_118, negated_conjecture, (app(esk51_0,esk51_0)!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_53]), c_0_44])])).
% 0.06/0.47  cnf(c_0_119, negated_conjecture, (~totalorderedP(X1)|~segmentP(X1,esk50_0)|~frontsegP(esk51_0,X1)|~neq(esk50_0,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
% 0.06/0.47  cnf(c_0_120, plain, (segmentP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
% 0.06/0.47  cnf(c_0_121, plain, (frontsegP(cons(X1,X2),cons(X1,nil))|~ssList(X2)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_43])])).
% 0.06/0.47  cnf(c_0_122, negated_conjecture, (cons(hd(esk51_0),tl(esk51_0))=esk51_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_89]), c_0_43]), c_0_44])]), c_0_98])).
% 0.06/0.47  cnf(c_0_123, negated_conjecture, (ssList(tl(esk51_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_43]), c_0_44])]), c_0_98])).
% 0.06/0.47  cnf(c_0_124, negated_conjecture, (ssItem(hd(esk51_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_117])]), c_0_118])).
% 0.06/0.47  cnf(c_0_125, negated_conjecture, (~totalorderedP(X1)|~frontsegP(esk51_0,X1)|~ssList(X1)|~neq(nil,X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_38]), c_0_38]), c_0_120])).
% 0.06/0.47  cnf(c_0_126, negated_conjecture, (frontsegP(esk51_0,cons(hd(esk51_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_123]), c_0_124])])).
% 0.06/0.47  fof(c_0_127, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
% 0.06/0.47  cnf(c_0_128, negated_conjecture, (~totalorderedP(cons(hd(esk51_0),nil))|~ssList(cons(hd(esk51_0),nil))|~neq(nil,cons(hd(esk51_0),nil))), inference(spm,[status(thm)],[c_0_125, c_0_126])).
% 0.06/0.47  cnf(c_0_129, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_127])).
% 0.06/0.47  fof(c_0_130, plain, (![X2]:(~ssItem(X2)|totalorderedP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])).
% 0.06/0.47  cnf(c_0_131, negated_conjecture, (~totalorderedP(cons(hd(esk51_0),nil))|~neq(nil,cons(hd(esk51_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_43]), c_0_124])])).
% 0.06/0.47  cnf(c_0_132, plain, (totalorderedP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_130])).
% 0.06/0.47  fof(c_0_133, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).
% 0.06/0.47  cnf(c_0_134, negated_conjecture, (~neq(nil,cons(hd(esk51_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_124])])).
% 0.06/0.47  cnf(c_0_135, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_133])).
% 0.06/0.47  fof(c_0_136, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
% 0.06/0.47  cnf(c_0_137, negated_conjecture, (cons(hd(esk51_0),nil)=nil|~ssList(cons(hd(esk51_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_43])])).
% 0.06/0.47  cnf(c_0_138, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_136])).
% 0.06/0.47  cnf(c_0_139, negated_conjecture, (cons(hd(esk51_0),nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_129]), c_0_43]), c_0_124])])).
% 0.06/0.47  cnf(c_0_140, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_43]), c_0_124])]), ['proof']).
% 0.06/0.47  # SZS output end CNFRefutation
% 0.06/0.47  # Training examples: 0 positive, 0 negative
% 0.06/0.47  
% 0.06/0.47  # -------------------------------------------------
% 0.06/0.47  # User time                : 0.235 s
% 0.06/0.47  # System time              : 0.008 s
% 0.06/0.47  # Total time               : 0.243 s
% 0.06/0.47  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
