<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC302+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC302+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC302+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC302%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC302+1.p">SWC302+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n144.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:24 EDT 2016

% Result     : Theorem 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of formulae       :   53 ( 343 expanded)
%              Number of clauses        :   34 ( 116 expanded)
%              Number of leaves         :    9 (  85 expanded)
%              Depth                    :   16
%              Number of atoms          :  203 (3061 expanded)
%              Number of equality atoms :   34 ( 586 expanded)
%              Maximal formula depth    :   26 (   5 average)
%              Maximal clause size      :   30 (   2 average)
%              Maximal term depth       :    6 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ~ segmentP(X4,X3)
                    | ~ totalorderedP(X3)
                    | ! [X5] :
                        ( ssItem(X5)
                       => ! [X6] :
                            ( ssItem(X6)
                           => ! [X7] :
                                ( ssList(X7)
                               => ! [X8] :
                                    ( ssList(X8)
                                   => ! [X9] :
                                        ( ssList(X9)
                                       => ( app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9) != X1
                                          | ~ lt(X6,X5) ) ) ) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( totalorderedP(X1)
      &lt;=> ! [X2] :
            ( ssItem(X2)
           => ! [X3] :
                ( ssItem(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ! [X5] :
                        ( ssList(X5)
                       => ! [X6] :
                            ( ssList(X6)
                           => ( app(app(X4,cons(X2,X5)),cons(X3,X6)) = X1
                             => leq(X2,X3) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax11)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( lt(X1,X2)
           => ~ lt(X2,X1) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax33)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( leq(X1,X2)
           => ( X1 = X2
              | lt(X1,X2) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax92)).

<A NAME="c_0_9"></A>fof(c_0_9,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ~ segmentP(X4,X3)
                      | ~ totalorderedP(X3)
                      | ! [X5] :
                          ( ssItem(X5)
                         => ! [X6] :
                              ( ssItem(X6)
                             => ! [X7] :
                                  ( ssList(X7)
                                 => ! [X8] :
                                      ( ssList(X8)
                                     => ! [X9] :
                                          ( ssList(X9)
                                         => ( app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9) != X1
                                            | ~ lt(X6,X5) ) ) ) ) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_10"></A>fof(c_0_10,negated_conjecture,
    ( ssList(esk48_0)
    & ssList(esk49_0)
    & ssList(esk50_0)
    & ssList(esk51_0)
    & esk49_0 = esk51_0
    & esk48_0 = esk50_0
    & segmentP(esk51_0,esk50_0)
    & totalorderedP(esk50_0)
    & ssItem(esk52_0)
    & ssItem(esk53_0)
    & ssList(esk54_0)
    & ssList(esk55_0)
    & ssList(esk56_0)
    & app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk48_0
    & lt(esk53_0,esk52_0) ),
    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_13"></A>fof(c_0_13,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ssList(esk55_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( ssList(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_18"></A>fof(c_0_18,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0) = esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_15]),c_0_16]),c_0_17])])).

<A NAME="c_0_20"></A>cnf(c_0_20,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( ssItem(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_23"></A>fof(c_0_23,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_20]),c_0_21]),c_0_22])])).

<A NAME="c_0_25"></A>cnf(c_0_25,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( app(app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,nil)),esk56_0) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_25]),c_0_16]),c_0_22])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( ssList(esk56_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0
    | ~ ssList(app(esk54_0,cons(esk52_0,esk55_0)))
    | ~ ssList(cons(esk53_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_26]),c_0_27])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( ssItem(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_30"></A>fof(c_0_30,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0
    | ~ ssList(app(esk54_0,cons(esk52_0,esk55_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_20]),c_0_21]),c_0_29])])).

<A NAME="c_0_32"></A>cnf(c_0_32,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_33"></A>fof(c_0_33,plain,(
    ! [X7,X8,X9,X10,X11,X12] :
      ( ( ~ totalorderedP(X7)
        | ~ ssItem(X8)
        | ~ ssItem(X9)
        | ~ ssList(X10)
        | ~ ssList(X11)
        | ~ ssList(X12)
        | app(app(X10,cons(X8,X11)),cons(X9,X12)) != X7
        | leq(X8,X9)
        | ~ ssList(X7) )
      & ( ssItem(esk25_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssItem(esk26_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk27_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk28_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk29_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( app(app(esk27_1(X7),cons(esk25_1(X7),esk28_1(X7))),cons(esk26_1(X7),esk29_1(X7))) = X7
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ~ leq(esk25_1(X7),esk26_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0
    | ~ ssList(cons(esk52_0,esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_32]),c_0_17])])).

<A NAME="c_0_35"></A>fof(c_0_35,plain,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( lt(X1,X2)
           => ~ lt(X2,X1) ) ) ) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( leq(X2,X3)
    | ~ ssList(X1)
    | app(app(X4,cons(X2,X5)),cons(X3,X6)) != X1
    | ~ ssList(X6)
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssItem(X3)
    | ~ ssItem(X2)
    | ~ totalorderedP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_20]),c_0_16]),c_0_22])])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_39"></A>fof(c_0_39,plain,(
    ! [X3,X4] :
      ( ~ ssItem(X3)
      | ~ ssItem(X4)
      | ~ lt(X3,X4)
      | ~ lt(X4,X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])).

<A NAME="c_0_40"></A>fof(c_0_40,plain,(
    ! [X3,X4] :
      ( ~ ssItem(X3)
      | ~ ssItem(X4)
      | ~ leq(X3,X4)
      | X3 = X4
      | lt(X3,X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( leq(X1,X2)
    | ~ totalorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,esk56_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_25]),c_0_27]),c_0_29])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( totalorderedP(esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( ~ lt(X1,X2)
    | ~ lt(X2,X1)
    | ~ ssItem(X1)
    | ~ ssItem(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( lt(esk53_0,esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_47"></A>cnf(c_0_47,plain,
    ( lt(X1,X2)
    | X1 = X2
    | ~ leq(X1,X2)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( leq(esk52_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_42]),c_0_43]),c_0_44]),c_0_27]),c_0_16]),c_0_17]),c_0_29]),c_0_22])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( ~ lt(esk52_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_46]),c_0_29]),c_0_22])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( esk53_0 = esk52_0 ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_48]),c_0_29]),c_0_22])]),<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( ~ lt(esk52_0,esk52_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_50]),<A HREF="#c_0_51">c_0_51</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC302+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n144.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:40:38 CDT 2016
% 0.02/0.22  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.29  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.02/0.29  # and selection function SelectCQPrecWNTNp.
% 0.02/0.29  #
% 0.02/0.29  # Preprocessing time       : 0.014 s
% 0.02/0.29  # Presaturation interreduction done
% 0.02/0.29  
% 0.02/0.29  # Proof found!
% 0.02/0.29  # SZS status Theorem
% 0.02/0.29  # SZS output start CNFRefutation
% 0.02/0.29  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|~(segmentP(X4,X3)))|~(totalorderedP(X3)))|![X5]:(ssItem(X5)=>![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>![X9]:(ssList(X9)=>(app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9)!=X1|~(lt(X6,X5)))))))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.02/0.29  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.02/0.29  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.02/0.29  fof(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.02/0.29  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.02/0.29  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.02/0.29  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>(totalorderedP(X1)&lt;=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(app(X4,cons(X2,X5)),cons(X3,X6))=X1=>leq(X2,X3)))))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax11)).
% 0.02/0.29  fof(c_0_7, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>(lt(X1,X2)=>~(lt(X2,X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax33)).
% 0.02/0.29  fof(c_0_8, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>(leq(X1,X2)=>(X1=X2|lt(X1,X2)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax92)).
% 0.02/0.29  fof(c_0_9, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|~segmentP(X4,X3))|~totalorderedP(X3))|![X5]:(ssItem(X5)=>![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>![X9]:(ssList(X9)=>(app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9)!=X1|~lt(X6,X5)))))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.02/0.29  fof(c_0_10, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&segmentP(esk51_0,esk50_0))&totalorderedP(esk50_0))&(ssItem(esk52_0)&(ssItem(esk53_0)&(ssList(esk54_0)&(ssList(esk55_0)&(ssList(esk56_0)&(app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk48_0&lt(esk53_0,esk52_0))))))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 0.02/0.29  cnf(c_0_11, negated_conjecture, (app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  cnf(c_0_12, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  fof(c_0_13, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.02/0.29  cnf(c_0_14, negated_conjecture, (app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk50_0), inference(rw,[status(thm)],[c_0_11, c_0_12])).
% 0.02/0.29  cnf(c_0_15, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
% 0.02/0.29  cnf(c_0_16, negated_conjecture, (ssList(esk55_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  cnf(c_0_17, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  fof(c_0_18, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.02/0.29  cnf(c_0_19, negated_conjecture, (app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0)=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17])])).
% 0.02/0.29  cnf(c_0_20, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
% 0.02/0.29  cnf(c_0_21, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.02/0.29  cnf(c_0_22, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  fof(c_0_23, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 0.02/0.29  cnf(c_0_24, negated_conjecture, (app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22])])).
% 0.02/0.29  cnf(c_0_25, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 0.02/0.29  cnf(c_0_26, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,nil)),esk56_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_16]), c_0_22])])).
% 0.02/0.29  cnf(c_0_27, negated_conjecture, (ssList(esk56_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  cnf(c_0_28, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0|~ssList(app(esk54_0,cons(esk52_0,esk55_0)))|~ssList(cons(esk53_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_26]), c_0_27])])).
% 0.02/0.29  cnf(c_0_29, negated_conjecture, (ssItem(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  fof(c_0_30, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 0.02/0.29  cnf(c_0_31, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0|~ssList(app(esk54_0,cons(esk52_0,esk55_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_20]), c_0_21]), c_0_29])])).
% 0.02/0.29  cnf(c_0_32, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
% 0.02/0.29  fof(c_0_33, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((~totalorderedP(X7)|(~ssItem(X8)|(~ssItem(X9)|(~ssList(X10)|(~ssList(X11)|(~ssList(X12)|(app(app(X10,cons(X8,X11)),cons(X9,X12))!=X7|leq(X8,X9))))))))|~ssList(X7))&(((ssItem(esk25_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssItem(esk26_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk27_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk28_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk29_1(X7))|totalorderedP(X7))|~ssList(X7))&(((app(app(esk27_1(X7),cons(esk25_1(X7),esk28_1(X7))),cons(esk26_1(X7),esk29_1(X7)))=X7|totalorderedP(X7))|~ssList(X7))&((~leq(esk25_1(X7),esk26_1(X7))|totalorderedP(X7))|~ssList(X7)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
% 0.02/0.29  cnf(c_0_34, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0|~ssList(cons(esk52_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_17])])).
% 0.02/0.29  fof(c_0_35, plain, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>(lt(X1,X2)=>~lt(X2,X1))))), inference(fof_simplification,[status(thm)],[c_0_7])).
% 0.02/0.29  cnf(c_0_36, plain, (leq(X2,X3)|~ssList(X1)|app(app(X4,cons(X2,X5)),cons(X3,X6))!=X1|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssItem(X3)|~ssItem(X2)|~totalorderedP(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
% 0.02/0.29  cnf(c_0_37, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_20]), c_0_16]), c_0_22])])).
% 0.02/0.29  cnf(c_0_38, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  fof(c_0_39, plain, (![X3]:![X4]:(~ssItem(X3)|(~ssItem(X4)|(~lt(X3,X4)|~lt(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])).
% 0.02/0.29  fof(c_0_40, plain, (![X3]:![X4]:(~ssItem(X3)|(~ssItem(X4)|(~leq(X3,X4)|(X3=X4|lt(X3,X4)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
% 0.02/0.29  cnf(c_0_41, plain, (leq(X1,X2)|~totalorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_36])).
% 0.02/0.29  cnf(c_0_42, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,esk56_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_25]), c_0_27]), c_0_29])])).
% 0.02/0.29  cnf(c_0_43, negated_conjecture, (totalorderedP(esk50_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  cnf(c_0_44, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_38, c_0_12])).
% 0.02/0.29  cnf(c_0_45, plain, (~lt(X1,X2)|~lt(X2,X1)|~ssItem(X1)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.02/0.29  cnf(c_0_46, negated_conjecture, (lt(esk53_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.29  cnf(c_0_47, plain, (lt(X1,X2)|X1=X2|~leq(X1,X2)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
% 0.02/0.29  cnf(c_0_48, negated_conjecture, (leq(esk52_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44]), c_0_27]), c_0_16]), c_0_17]), c_0_29]), c_0_22])])).
% 0.02/0.29  cnf(c_0_49, negated_conjecture, (~lt(esk52_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_29]), c_0_22])])).
% 0.02/0.29  cnf(c_0_50, negated_conjecture, (esk53_0=esk52_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_29]), c_0_22])]), c_0_49])).
% 0.02/0.29  cnf(c_0_51, negated_conjecture, (~lt(esk52_0,esk52_0)), inference(rw,[status(thm)],[c_0_49, c_0_50])).
% 0.02/0.29  cnf(c_0_52, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_50]), c_0_51]), ['proof']).
% 0.02/0.29  # SZS output end CNFRefutation
% 0.02/0.29  # Training examples: 0 positive, 0 negative
% 0.02/0.29  
% 0.02/0.29  # -------------------------------------------------
% 0.02/0.29  # User time                : 0.057 s
% 0.02/0.29  # System time              : 0.005 s
% 0.02/0.29  # Total time               : 0.062 s
% 0.02/0.29  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
