<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC090+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC090+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC090+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC090%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC090+1.p">SWC090+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n145.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:37 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :   59 ( 442 expanded)
%              Number of clauses        :   39 ( 160 expanded)
%              Number of leaves         :   10 ( 109 expanded)
%              Depth                    :   11
%              Number of atoms          :  198 (3851 expanded)
%              Number of equality atoms :   51 (1022 expanded)
%              Maximal formula depth    :   20 (   4 average)
%              Maximal clause size      :   18 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ~ neq(X2,nil)
                    | ? [X5] :
                        ( ssList(X5)
                        & neq(X5,nil)
                        & segmentP(X2,X5)
                        & segmentP(X1,X5) )
                    | ! [X6] :
                        ( ssList(X6)
                       => ! [X7] :
                            ( ssList(X7)
                           => ( app(X6,X7) != X4
                              | app(X7,X6) != X3 ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( segmentP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax7)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X3,X2) = app(X1,X2)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax79)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax55)).

<A NAME="c_0_10"></A>fof(c_0_10,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ~ neq(X2,nil)
                      | ? [X5] :
                          ( ssList(X5)
                          & neq(X5,nil)
                          & segmentP(X2,X5)
                          & segmentP(X1,X5) )
                      | ! [X6] :
                          ( ssList(X6)
                         => ! [X7] :
                              ( ssList(X7)
                             => ( app(X6,X7) != X4
                                | app(X7,X6) != X3 ) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_11"></A>fof(c_0_11,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk8_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ssList(esk9_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(app(X9,X6),X10) != X5
        | segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

<A NAME="c_0_12"></A>fof(c_0_12,negated_conjecture,(
    ! [X12] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & neq(esk49_0,nil)
      & ( ~ ssList(X12)
        | ~ neq(X12,nil)
        | ~ segmentP(esk49_0,X12)
        | ~ segmentP(esk48_0,X12) )
      & ssList(esk52_0)
      & ssList(esk53_0)
      & app(esk52_0,esk53_0) = esk51_0
      & app(esk53_0,esk52_0) = esk50_0 ) ),
    inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).

<A NAME="c_0_13"></A>cnf(c_0_13,plain,
    ( segmentP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(app(X3,X2),X4) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_14"></A>fof(c_0_14,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).

<A NAME="c_0_15"></A>fof(c_0_15,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( app(esk52_0,esk53_0) = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_18"></A>fof(c_0_18,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X6,X5) != app(X4,X5)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_19"></A>cnf(c_0_19,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( app(esk53_0,esk52_0) = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( app(esk52_0,esk53_0) = esk49_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( ssList(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( ssList(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_28"></A>fof(c_0_28,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).

<A NAME="c_0_29"></A>cnf(c_0_29,plain,
    ( X1 = X2
    | app(X1,X3) != app(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_30"></A>fof(c_0_30,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( segmentP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_20]),c_0_21])]),<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( app(esk53_0,esk52_0) = esk48_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( segmentP(app(esk49_0,X1),esk53_0)
    | ~ ssList(app(esk49_0,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_25]),c_0_26]),c_0_27])])).

<A NAME="c_0_34"></A>cnf(c_0_34,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( X1 = esk52_0
    | app(X1,esk53_0) != esk49_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_25]),c_0_27]),c_0_26])])).

<A NAME="c_0_37"></A>cnf(c_0_37,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( segmentP(esk48_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_32]),c_0_26]),c_0_27])])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( segmentP(esk49_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_34]),c_0_35]),c_0_21])])).

<A NAME="c_0_41"></A>fof(c_0_41,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( segmentP(app(esk48_0,X1),esk52_0)
    | ~ ssList(app(esk48_0,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_32]),c_0_27]),c_0_26])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( app(X1,X2) = esk52_0
    | app(X1,app(X2,esk53_0)) != esk49_0
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_37]),c_0_27])]),<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( app(esk52_0,app(esk53_0,X1)) = app(esk49_0,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_25]),c_0_27]),c_0_26])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( ~ neq(esk53_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_39]),c_0_40]),c_0_27])])).

<A NAME="c_0_47"></A>cnf(c_0_47,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( segmentP(esk48_0,esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_34]),c_0_43]),c_0_21])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( esk52_0 = esk49_0
    | app(esk49_0,esk53_0) != esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_45]),c_0_25]),c_0_27]),c_0_26])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( esk53_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_47]),c_0_21]),c_0_27])])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( app(esk49_0,nil) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_34]),c_0_25]),c_0_21]),c_0_27])])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( ~ segmentP(esk49_0,esk52_0)
    | ~ neq(esk52_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_48]),c_0_26])])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( esk52_0 = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_50]),c_0_51])])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( neq(esk49_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_55"></A>fof(c_0_55,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( ~ segmentP(esk49_0,esk49_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_53]),c_0_53]),c_0_54])])).

<A NAME="c_0_57"></A>cnf(c_0_57,plain,
    ( segmentP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_57]),c_0_35])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC090+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.22  % Computer   : n145.star.cs.uiowa.edu
% 0.03/0.22  % Model      : x86_64 x86_64
% 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.22  % Memory     : 32218.75MB
% 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.22  % CPULimit   : 300
% 0.03/0.22  % DateTime   : Tue May 10 20:35:53 CDT 2016
% 0.03/0.22  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.44  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.44  # and selection function SelectCQPrecWNTNp.
% 0.06/0.44  #
% 0.06/0.44  # Preprocessing time       : 0.014 s
% 0.06/0.44  # Presaturation interreduction done
% 0.06/0.44  
% 0.06/0.44  # Proof found!
% 0.06/0.44  # SZS status Theorem
% 0.06/0.44  # SZS output start CNFRefutation
% 0.06/0.44  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|~(neq(X2,nil)))|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))|![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>(app(X6,X7)!=X4|app(X7,X6)!=X3))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.06/0.44  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax7)).
% 0.06/0.44  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
% 0.06/0.44  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.06/0.44  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax79)).
% 0.06/0.44  fof(c_0_5, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.44  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
% 0.06/0.44  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.06/0.44  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.06/0.44  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>segmentP(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax55)).
% 0.06/0.44  fof(c_0_10, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|~neq(X2,nil))|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))|![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>(app(X6,X7)!=X4|app(X7,X6)!=X3)))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.44  fof(c_0_11, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
% 0.06/0.44  fof(c_0_12, negated_conjecture, (![X12]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&neq(esk49_0,nil))&(((~ssList(X12)|~neq(X12,nil))|~segmentP(esk49_0,X12))|~segmentP(esk48_0,X12)))&(ssList(esk52_0)&(ssList(esk53_0)&(app(esk52_0,esk53_0)=esk51_0&app(esk53_0,esk52_0)=esk50_0))))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).
% 0.06/0.44  cnf(c_0_13, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.06/0.44  fof(c_0_14, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
% 0.06/0.44  fof(c_0_15, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
% 0.06/0.44  cnf(c_0_16, negated_conjecture, (app(esk52_0,esk53_0)=esk51_0), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.06/0.44  cnf(c_0_17, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.06/0.44  fof(c_0_18, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 0.06/0.44  cnf(c_0_19, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_13])).
% 0.06/0.44  cnf(c_0_20, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
% 0.06/0.44  cnf(c_0_21, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.06/0.44  cnf(c_0_22, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
% 0.06/0.44  cnf(c_0_23, negated_conjecture, (app(esk53_0,esk52_0)=esk50_0), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.06/0.44  cnf(c_0_24, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.06/0.44  cnf(c_0_25, negated_conjecture, (app(esk52_0,esk53_0)=esk49_0), inference(rw,[status(thm)],[c_0_16, c_0_17])).
% 0.06/0.44  cnf(c_0_26, negated_conjecture, (ssList(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.06/0.44  cnf(c_0_27, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.06/0.44  fof(c_0_28, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
% 0.06/0.44  cnf(c_0_29, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
% 0.06/0.44  fof(c_0_30, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
% 0.06/0.44  cnf(c_0_31, plain, (segmentP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])]), c_0_22])).
% 0.06/0.44  cnf(c_0_32, negated_conjecture, (app(esk53_0,esk52_0)=esk48_0), inference(rw,[status(thm)],[c_0_23, c_0_24])).
% 0.06/0.44  cnf(c_0_33, negated_conjecture, (segmentP(app(esk49_0,X1),esk53_0)|~ssList(app(esk49_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_25]), c_0_26]), c_0_27])])).
% 0.06/0.44  cnf(c_0_34, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
% 0.06/0.44  cnf(c_0_35, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.06/0.44  cnf(c_0_36, negated_conjecture, (X1=esk52_0|app(X1,esk53_0)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_25]), c_0_27]), c_0_26])])).
% 0.06/0.44  cnf(c_0_37, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
% 0.06/0.44  cnf(c_0_38, negated_conjecture, (~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.06/0.44  cnf(c_0_39, negated_conjecture, (segmentP(esk48_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_26]), c_0_27])])).
% 0.06/0.44  cnf(c_0_40, negated_conjecture, (segmentP(esk49_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_21])])).
% 0.06/0.44  fof(c_0_41, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
% 0.06/0.44  cnf(c_0_42, negated_conjecture, (segmentP(app(esk48_0,X1),esk52_0)|~ssList(app(esk48_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_32]), c_0_27]), c_0_26])])).
% 0.06/0.44  cnf(c_0_43, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.06/0.44  cnf(c_0_44, negated_conjecture, (app(X1,X2)=esk52_0|app(X1,app(X2,esk53_0))!=esk49_0|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_27])]), c_0_22])).
% 0.06/0.44  cnf(c_0_45, negated_conjecture, (app(esk52_0,app(esk53_0,X1))=app(esk49_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_25]), c_0_27]), c_0_26])])).
% 0.06/0.44  cnf(c_0_46, negated_conjecture, (~neq(esk53_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_27])])).
% 0.06/0.44  cnf(c_0_47, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 0.06/0.44  cnf(c_0_48, negated_conjecture, (segmentP(esk48_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_34]), c_0_43]), c_0_21])])).
% 0.06/0.44  cnf(c_0_49, negated_conjecture, (esk52_0=esk49_0|app(esk49_0,esk53_0)!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_25]), c_0_27]), c_0_26])])).
% 0.06/0.44  cnf(c_0_50, negated_conjecture, (esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_21]), c_0_27])])).
% 0.06/0.44  cnf(c_0_51, negated_conjecture, (app(esk49_0,nil)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_34]), c_0_25]), c_0_21]), c_0_27])])).
% 0.06/0.44  cnf(c_0_52, negated_conjecture, (~segmentP(esk49_0,esk52_0)|~neq(esk52_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_48]), c_0_26])])).
% 0.06/0.44  cnf(c_0_53, negated_conjecture, (esk52_0=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50]), c_0_51])])).
% 0.06/0.44  cnf(c_0_54, negated_conjecture, (neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_12])).
% 0.06/0.44  fof(c_0_55, plain, (![X2]:(~ssList(X2)|segmentP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).
% 0.06/0.44  cnf(c_0_56, negated_conjecture, (~segmentP(esk49_0,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_53]), c_0_53]), c_0_54])])).
% 0.06/0.44  cnf(c_0_57, plain, (segmentP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
% 0.06/0.44  cnf(c_0_58, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_35])]), ['proof']).
% 0.06/0.44  # SZS output end CNFRefutation
% 0.06/0.44  # Training examples: 0 positive, 0 negative
% 0.06/0.44  
% 0.06/0.44  # -------------------------------------------------
% 0.06/0.44  # User time                : 0.207 s
% 0.06/0.44  # System time              : 0.010 s
% 0.06/0.44  # Total time               : 0.217 s
% 0.06/0.44  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
