<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC014+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC014+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC014+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC014%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC014+1.p">SWC014+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n151.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:18 EDT 2016

% Result     : Theorem 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of formulae       :   52 ( 492 expanded)
%              Number of clauses        :   34 ( 176 expanded)
%              Number of leaves         :    9 ( 112 expanded)
%              Depth                    :   11
%              Number of atoms          :  210 (7875 expanded)
%              Number of equality atoms :   73 (2641 expanded)
%              Maximal formula depth    :   22 (   5 average)
%              Maximal clause size      :   40 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ( ( ~ neq(X2,nil)
                        | ? [X5] :
                            ( ssList(X5)
                            & X2 = X5
                            & ? [X6] :
                                ( ssList(X6)
                                & tl(X2) = X6
                                & app(X1,X6) = X5
                                & neq(nil,X2) ) )
                        | ! [X7] :
                            ( ssItem(X7)
                           => ! [X8] :
                                ( ssList(X8)
                               => ( cons(X7,nil) != X3
                                  | app(cons(X7,nil),X8) != X4 ) ) ) )
                      & ( ~ neq(X2,nil)
                        | neq(X4,nil) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) != X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax18)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssList(tl(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_9"></A>fof(c_0_9,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ( ( ~ neq(X2,nil)
                          | ? [X5] :
                              ( ssList(X5)
                              & X2 = X5
                              & ? [X6] :
                                  ( ssList(X6)
                                  & tl(X2) = X6
                                  & app(X1,X6) = X5
                                  & neq(nil,X2) ) )
                          | ! [X7] :
                              ( ssItem(X7)
                             => ! [X8] :
                                  ( ssList(X8)
                                 => ( cons(X7,nil) != X3
                                    | app(cons(X7,nil),X8) != X4 ) ) ) )
                        & ( ~ neq(X2,nil)
                          | neq(X4,nil) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_10"></A>fof(c_0_10,negated_conjecture,(
    ! [X13,X14] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ( neq(esk49_0,nil)
        | neq(esk49_0,nil) )
      & ( ~ neq(esk51_0,nil)
        | neq(esk49_0,nil) )
      & ( neq(esk49_0,nil)
        | ~ ssList(X13)
        | esk49_0 != X13
        | ~ ssList(X14)
        | tl(esk49_0) != X14
        | app(esk48_0,X14) != X13
        | ~ neq(nil,esk49_0) )
      & ( ~ neq(esk51_0,nil)
        | ~ ssList(X13)
        | esk49_0 != X13
        | ~ ssList(X14)
        | tl(esk49_0) != X14
        | app(esk48_0,X14) != X13
        | ~ neq(nil,esk49_0) )
      & ( neq(esk49_0,nil)
        | ssItem(esk52_0) )
      & ( ~ neq(esk51_0,nil)
        | ssItem(esk52_0) )
      & ( neq(esk49_0,nil)
        | ssList(esk53_0) )
      & ( ~ neq(esk51_0,nil)
        | ssList(esk53_0) )
      & ( neq(esk49_0,nil)
        | cons(esk52_0,nil) = esk50_0 )
      & ( ~ neq(esk51_0,nil)
        | cons(esk52_0,nil) = esk50_0 )
      & ( neq(esk49_0,nil)
        | app(cons(esk52_0,nil),esk53_0) = esk51_0 )
      & ( ~ neq(esk51_0,nil)
        | app(cons(esk52_0,nil),esk53_0) = esk51_0 ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( neq(esk49_0,nil)
    | neq(esk49_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( cons(esk52_0,nil) = esk50_0
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( neq(esk49_0,nil) ),
    inference(cn,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_16"></A>fof(c_0_16,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) != X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( ssItem(esk52_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( ~ neq(nil,esk49_0)
    | app(esk48_0,X1) != X2
    | tl(esk49_0) != X1
    | ~ ssList(X1)
    | esk49_0 != X2
    | ~ ssList(X2)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_19"></A>fof(c_0_19,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssList(tl(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).

<A NAME="c_0_20"></A>fof(c_0_20,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( app(cons(esk52_0,nil),esk53_0) = esk51_0
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( cons(esk52_0,nil) = esk48_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_13]),c_0_14]),c_0_15])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( ssList(esk53_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,plain,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( ssItem(esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_14]),c_0_15])])).

<A NAME="c_0_27"></A>fof(c_0_27,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( app(esk48_0,tl(esk49_0)) != esk49_0
    | ~ ssList(app(esk48_0,tl(esk49_0)))
    | ~ ssList(tl(esk49_0))
    | ~ neq(nil,esk49_0) ),
    inference(er,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_14]),c_0_15])])])])).

<A NAME="c_0_29"></A>cnf(c_0_29,plain,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_31"></A>fof(c_0_31,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).

<A NAME="c_0_32"></A>cnf(c_0_32,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( app(esk48_0,esk53_0) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_22]),c_0_14]),c_0_14]),c_0_15])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( ssList(esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_14]),c_0_15])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( esk48_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_22]),c_0_25]),c_0_26])])).

<A NAME="c_0_37"></A>cnf(c_0_37,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( esk49_0 = nil
    | app(esk48_0,tl(esk49_0)) != esk49_0
    | ~ ssList(app(esk48_0,tl(esk49_0)))
    | ~ neq(nil,esk49_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_29]),c_0_30])])).

<A NAME="c_0_39"></A>cnf(c_0_39,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( esk49_0 != nil ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_33]),c_0_34]),c_0_35])]),<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_41"></A>fof(c_0_41,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( app(esk48_0,X1) = cons(esk52_0,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_22]),c_0_26])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( app(esk48_0,tl(esk49_0)) != esk49_0
    | ~ ssList(tl(esk49_0))
    | ~ neq(nil,esk49_0) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_39]),c_0_35])]),<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( cons(esk52_0,esk53_0) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_42]),c_0_34])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( app(esk48_0,tl(esk49_0)) != esk49_0
    | ~ neq(nil,esk49_0) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_29]),c_0_30])]),<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( tl(esk49_0) = esk53_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_45]),c_0_34]),c_0_26])])).

<A NAME="c_0_48"></A>fof(c_0_48,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( ~ neq(nil,esk49_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_47]),c_0_33])])).

<A NAME="c_0_50"></A>cnf(c_0_50,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_50]),c_0_30]),c_0_25])]),<A HREF="#c_0_40">c_0_40</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC014+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n151.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:34:38 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.26  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.03/0.26  # and selection function SelectCQPrecWNTNp.
% 0.03/0.26  #
% 0.03/0.26  # Preprocessing time       : 0.014 s
% 0.03/0.26  # Presaturation interreduction done
% 0.03/0.26  
% 0.03/0.26  # Proof found!
% 0.03/0.26  # SZS status Theorem
% 0.03/0.26  # SZS output start CNFRefutation
% 0.03/0.26  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((~(neq(X2,nil))|?[X5]:((ssList(X5)&X2=X5)&?[X6]:(((ssList(X6)&tl(X2)=X6)&app(X1,X6)=X5)&neq(nil,X2))))|![X7]:(ssItem(X7)=>![X8]:(ssList(X8)=>(cons(X7,nil)!=X3|app(cons(X7,nil),X8)!=X4))))&(~(neq(X2,nil))|neq(X4,nil))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.03/0.26  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax18)).
% 0.03/0.26  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
% 0.03/0.26  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.03/0.26  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.03/0.26  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.03/0.26  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.03/0.26  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
% 0.03/0.26  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.03/0.26  fof(c_0_9, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((~neq(X2,nil)|?[X5]:((ssList(X5)&X2=X5)&?[X6]:(((ssList(X6)&tl(X2)=X6)&app(X1,X6)=X5)&neq(nil,X2))))|![X7]:(ssItem(X7)=>![X8]:(ssList(X8)=>(cons(X7,nil)!=X3|app(cons(X7,nil),X8)!=X4))))&(~neq(X2,nil)|neq(X4,nil)))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.03/0.26  fof(c_0_10, negated_conjecture, (![X13]:![X14]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((esk49_0=esk51_0&esk48_0=esk50_0)&((((neq(esk49_0,nil)|neq(esk49_0,nil))&(~neq(esk51_0,nil)|neq(esk49_0,nil)))&((neq(esk49_0,nil)|((~ssList(X13)|esk49_0!=X13)|(((~ssList(X14)|tl(esk49_0)!=X14)|app(esk48_0,X14)!=X13)|~neq(nil,esk49_0))))&(~neq(esk51_0,nil)|((~ssList(X13)|esk49_0!=X13)|(((~ssList(X14)|tl(esk49_0)!=X14)|app(esk48_0,X14)!=X13)|~neq(nil,esk49_0))))))&(((neq(esk49_0,nil)|ssItem(esk52_0))&(~neq(esk51_0,nil)|ssItem(esk52_0)))&(((neq(esk49_0,nil)|ssList(esk53_0))&(~neq(esk51_0,nil)|ssList(esk53_0)))&(((neq(esk49_0,nil)|cons(esk52_0,nil)=esk50_0)&(~neq(esk51_0,nil)|cons(esk52_0,nil)=esk50_0))&((neq(esk49_0,nil)|app(cons(esk52_0,nil),esk53_0)=esk51_0)&(~neq(esk51_0,nil)|app(cons(esk52_0,nil),esk53_0)=esk51_0)))))))))))), 
% inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])])).
% 0.03/0.26  cnf(c_0_11, negated_conjecture, (neq(esk49_0,nil)|neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  cnf(c_0_12, negated_conjecture, (cons(esk52_0,nil)=esk50_0|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  cnf(c_0_13, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  cnf(c_0_14, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  cnf(c_0_15, negated_conjecture, (neq(esk49_0,nil)), inference(cn,[status(thm)],[c_0_11])).
% 0.03/0.26  fof(c_0_16, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.03/0.26  cnf(c_0_17, negated_conjecture, (ssItem(esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  cnf(c_0_18, negated_conjecture, (~neq(nil,esk49_0)|app(esk48_0,X1)!=X2|tl(esk49_0)!=X1|~ssList(X1)|esk49_0!=X2|~ssList(X2)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  fof(c_0_19, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
% 0.03/0.26  fof(c_0_20, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
% 0.03/0.26  cnf(c_0_21, negated_conjecture, (app(cons(esk52_0,nil),esk53_0)=esk51_0|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  cnf(c_0_22, negated_conjecture, (cons(esk52_0,nil)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15])])).
% 0.03/0.26  cnf(c_0_23, negated_conjecture, (ssList(esk53_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  cnf(c_0_24, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.03/0.26  cnf(c_0_25, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.03/0.26  cnf(c_0_26, negated_conjecture, (ssItem(esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_14]), c_0_15])])).
% 0.03/0.26  fof(c_0_27, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 0.03/0.26  cnf(c_0_28, negated_conjecture, (app(esk48_0,tl(esk49_0))!=esk49_0|~ssList(app(esk48_0,tl(esk49_0)))|~ssList(tl(esk49_0))|~neq(nil,esk49_0)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_14]), c_0_15])])])])).
% 0.03/0.26  cnf(c_0_29, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.03/0.26  cnf(c_0_30, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  fof(c_0_31, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
% 0.03/0.26  cnf(c_0_32, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
% 0.03/0.26  cnf(c_0_33, negated_conjecture, (app(esk48_0,esk53_0)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22]), c_0_14]), c_0_14]), c_0_15])])).
% 0.03/0.26  cnf(c_0_34, negated_conjecture, (ssList(esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_14]), c_0_15])])).
% 0.03/0.26  cnf(c_0_35, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.03/0.26  cnf(c_0_36, negated_conjecture, (esk48_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_22]), c_0_25]), c_0_26])])).
% 0.03/0.26  cnf(c_0_37, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
% 0.03/0.26  cnf(c_0_38, negated_conjecture, (esk49_0=nil|app(esk48_0,tl(esk49_0))!=esk49_0|~ssList(app(esk48_0,tl(esk49_0)))|~neq(nil,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
% 0.03/0.26  cnf(c_0_39, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
% 0.03/0.26  cnf(c_0_40, negated_conjecture, (esk49_0!=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35])]), c_0_36])).
% 0.03/0.26  fof(c_0_41, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
% 0.03/0.26  cnf(c_0_42, negated_conjecture, (app(esk48_0,X1)=cons(esk52_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_22]), c_0_26])])).
% 0.03/0.26  cnf(c_0_43, negated_conjecture, (app(esk48_0,tl(esk49_0))!=esk49_0|~ssList(tl(esk49_0))|~neq(nil,esk49_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_35])]), c_0_40])).
% 0.03/0.26  cnf(c_0_44, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 0.03/0.26  cnf(c_0_45, negated_conjecture, (cons(esk52_0,esk53_0)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_42]), c_0_34])])).
% 0.03/0.26  cnf(c_0_46, negated_conjecture, (app(esk48_0,tl(esk49_0))!=esk49_0|~neq(nil,esk49_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_29]), c_0_30])]), c_0_40])).
% 0.03/0.26  cnf(c_0_47, negated_conjecture, (tl(esk49_0)=esk53_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_34]), c_0_26])])).
% 0.03/0.26  fof(c_0_48, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
% 0.03/0.26  cnf(c_0_49, negated_conjecture, (~neq(nil,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_33])])).
% 0.03/0.26  cnf(c_0_50, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
% 0.03/0.26  cnf(c_0_51, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_30]), c_0_25])]), c_0_40]), ['proof']).
% 0.03/0.26  # SZS output end CNFRefutation
% 0.03/0.26  # Training examples: 0 positive, 0 negative
% 0.03/0.26  
% 0.03/0.26  # -------------------------------------------------
% 0.03/0.26  # User time                : 0.028 s
% 0.03/0.26  # System time              : 0.005 s
% 0.03/0.26  # Total time               : 0.033 s
% 0.03/0.26  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
