<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC402-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC402-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC402-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC402-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC402-1.p">SWC402-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n129.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:48 EDT 2016

% Result     : Unsatisfiable 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of clauses        :   96 ( 539 expanded)
%              Number of leaves         :   37 ( 355 expanded)
%              Depth                    :   20
%              Number of atoms          :  213 ( 934 expanded)
%              Number of equality atoms :   50 ( 225 expanded)
%              Maximal clause size      :    7 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,axiom,
    ( rearsegP(X3,X2)
    | app(X1,X2) != X3
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause143)).

<A NAME="c_0_1"></A>cnf(c_0_1,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( memberP(sk1,sk7) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_14)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( app(app(sk5,sk3),sk6) = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ssList(sk6) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
    | ~ memberP(X1,X2)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause169)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ssItem(sk7) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( ssList(sk5) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( frontsegP(X3,X1)
    | app(X1,X2) != X3
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause144)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( ssList(skaf43(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause52)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( X1 = X3
    | app(X1,X2) != app(X3,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause151)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( app(skaf46(X1,X2),X2) = X1
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause131)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( ssList(skaf46(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause50)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( X2 = X3
    | app(X1,X2) != app(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause150)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_23"></A>cnf(c_0_23,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_24"></A>cnf(c_0_24,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_25"></A>cnf(c_0_25,axiom,
    ( rearsegP(app(X3,X1),X2)
    | ~ rearsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause136)).

<A NAME="c_0_26"></A>cnf(c_0_26,axiom,
    ( rearsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause58)).

<A NAME="c_0_27"></A>cnf(c_0_27,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_28"></A>cnf(c_0_28,axiom,
    ( frontsegP(X1,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause61)).

<A NAME="c_0_29"></A>cnf(c_0_29,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_30"></A>cnf(c_0_30,axiom,
    ( frontsegP(app(X1,X3),X2)
    | ~ frontsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause137)).

<A NAME="c_0_31"></A>cnf(c_0_31,axiom,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause60)).

<A NAME="c_0_32"></A>cnf(c_0_32,axiom,
    ( app(X1,X2) = nil
    | nil != X1
    | nil != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause145)).

<A NAME="c_0_33"></A>cnf(c_0_33,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_34"></A>cnf(c_0_34,axiom,
    ( memberP(app(X1,X3),X2)
    | ~ memberP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause140)).

<A NAME="c_0_35"></A>cnf(c_0_35,axiom,
    ( memberP(app(X3,X1),X2)
    | ~ memberP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssItem(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause141)).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( ~ memberP(sk2,sk7) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_15)).

<A NAME="c_0_37"></A>cnf(c_0_37,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_0]),<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( memberP(sk3,sk7) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( rearsegP(sk4,sk6)
    | ~ ssList(app(sk5,sk3)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_5]),c_0_6])])).

<A NAME="c_0_41"></A>cnf(c_0_41,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( app(skaf42(sk3,sk7),cons(sk7,skaf43(sk7,sk3))) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_38]),c_0_9]),c_0_39])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( rearsegP(sk4,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_1]),c_0_10]),c_0_39])])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]),<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( app(sk5,app(sk3,sk6)) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,c_0_14]),c_0_10]),c_0_39]),c_0_6])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( ~ ssList(app(sk3,cons(sk7,X1)))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(sk3,cons(sk7,X1))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_42]),c_0_9]),c_0_15]),c_0_16])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( X1 = app(sk5,sk3)
    | app(X1,sk6) != sk4
    | ~ ssList(app(sk5,sk3))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_5]),c_0_6])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( app(skaf46(sk4,sk6),sk6) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_43]),c_0_6]),c_0_44])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( frontsegP(sk4,sk5)
    | ~ ssList(app(sk3,sk6)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_46]),c_0_10])])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(app(sk3,cons(sk7,X2)))
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( skaf46(sk4,sk6) = app(sk5,sk3)
    | ~ ssList(app(sk5,sk3)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_49]),c_0_20])])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( frontsegP(sk4,sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_1]),c_0_39]),c_0_6])])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(cons(sk7,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,c_0_1]),c_0_39])])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( skaf46(sk4,sk6) = app(sk5,sk3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_1]),c_0_10]),c_0_39])])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( X1 = app(sk3,sk6)
    | app(sk5,X1) != sk4
    | ~ ssList(app(sk3,sk6))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_46]),c_0_10])])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( app(sk5,skaf45(sk4,sk5)) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_53]),c_0_10]),c_0_44])])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_23]),c_0_9])])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( ssList(app(sk5,sk3)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( skaf45(sk4,sk5) = app(sk3,sk6)
    | ~ ssList(app(sk3,sk6)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_57]),c_0_24])])).

<A NAME="c_0_61"></A>cnf(c_0_61,plain,
    ( rearsegP(app(X1,X2),nil)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26]),c_0_27])])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,<A HREF="#c_0_59">c_0_59</A>])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( skaf45(sk4,sk5) = app(sk3,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>,c_0_1]),c_0_39]),c_0_6])])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( rearsegP(sk3,nil)
    | ~ ssList(cons(sk7,skaf43(sk7,sk3))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_42]),c_0_16])])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( ssList(cons(X1,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_62])])).

<A NAME="c_0_66"></A>cnf(c_0_66,negated_conjecture,
    ( app(sk5,app(app(sk3,sk6),X1)) = app(sk4,X1)
    | ~ ssList(app(sk3,sk6))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_46]),c_0_10])])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( ssList(app(sk3,sk6)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( rearsegP(sk3,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,c_0_65]),c_0_15])])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( app(skaf46(X1,nil),nil) = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_26]),c_0_27])])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( app(sk5,app(app(sk3,sk6),X1)) = app(sk4,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>,c_0_67])])).

<A NAME="c_0_71"></A>cnf(c_0_71,plain,
    ( app(X1,skaf45(X1,X1)) = X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_72"></A>cnf(c_0_72,negated_conjecture,
    ( app(skaf46(sk3,nil),nil) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_68]),c_0_27]),c_0_39])])).

<A NAME="c_0_73"></A>cnf(c_0_73,plain,
    ( skaf46(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_69]),c_0_20])])).

<A NAME="c_0_74"></A>cnf(c_0_74,plain,
    ( X1 = nil
    | app(X2,X1) != X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_29]),c_0_27])])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( app(sk4,skaf45(app(sk3,sk6),app(sk3,sk6))) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>,c_0_71]),c_0_46]),c_0_24]),c_0_67])])).

<A NAME="c_0_76"></A>cnf(c_0_76,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31]),c_0_27])])).

<A NAME="c_0_77"></A>cnf(c_0_77,negated_conjecture,
    ( app(sk3,nil) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_73]),c_0_39])])).

<A NAME="c_0_78"></A>cnf(c_0_78,negated_conjecture,
    ( skaf45(app(sk3,sk6),app(sk3,sk6)) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,c_0_75]),c_0_24]),c_0_44])])).

<A NAME="c_0_79"></A>cnf(c_0_79,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_32])]),c_0_27])])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( frontsegP(sk3,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>,c_0_77]),c_0_27]),c_0_39])])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( app(nil,app(X1,X2)) = app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_33]),c_0_27])])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( app(app(sk3,sk6),nil) = app(sk3,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>,c_0_78]),c_0_67])])).

<A NAME="c_0_83"></A>cnf(c_0_83,plain,
    ( app(nil,app(nil,X1)) = app(nil,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_79]),c_0_27])])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( app(nil,skaf45(sk3,nil)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_80]),c_0_27]),c_0_39])])).

<A NAME="c_0_85"></A>cnf(c_0_85,plain,
    ( frontsegP(app(X1,app(X2,X3)),app(X1,X2))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_14]),<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( app(nil,app(sk3,sk6)) = app(sk3,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_82]),c_0_67]),c_0_27])])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( app(nil,sk3) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_84]),c_0_24])])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( frontsegP(app(sk3,sk6),sk3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_86]),c_0_87]),c_0_6]),c_0_27]),c_0_39])])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( memberP(app(sk3,X1),sk7)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_38]),c_0_9]),c_0_39])])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( app(sk3,skaf45(app(sk3,sk6),sk3)) = app(sk3,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_88]),c_0_39]),c_0_67])])).

<A NAME="c_0_91"></A>cnf(c_0_91,plain,
    ( memberP(app(X1,X2),X3)
    | ~ memberP(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_62])])).

<A NAME="c_0_92"></A>cnf(c_0_92,negated_conjecture,
    ( memberP(app(sk3,sk6),sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>,c_0_90]),c_0_24])])).

<A NAME="c_0_93"></A>cnf(c_0_93,negated_conjecture,
    ( memberP(app(X1,app(sk3,sk6)),sk7)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_92]),c_0_67])])).

<A NAME="c_0_94"></A>cnf(c_0_94,negated_conjecture,
    ( ~ memberP(sk4,sk7) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_95"></A>cnf(c_0_95,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_46]),c_0_10])]),<A HREF="#c_0_94">c_0_94</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC402-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n129.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:43:53 CDT 2016
% 0.02/0.22  % CPUTime    : 
% 0.02/0.22  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.45  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.06/0.45  # and selection function SelectCQArNTEqFirst.
% 0.06/0.45  #
% 0.06/0.45  # Preprocessing time       : 0.009 s
% 0.06/0.45  # Presaturation interreduction done
% 0.06/0.45  
% 0.06/0.45  # Proof found!
% 0.06/0.45  # SZS status Unsatisfiable
% 0.06/0.45  # SZS output start CNFRefutation
% 0.06/0.45  cnf(c_0_0, axiom, (rearsegP(X3,X2)|app(X1,X2)!=X3|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause143)).
% 0.06/0.45  cnf(c_0_1, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.06/0.45  cnf(c_0_2, negated_conjecture, (memberP(sk1,sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_14)).
% 0.06/0.45  cnf(c_0_3, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.06/0.45  cnf(c_0_4, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.06/0.45  cnf(c_0_5, negated_conjecture, (app(app(sk5,sk3),sk6)=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
% 0.06/0.45  cnf(c_0_6, negated_conjecture, (ssList(sk6)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
% 0.06/0.45  cnf(c_0_7, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
% 0.06/0.45  cnf(c_0_8, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause169)).
% 0.06/0.45  cnf(c_0_9, negated_conjecture, (ssItem(sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
% 0.06/0.45  cnf(c_0_10, negated_conjecture, (ssList(sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 0.06/0.45  cnf(c_0_11, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 0.06/0.45  cnf(c_0_12, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 0.06/0.45  cnf(c_0_13, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause144)).
% 0.06/0.45  cnf(c_0_14, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
% 0.06/0.45  cnf(c_0_15, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause52)).
% 0.06/0.45  cnf(c_0_16, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause53)).
% 0.06/0.45  cnf(c_0_17, axiom, (X1=X3|app(X1,X2)!=app(X3,X2)|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause151)).
% 0.06/0.45  cnf(c_0_18, axiom, (app(skaf46(X1,X2),X2)=X1|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause131)).
% 0.06/0.45  cnf(c_0_19, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
% 0.06/0.45  cnf(c_0_20, axiom, (ssList(skaf46(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause50)).
% 0.06/0.45  cnf(c_0_21, axiom, (X2=X3|app(X1,X2)!=app(X1,X3)|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause150)).
% 0.06/0.45  cnf(c_0_22, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause132)).
% 0.06/0.45  cnf(c_0_23, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.06/0.45  cnf(c_0_24, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause51)).
% 0.06/0.45  cnf(c_0_25, axiom, (rearsegP(app(X3,X1),X2)|~rearsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause136)).
% 0.06/0.45  cnf(c_0_26, axiom, (rearsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause58)).
% 0.06/0.45  cnf(c_0_27, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.06/0.45  cnf(c_0_28, axiom, (frontsegP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause61)).
% 0.06/0.45  cnf(c_0_29, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause73)).
% 0.06/0.45  cnf(c_0_30, axiom, (frontsegP(app(X1,X3),X2)|~frontsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause137)).
% 0.06/0.45  cnf(c_0_31, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause60)).
% 0.06/0.45  cnf(c_0_32, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause145)).
% 0.06/0.45  cnf(c_0_33, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause74)).
% 0.06/0.45  cnf(c_0_34, axiom, (memberP(app(X1,X3),X2)|~memberP(X1,X2)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause140)).
% 0.06/0.45  cnf(c_0_35, axiom, (memberP(app(X3,X1),X2)|~memberP(X1,X2)|~ssList(X1)|~ssList(X3)|~ssItem(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause141)).
% 0.06/0.45  cnf(c_0_36, negated_conjecture, (~memberP(sk2,sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_15)).
% 0.06/0.45  cnf(c_0_37, plain, (rearsegP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_0]), c_0_1])).
% 0.06/0.45  cnf(c_0_38, negated_conjecture, (memberP(sk3,sk7)), inference(rw,[status(thm)],[c_0_2, c_0_3])).
% 0.06/0.45  cnf(c_0_39, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_4, c_0_3])).
% 0.06/0.45  cnf(c_0_40, negated_conjecture, (rearsegP(sk4,sk6)|~ssList(app(sk5,sk3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_5]), c_0_6])])).
% 0.06/0.45  cnf(c_0_41, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_7])).
% 0.06/0.45  cnf(c_0_42, negated_conjecture, (app(skaf42(sk3,sk7),cons(sk7,skaf43(sk7,sk3)))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_38]), c_0_9]), c_0_39])])).
% 0.06/0.45  cnf(c_0_43, negated_conjecture, (rearsegP(sk4,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1]), c_0_10]), c_0_39])])).
% 0.06/0.45  cnf(c_0_44, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_11, c_0_12])).
% 0.06/0.45  cnf(c_0_45, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_1])).
% 0.06/0.45  cnf(c_0_46, negated_conjecture, (app(sk5,app(sk3,sk6))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_14]), c_0_10]), c_0_39]), c_0_6])])).
% 0.06/0.45  cnf(c_0_47, negated_conjecture, (~ssList(app(sk3,cons(sk7,X1)))|~ssList(X1)|~duplicatefreeP(app(sk3,cons(sk7,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_9]), c_0_15]), c_0_16])])).
% 0.06/0.45  cnf(c_0_48, negated_conjecture, (X1=app(sk5,sk3)|app(X1,sk6)!=sk4|~ssList(app(sk5,sk3))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_5]), c_0_6])])).
% 0.06/0.45  cnf(c_0_49, negated_conjecture, (app(skaf46(sk4,sk6),sk6)=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_43]), c_0_6]), c_0_44])])).
% 0.06/0.45  cnf(c_0_50, negated_conjecture, (frontsegP(sk4,sk5)|~ssList(app(sk3,sk6))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_10])])).
% 0.06/0.45  cnf(c_0_51, negated_conjecture, (ssItem(X1)|~ssList(app(sk3,cons(sk7,X2)))|~ssList(X2)), inference(spm,[status(thm)],[c_0_47, c_0_19])).
% 0.06/0.45  cnf(c_0_52, negated_conjecture, (skaf46(sk4,sk6)=app(sk5,sk3)|~ssList(app(sk5,sk3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_20])])).
% 0.06/0.45  cnf(c_0_53, negated_conjecture, (frontsegP(sk4,sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1]), c_0_39]), c_0_6])])).
% 0.06/0.45  cnf(c_0_54, negated_conjecture, (ssItem(X1)|~ssList(cons(sk7,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_1]), c_0_39])])).
% 0.06/0.45  cnf(c_0_55, negated_conjecture, (skaf46(sk4,sk6)=app(sk5,sk3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1]), c_0_10]), c_0_39])])).
% 0.06/0.45  cnf(c_0_56, negated_conjecture, (X1=app(sk3,sk6)|app(sk5,X1)!=sk4|~ssList(app(sk3,sk6))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_46]), c_0_10])])).
% 0.06/0.45  cnf(c_0_57, negated_conjecture, (app(sk5,skaf45(sk4,sk5))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_53]), c_0_10]), c_0_44])])).
% 0.06/0.45  cnf(c_0_58, negated_conjecture, (ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_23]), c_0_9])])).
% 0.06/0.45  cnf(c_0_59, negated_conjecture, (ssList(app(sk5,sk3))), inference(spm,[status(thm)],[c_0_20, c_0_55])).
% 0.06/0.45  cnf(c_0_60, negated_conjecture, (skaf45(sk4,sk5)=app(sk3,sk6)|~ssList(app(sk3,sk6))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_24])])).
% 0.06/0.45  cnf(c_0_61, plain, (rearsegP(app(X1,X2),nil)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
% 0.06/0.45  cnf(c_0_62, negated_conjecture, (ssItem(X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
% 0.06/0.45  cnf(c_0_63, negated_conjecture, (skaf45(sk4,sk5)=app(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_1]), c_0_39]), c_0_6])])).
% 0.06/0.45  cnf(c_0_64, negated_conjecture, (rearsegP(sk3,nil)|~ssList(cons(sk7,skaf43(sk7,sk3)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_42]), c_0_16])])).
% 0.06/0.45  cnf(c_0_65, plain, (ssList(cons(X1,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_62])])).
% 0.06/0.45  cnf(c_0_66, negated_conjecture, (app(sk5,app(app(sk3,sk6),X1))=app(sk4,X1)|~ssList(app(sk3,sk6))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_46]), c_0_10])])).
% 0.06/0.45  cnf(c_0_67, negated_conjecture, (ssList(app(sk3,sk6))), inference(spm,[status(thm)],[c_0_24, c_0_63])).
% 0.06/0.45  cnf(c_0_68, negated_conjecture, (rearsegP(sk3,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_15])])).
% 0.06/0.45  cnf(c_0_69, plain, (app(skaf46(X1,nil),nil)=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_26]), c_0_27])])).
% 0.06/0.45  cnf(c_0_70, negated_conjecture, (app(sk5,app(app(sk3,sk6),X1))=app(sk4,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_67])])).
% 0.06/0.45  cnf(c_0_71, plain, (app(X1,skaf45(X1,X1))=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_22, c_0_28])).
% 0.06/0.45  cnf(c_0_72, negated_conjecture, (app(skaf46(sk3,nil),nil)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_68]), c_0_27]), c_0_39])])).
% 0.06/0.45  cnf(c_0_73, plain, (skaf46(X1,nil)=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_69]), c_0_20])])).
% 0.06/0.45  cnf(c_0_74, plain, (X1=nil|app(X2,X1)!=X2|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_29]), c_0_27])])).
% 0.06/0.45  cnf(c_0_75, negated_conjecture, (app(sk4,skaf45(app(sk3,sk6),app(sk3,sk6)))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_46]), c_0_24]), c_0_67])])).
% 0.06/0.45  cnf(c_0_76, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_27])])).
% 0.06/0.45  cnf(c_0_77, negated_conjecture, (app(sk3,nil)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_39])])).
% 0.06/0.45  cnf(c_0_78, negated_conjecture, (skaf45(app(sk3,sk6),app(sk3,sk6))=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_24]), c_0_44])])).
% 0.06/0.45  cnf(c_0_79, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_32])]), c_0_27])])).
% 0.06/0.45  cnf(c_0_80, negated_conjecture, (frontsegP(sk3,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_27]), c_0_39])])).
% 0.06/0.45  cnf(c_0_81, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_33]), c_0_27])])).
% 0.06/0.45  cnf(c_0_82, negated_conjecture, (app(app(sk3,sk6),nil)=app(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_78]), c_0_67])])).
% 0.06/0.45  cnf(c_0_83, plain, (app(nil,app(nil,X1))=app(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_79]), c_0_27])])).
% 0.06/0.45  cnf(c_0_84, negated_conjecture, (app(nil,skaf45(sk3,nil))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_80]), c_0_27]), c_0_39])])).
% 0.06/0.45  cnf(c_0_85, plain, (frontsegP(app(X1,app(X2,X3)),app(X1,X2))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_14]), c_0_1])).
% 0.06/0.45  cnf(c_0_86, negated_conjecture, (app(nil,app(sk3,sk6))=app(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_67]), c_0_27])])).
% 0.06/0.45  cnf(c_0_87, negated_conjecture, (app(nil,sk3)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_24])])).
% 0.06/0.45  cnf(c_0_88, negated_conjecture, (frontsegP(app(sk3,sk6),sk3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87]), c_0_6]), c_0_27]), c_0_39])])).
% 0.06/0.45  cnf(c_0_89, negated_conjecture, (memberP(app(sk3,X1),sk7)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_38]), c_0_9]), c_0_39])])).
% 0.06/0.45  cnf(c_0_90, negated_conjecture, (app(sk3,skaf45(app(sk3,sk6),sk3))=app(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_88]), c_0_39]), c_0_67])])).
% 0.06/0.45  cnf(c_0_91, plain, (memberP(app(X1,X2),X3)|~memberP(X2,X3)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_62])])).
% 0.06/0.45  cnf(c_0_92, negated_conjecture, (memberP(app(sk3,sk6),sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_24])])).
% 0.06/0.45  cnf(c_0_93, negated_conjecture, (memberP(app(X1,app(sk3,sk6)),sk7)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_67])])).
% 0.06/0.45  cnf(c_0_94, negated_conjecture, (~memberP(sk4,sk7)), inference(rw,[status(thm)],[c_0_36, c_0_12])).
% 0.06/0.45  cnf(c_0_95, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_46]), c_0_10])]), c_0_94]), ['proof']).
% 0.06/0.45  # SZS output end CNFRefutation
% 0.06/0.45  # Training examples: 0 positive, 0 negative
% 0.06/0.45  
% 0.06/0.45  # -------------------------------------------------
% 0.06/0.45  # User time                : 0.220 s
% 0.06/0.45  # System time              : 0.006 s
% 0.06/0.45  # Total time               : 0.226 s
% 0.06/0.45  # Maximum resident set size: 1864 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
