<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC378-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC378-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC378-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC378-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC378-1.p">SWC378-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n114.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:42 EDT 2016

% Result     : Unsatisfiable 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of clauses        :   32 ( 135 expanded)
%              Number of leaves         :   15 (  93 expanded)
%              Depth                    :   12
%              Number of atoms          :   60 ( 228 expanded)
%              Number of equality atoms :   11 (  48 expanded)
%              Maximal clause size      :    6 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( app(sk6,sk5) = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( ssList(sk6) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssList(sk5) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( app(sk5,sk6) = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( memberP(app(X3,X1),X2)
    | ~ memberP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssItem(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause141)).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( memberP(sk2,sk7)
    | memberP(sk1,sk7) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ssItem(sk7) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( memberP(X2,X3)
    | memberP(X1,X3)
    | ~ memberP(app(X1,X2),X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssItem(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause158)).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( memberP(app(X1,X3),X2)
    | ~ memberP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause140)).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( ~ memberP(sk2,sk7)
    | ~ memberP(sk1,sk7) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( app(sk6,sk5) = sk1 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( app(sk6,app(sk5,X1)) = app(sk1,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_15]),c_0_3]),c_0_4])])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( app(sk5,sk6) = sk2 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( memberP(app(X1,sk2),sk7)
    | memberP(sk1,sk7)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_8]),c_0_9]),c_0_10])])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( app(sk6,sk2) = app(sk1,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_17]),c_0_3])])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( memberP(app(sk1,sk6),sk7)
    | memberP(sk1,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_19]),c_0_3])])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( memberP(sk1,sk7)
    | memberP(sk6,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_20]),c_0_9]),c_0_3]),c_0_12])])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( memberP(app(sk6,X1),sk7)
    | memberP(sk1,sk7)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_21]),c_0_9]),c_0_3])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( memberP(sk1,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_15]),c_0_4])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( app(sk5,app(sk6,X1)) = app(sk2,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_17]),c_0_4]),c_0_3])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( memberP(app(X1,sk1),sk7)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_23]),c_0_9]),c_0_12])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( app(sk5,sk1) = app(sk2,sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_15]),c_0_4])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( memberP(app(sk2,sk5),sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26]),c_0_4])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( ~ memberP(sk2,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_23])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( memberP(sk5,sk7) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_27]),c_0_9]),c_0_4]),c_0_10])]),<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( memberP(app(sk5,X1),sk7)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_29]),c_0_9]),c_0_4])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_17]),c_0_3])]),<A HREF="#c_0_28">c_0_28</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC378-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n114.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:42:53 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.28  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.02/0.28  # and selection function SelectCQArNTEqFirst.
% 0.02/0.28  #
% 0.02/0.28  # Preprocessing time       : 0.009 s
% 0.02/0.28  # Presaturation interreduction done
% 0.02/0.28  
% 0.02/0.28  # Proof found!
% 0.02/0.28  # SZS status Unsatisfiable
% 0.02/0.28  # SZS output start CNFRefutation
% 0.02/0.28  cnf(c_0_0, negated_conjecture, (app(sk6,sk5)=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
% 0.02/0.28  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 0.02/0.28  cnf(c_0_2, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause149)).
% 0.02/0.28  cnf(c_0_3, negated_conjecture, (ssList(sk6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
% 0.02/0.28  cnf(c_0_4, negated_conjecture, (ssList(sk5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 0.02/0.28  cnf(c_0_5, negated_conjecture, (app(sk5,sk6)=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
% 0.02/0.28  cnf(c_0_6, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 0.02/0.28  cnf(c_0_7, axiom, (memberP(app(X3,X1),X2)|~memberP(X1,X2)|~ssList(X1)|~ssList(X3)|~ssItem(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause141)).
% 0.02/0.28  cnf(c_0_8, negated_conjecture, (memberP(sk2,sk7)|memberP(sk1,sk7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_12)).
% 0.02/0.28  cnf(c_0_9, negated_conjecture, (ssItem(sk7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_11)).
% 0.02/0.28  cnf(c_0_10, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
% 0.02/0.28  cnf(c_0_11, axiom, (memberP(X2,X3)|memberP(X1,X3)|~memberP(app(X1,X2),X3)|~ssList(X2)|~ssList(X1)|~ssItem(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause158)).
% 0.02/0.28  cnf(c_0_12, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 0.02/0.28  cnf(c_0_13, axiom, (memberP(app(X1,X3),X2)|~memberP(X1,X2)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause140)).
% 0.02/0.28  cnf(c_0_14, negated_conjecture, (~memberP(sk2,sk7)|~memberP(sk1,sk7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_13)).
% 0.02/0.28  cnf(c_0_15, negated_conjecture, (app(sk6,sk5)=sk1), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 0.02/0.28  cnf(c_0_16, negated_conjecture, (app(sk6,app(sk5,X1))=app(sk1,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_15]), c_0_3]), c_0_4])])).
% 0.02/0.28  cnf(c_0_17, negated_conjecture, (app(sk5,sk6)=sk2), inference(rw,[status(thm)],[c_0_5, c_0_6])).
% 0.02/0.28  cnf(c_0_18, negated_conjecture, (memberP(app(X1,sk2),sk7)|memberP(sk1,sk7)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9]), c_0_10])])).
% 0.02/0.28  cnf(c_0_19, negated_conjecture, (app(sk6,sk2)=app(sk1,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_3])])).
% 0.02/0.28  cnf(c_0_20, negated_conjecture, (memberP(app(sk1,sk6),sk7)|memberP(sk1,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_3])])).
% 0.02/0.28  cnf(c_0_21, negated_conjecture, (memberP(sk1,sk7)|memberP(sk6,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_20]), c_0_9]), c_0_3]), c_0_12])])).
% 0.02/0.28  cnf(c_0_22, negated_conjecture, (memberP(app(sk6,X1),sk7)|memberP(sk1,sk7)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_21]), c_0_9]), c_0_3])])).
% 0.02/0.28  cnf(c_0_23, negated_conjecture, (memberP(sk1,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_15]), c_0_4])])).
% 0.02/0.28  cnf(c_0_24, negated_conjecture, (app(sk5,app(sk6,X1))=app(sk2,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_17]), c_0_4]), c_0_3])])).
% 0.02/0.28  cnf(c_0_25, negated_conjecture, (memberP(app(X1,sk1),sk7)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_23]), c_0_9]), c_0_12])])).
% 0.02/0.28  cnf(c_0_26, negated_conjecture, (app(sk5,sk1)=app(sk2,sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_15]), c_0_4])])).
% 0.02/0.28  cnf(c_0_27, negated_conjecture, (memberP(app(sk2,sk5),sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_4])])).
% 0.02/0.28  cnf(c_0_28, negated_conjecture, (~memberP(sk2,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_23])])).
% 0.02/0.28  cnf(c_0_29, negated_conjecture, (memberP(sk5,sk7)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_27]), c_0_9]), c_0_4]), c_0_10])]), c_0_28])).
% 0.02/0.28  cnf(c_0_30, negated_conjecture, (memberP(app(sk5,X1),sk7)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_29]), c_0_9]), c_0_4])])).
% 0.02/0.28  cnf(c_0_31, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_17]), c_0_3])]), c_0_28]), ['proof']).
% 0.02/0.28  # SZS output end CNFRefutation
% 0.02/0.28  # Training examples: 0 positive, 0 negative
% 0.02/0.28  
% 0.02/0.28  # -------------------------------------------------
% 0.02/0.28  # User time                : 0.046 s
% 0.02/0.28  # System time              : 0.004 s
% 0.02/0.28  # Total time               : 0.050 s
% 0.02/0.28  # Maximum resident set size: 1864 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
