<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC218-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC218-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC218-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC218-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC218-1.p">SWC218-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n130.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:05 EDT 2016

% Result     : Unsatisfiable 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of clauses        :   37 ( 159 expanded)
%              Number of leaves         :   18 ( 106 expanded)
%              Depth                    :   16
%              Number of atoms          :   63 ( 261 expanded)
%              Number of equality atoms :   30 ( 101 expanded)
%              Maximal clause size      :    4 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( app(app(sk5,cons(sk7,nil)),sk6) = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( nil = X2
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause119)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( ssList(sk6) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ~ neq(sk1,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ssList(sk5) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ssItem(sk7) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( app(sk5,sk6) = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( nil = X1
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause118)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause99)).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( app(app(sk5,cons(sk7,nil)),sk6) = sk1 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( sk6 = nil
    | sk1 != nil
    | ~ ssList(app(sk5,cons(sk7,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_18]),c_0_3])])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,c_0_5]),c_0_6]),c_0_7])])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( sk6 = nil
    | ~ ssList(app(sk5,cons(sk7,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_20])])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( sk6 = nil
    | ~ ssList(cons(sk7,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_8]),c_0_9])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( sk6 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_10]),c_0_11]),c_0_7])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( app(sk5,nil) = sk4 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( app(sk5,cons(sk7,nil)) = nil
    | sk1 != nil
    | ~ ssList(app(sk5,cons(sk7,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_18]),c_0_3])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( sk5 = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_24]),c_0_9])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( app(sk4,cons(sk7,nil)) = nil
    | ~ ssList(app(sk4,cons(sk7,nil))) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_20])]),c_0_26]),<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( app(sk4,cons(sk7,nil)) = nil
    | ~ ssList(cons(sk7,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_8]),c_0_28])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( app(sk4,cons(sk7,nil)) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_10]),c_0_11]),c_0_7])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( sk4 = nil
    | ~ ssList(cons(sk7,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_30]),c_0_28])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_10]),c_0_11]),c_0_7])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( app(nil,cons(sk7,nil)) = nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( cons(sk7,nil) = nil
    | ~ ssList(cons(sk7,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_33]),c_0_7])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( cons(sk7,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_10]),c_0_11]),c_0_7])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_35]),c_0_11]),c_0_7])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC218-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n130.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:37:23 CDT 2016
% 0.02/0.22  % CPUTime    : 
% 0.02/0.22  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.02/0.25  # and selection function SelectCQArNTEqFirst.
% 0.02/0.25  #
% 0.02/0.25  # Preprocessing time       : 0.009 s
% 0.02/0.25  # Presaturation interreduction done
% 0.02/0.25  
% 0.02/0.25  # Proof found!
% 0.02/0.25  # SZS status Unsatisfiable
% 0.02/0.25  # SZS output start CNFRefutation
% 0.02/0.25  cnf(c_0_0, negated_conjecture, (app(app(sk5,cons(sk7,nil)),sk6)=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_11)).
% 0.02/0.25  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.02/0.25  cnf(c_0_2, axiom, (nil=X2|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause119)).
% 0.02/0.25  cnf(c_0_3, negated_conjecture, (ssList(sk6)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
% 0.02/0.25  cnf(c_0_4, negated_conjecture, (~neq(sk1,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_12)).
% 0.02/0.25  cnf(c_0_5, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause100)).
% 0.02/0.25  cnf(c_0_6, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.02/0.25  cnf(c_0_7, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.02/0.25  cnf(c_0_8, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.02/0.25  cnf(c_0_9, negated_conjecture, (ssList(sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 0.02/0.25  cnf(c_0_10, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.02/0.25  cnf(c_0_11, negated_conjecture, (ssItem(sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_10)).
% 0.02/0.25  cnf(c_0_12, negated_conjecture, (app(sk5,sk6)=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
% 0.02/0.25  cnf(c_0_13, axiom, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause118)).
% 0.02/0.25  cnf(c_0_14, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause73)).
% 0.02/0.25  cnf(c_0_15, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 0.02/0.25  cnf(c_0_16, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 0.02/0.25  cnf(c_0_17, axiom, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause99)).
% 0.02/0.25  cnf(c_0_18, negated_conjecture, (app(app(sk5,cons(sk7,nil)),sk6)=sk1), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 0.02/0.25  cnf(c_0_19, negated_conjecture, (sk6=nil|sk1!=nil|~ssList(app(sk5,cons(sk7,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_18]), c_0_3])])).
% 0.02/0.25  cnf(c_0_20, negated_conjecture, (sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_5]), c_0_6]), c_0_7])])).
% 0.02/0.25  cnf(c_0_21, negated_conjecture, (sk6=nil|~ssList(app(sk5,cons(sk7,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20])])).
% 0.02/0.25  cnf(c_0_22, negated_conjecture, (sk6=nil|~ssList(cons(sk7,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_8]), c_0_9])])).
% 0.02/0.25  cnf(c_0_23, negated_conjecture, (sk6=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_10]), c_0_11]), c_0_7])])).
% 0.02/0.25  cnf(c_0_24, negated_conjecture, (app(sk5,nil)=sk4), inference(rw,[status(thm)],[c_0_12, c_0_23])).
% 0.02/0.25  cnf(c_0_25, negated_conjecture, (app(sk5,cons(sk7,nil))=nil|sk1!=nil|~ssList(app(sk5,cons(sk7,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_18]), c_0_3])])).
% 0.02/0.25  cnf(c_0_26, negated_conjecture, (sk5=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_24]), c_0_9])])).
% 0.02/0.25  cnf(c_0_27, negated_conjecture, (app(sk4,cons(sk7,nil))=nil|~ssList(app(sk4,cons(sk7,nil)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_20])]), c_0_26]), c_0_26])).
% 0.02/0.25  cnf(c_0_28, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
% 0.02/0.25  cnf(c_0_29, negated_conjecture, (app(sk4,cons(sk7,nil))=nil|~ssList(cons(sk7,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_8]), c_0_28])])).
% 0.02/0.25  cnf(c_0_30, negated_conjecture, (app(sk4,cons(sk7,nil))=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_10]), c_0_11]), c_0_7])])).
% 0.02/0.25  cnf(c_0_31, negated_conjecture, (sk4=nil|~ssList(cons(sk7,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_30]), c_0_28])])).
% 0.02/0.25  cnf(c_0_32, negated_conjecture, (sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_10]), c_0_11]), c_0_7])])).
% 0.02/0.25  cnf(c_0_33, negated_conjecture, (app(nil,cons(sk7,nil))=nil), inference(rw,[status(thm)],[c_0_30, c_0_32])).
% 0.02/0.25  cnf(c_0_34, negated_conjecture, (cons(sk7,nil)=nil|~ssList(cons(sk7,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_33]), c_0_7])])).
% 0.02/0.25  cnf(c_0_35, negated_conjecture, (cons(sk7,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_10]), c_0_11]), c_0_7])])).
% 0.02/0.25  cnf(c_0_36, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_35]), c_0_11]), c_0_7])]), ['proof']).
% 0.02/0.25  # SZS output end CNFRefutation
% 0.02/0.25  # Training examples: 0 positive, 0 negative
% 0.02/0.25  
% 0.02/0.25  # -------------------------------------------------
% 0.02/0.25  # User time                : 0.021 s
% 0.02/0.25  # System time              : 0.002 s
% 0.02/0.25  # Total time               : 0.023 s
% 0.02/0.25  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
