<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC107-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC107-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC107-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC107-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC107-1.p">SWC107-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n106.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:41 EDT 2016

% Result     : Unsatisfiable 159.29s
% Output     : CNFRefutation 159.29s
% Verified   : 
% Statistics : Number of clauses        :   90 ( 358 expanded)
%              Number of leaves         :   43 ( 241 expanded)
%              Depth                    :   17
%              Number of atoms          :  251 ( 784 expanded)
%              Number of equality atoms :   98 ( 340 expanded)
%              Maximal clause size      :    6 (   3 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( cons(sk5,nil) = sk3
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( ssItem(sk5)
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_3"></A>cnf(c_0_3,axiom,
    ( memberP(cons(X2,X3),X1)
    | X1 != X2
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause138)).

<A NAME="c_0_4"></A>cnf(c_0_4,axiom,
    ( memberP(X2,X3)
    | X3 = X1
    | ~ memberP(cons(X1,X2),X3)
    | ~ ssList(X2)
    | ~ ssItem(X1)
    | ~ ssItem(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause161)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause71)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( cons(skaf83(X1),skaf82(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause109)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( ssItem(skaf83(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause12)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( ssList(skaf82(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause13)).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( X4 = X2
    | cons(X1,X2) != cons(X3,X4)
    | ~ ssItem(X3)
    | ~ ssItem(X1)
    | ~ ssList(X4)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause171)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( cons(sk5,nil) = sk3
    | nil = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( cons(skaf44(X1),nil) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause101)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( ssItem(skaf44(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause47)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( singletonP(X2)
    | cons(X1,nil) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause116)).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( nil = sk2
    | ~ neq(sk1,nil)
    | ~ segmentP(sk2,sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_16)).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( cons(X1,X2) != nil
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause98)).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( ssItem(sk5)
    | nil = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( nil = X2
    | tl(app(X2,X1)) = app(tl(X2),X1)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause133)).

<A NAME="c_0_23"></A>cnf(c_0_23,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_25"></A>cnf(c_0_25,axiom,
    ( memberP(app(X3,X1),X2)
    | ~ memberP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssItem(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause141)).

<A NAME="c_0_26"></A>cnf(c_0_26,axiom,
    ( app(skaf46(X1,X2),X2) = X1
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause131)).

<A NAME="c_0_27"></A>cnf(c_0_27,axiom,
    ( ssList(skaf46(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause50)).

<A NAME="c_0_28"></A>cnf(c_0_28,axiom,
    ( rearsegP(X3,X2)
    | app(X1,X2) != X3
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause143)).

<A NAME="c_0_29"></A>cnf(c_0_29,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_30"></A>cnf(c_0_30,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_31"></A>cnf(c_0_31,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_32"></A>cnf(c_0_32,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_33"></A>cnf(c_0_33,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_34"></A>cnf(c_0_34,axiom,
    ( X2 = X3
    | app(X1,X2) != app(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause150)).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( memberP(sk4,sk5)
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_36"></A>cnf(c_0_36,axiom,
    ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
    | ~ memberP(X1,X2)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause169)).

<A NAME="c_0_37"></A>cnf(c_0_37,axiom,
    ( ssList(skaf43(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause52)).

<A NAME="c_0_38"></A>cnf(c_0_38,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( memberP(sk4,sk5)
    | nil = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_40"></A>cnf(c_0_40,axiom,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause60)).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( neq(sk2,nil)
    | nil != sk1 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_17)).

<A NAME="c_0_42"></A>cnf(c_0_42,axiom,
    ( X1 != X2
    | ~ neq(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause115)).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( cons(sk5,nil) = sk1
    | sk1 = nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,c_0_1]),<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( sk1 = nil
    | ssItem(sk5) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( sk1 = nil
    | sk5 = X1
    | ~ memberP(sk1,X1)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,c_0_43]),c_0_5])]),c_0_44]),<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_47"></A>cnf(c_0_47,plain,
    ( nil = X1
    | memberP(X1,skaf83(X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_7]),c_0_8]),c_0_9])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( skaf83(sk1) = sk5
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_47]),c_0_8]),c_0_10])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( sk1 = nil
    | X1 = nil
    | cons(X2,X1) != sk1
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_43]),c_0_5])]),<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( cons(sk5,skaf82(sk1)) = sk1
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_48]),c_0_10])])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( skaf82(sk1) = nil
    | sk1 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_50]),c_0_9])]),<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( cons(skaf83(sk1),nil) = sk1
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_51]),c_0_10])])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( cons(skaf83(sk1),X1) = app(sk1,X1)
    | sk1 = nil
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_52]),c_0_8])])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( cons(sk5,nil) = sk1
    | sk4 = nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( app(sk1,skaf82(sk1)) = sk1
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_53]),c_0_10]),c_0_9])])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( tl(X1) = nil
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_15]),c_0_16]),c_0_5])])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( sk1 = nil
    | singletonP(X1)
    | sk1 != X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_52]),c_0_8])])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( sk4 = nil
    | ~ neq(sk1,nil)
    | ~ segmentP(sk4,sk1) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_19]),<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( sk4 = nil
    | sk1 != nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_54]),c_0_5])]),<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( app(tl(sk1),skaf82(sk1)) = tl(sk1)
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_55]),c_0_10]),c_0_9])])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( tl(X1) = nil
    | sk1 = nil
    | sk1 != X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( sk4 = nil
    | ~ segmentP(sk4,sk1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_23]),c_0_10]),c_0_5])]),<A HREF="#c_0_59">c_0_59</A>])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_64"></A>cnf(c_0_64,plain,
    ( memberP(X1,X2)
    | ~ memberP(X3,X2)
    | ~ rearsegP(X1,X3)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26]),c_0_27])])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_28]),<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,negated_conjecture,
    ( app(nil,skaf82(sk1)) = nil
    | sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>,c_0_61]),c_0_10])])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( sk4 = nil
    | app(app(X1,sk1),X2) != sk4
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_30]),c_0_63]),c_0_10])])).

<A NAME="c_0_68"></A>cnf(c_0_68,plain,
    ( ~ memberP(X1,X2)
    | ~ rearsegP(nil,X1)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>,c_0_64]),c_0_5])])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( rearsegP(X1,skaf45(X1,X2))
    | ~ frontsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_31]),c_0_32])])])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( sk1 = nil
    | rearsegP(nil,skaf82(sk1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_66]),c_0_9]),c_0_5])])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( sk4 = nil
    | app(X1,app(sk1,X2)) != sk4
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,c_0_33]),c_0_10])])).

<A NAME="c_0_72"></A>cnf(c_0_72,negated_conjecture,
    ( app(sk1,X1) = cons(sk5,X1)
    | sk4 = nil
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_54]),<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,plain,
    ( ~ memberP(skaf45(nil,X1),X2)
    | ~ frontsegP(nil,X1)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_69]),c_0_32]),c_0_5])])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( sk1 = nil
    | ~ memberP(skaf82(sk1),X1)
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_70]),c_0_9])])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( X1 = skaf82(sk1)
    | sk1 = nil
    | app(nil,X1) != nil
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_66]),c_0_9]),c_0_5])])).

<A NAME="c_0_76"></A>cnf(c_0_76,negated_conjecture,
    ( sk4 = nil
    | app(X1,cons(sk5,X2)) != sk4
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>,<A HREF="#c_0_72">c_0_72</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,plain,
    ( skaf45(nil,X1) = nil
    | ~ frontsegP(nil,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>,c_0_47]),c_0_8]),c_0_32])])).

<A NAME="c_0_78"></A>cnf(c_0_78,negated_conjecture,
    ( sk1 = nil
    | app(nil,X1) != nil
    | ~ memberP(X1,X2)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,<A HREF="#c_0_75">c_0_75</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( sk1 = nil
    | memberP(sk4,sk5) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( sk4 = nil
    | X1 != sk4
    | ~ memberP(X1,sk5)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>,c_0_36]),c_0_37]),c_0_38])]),<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( app(X1,nil) = nil
    | ~ frontsegP(nil,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_77]),c_0_5])])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( sk1 = nil
    | app(nil,sk4) != nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_79]),c_0_63])]),<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,c_0_39]),c_0_63])])).

<A NAME="c_0_84"></A>cnf(c_0_84,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_40]),c_0_5])])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( neq(sk4,nil)
    | sk1 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,c_0_83]),c_0_84])])).

<A NAME="c_0_87"></A>cnf(c_0_87,plain,
    ( ~ neq(X1,X1)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_83]),c_0_86])])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,c_0_88]),c_0_5])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC107-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n106.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:35:23 CDT 2016
% 0.02/0.22  % CPUTime    : 
% 0.02/0.22  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 151.10/151.24  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 151.10/151.24  # and selection function SelectCQArNTEqFirst.
% 151.10/151.24  #
% 151.10/151.24  # Preprocessing time       : 0.010 s
% 151.10/151.24  # Presaturation interreduction done
% 151.40/151.57  # No success with AutoSched0
% 151.40/151.57  # Trying AutoSched1 for 73 seconds
% 159.29/159.50  # AutoSched1-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S00
% 159.29/159.50  # and selection function NoSelection.
% 159.29/159.50  #
% 159.29/159.50  # Preprocessing time       : 0.010 s
% 159.29/159.50  # Presaturation interreduction done
% 159.29/159.50  
% 159.29/159.50  # Proof found!
% 159.29/159.50  # SZS status Unsatisfiable
% 159.29/159.50  # SZS output start CNFRefutation
% 159.29/159.50  cnf(c_0_0, negated_conjecture, (cons(sk5,nil)=sk3|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_12)).
% 159.29/159.50  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 159.29/159.50  cnf(c_0_2, negated_conjecture, (ssItem(sk5)|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
% 159.29/159.50  cnf(c_0_3, axiom, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause138)).
% 159.29/159.50  cnf(c_0_4, axiom, (memberP(X2,X3)|X3=X1|~memberP(cons(X1,X2),X3)|~ssList(X2)|~ssItem(X1)|~ssItem(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause161)).
% 159.29/159.50  cnf(c_0_5, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 159.29/159.50  cnf(c_0_6, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause71)).
% 159.29/159.50  cnf(c_0_7, axiom, (cons(skaf83(X1),skaf82(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause109)).
% 159.29/159.50  cnf(c_0_8, axiom, (ssItem(skaf83(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause12)).
% 159.29/159.50  cnf(c_0_9, axiom, (ssList(skaf82(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause13)).
% 159.29/159.50  cnf(c_0_10, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 159.29/159.50  cnf(c_0_11, axiom, (X4=X2|cons(X1,X2)!=cons(X3,X4)|~ssItem(X3)|~ssItem(X1)|~ssList(X4)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause171)).
% 159.29/159.50  cnf(c_0_12, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
% 159.29/159.50  cnf(c_0_13, negated_conjecture, (cons(sk5,nil)=sk3|nil=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
% 159.29/159.50  cnf(c_0_14, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause96)).
% 159.29/159.50  cnf(c_0_15, axiom, (cons(skaf44(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause101)).
% 159.29/159.50  cnf(c_0_16, axiom, (ssItem(skaf44(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause47)).
% 159.29/159.50  cnf(c_0_17, axiom, (singletonP(X2)|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause116)).
% 159.29/159.50  cnf(c_0_18, negated_conjecture, (nil=sk2|~neq(sk1,nil)|~segmentP(sk2,sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_16)).
% 159.29/159.50  cnf(c_0_19, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 159.29/159.50  cnf(c_0_20, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause98)).
% 159.29/159.50  cnf(c_0_21, negated_conjecture, (ssItem(sk5)|nil=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 159.29/159.50  cnf(c_0_22, axiom, (nil=X2|tl(app(X2,X1))=app(tl(X2),X1)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause133)).
% 159.29/159.50  cnf(c_0_23, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause100)).
% 159.29/159.50  cnf(c_0_24, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 159.29/159.50  cnf(c_0_25, axiom, (memberP(app(X3,X1),X2)|~memberP(X1,X2)|~ssList(X1)|~ssList(X3)|~ssItem(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause141)).
% 159.29/159.50  cnf(c_0_26, axiom, (app(skaf46(X1,X2),X2)=X1|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause131)).
% 159.29/159.50  cnf(c_0_27, axiom, (ssList(skaf46(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause50)).
% 159.29/159.50  cnf(c_0_28, axiom, (rearsegP(X3,X2)|app(X1,X2)!=X3|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause143)).
% 159.29/159.50  cnf(c_0_29, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 159.29/159.50  cnf(c_0_30, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause173)).
% 159.29/159.50  cnf(c_0_31, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause132)).
% 159.29/159.50  cnf(c_0_32, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause51)).
% 159.29/159.50  cnf(c_0_33, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
% 159.29/159.50  cnf(c_0_34, axiom, (X2=X3|app(X1,X2)!=app(X1,X3)|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause150)).
% 159.29/159.50  cnf(c_0_35, negated_conjecture, (memberP(sk4,sk5)|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
% 159.29/159.50  cnf(c_0_36, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause169)).
% 159.29/159.50  cnf(c_0_37, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause52)).
% 159.29/159.50  cnf(c_0_38, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause53)).
% 159.29/159.50  cnf(c_0_39, negated_conjecture, (memberP(sk4,sk5)|nil=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_10)).
% 159.29/159.50  cnf(c_0_40, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause60)).
% 159.29/159.50  cnf(c_0_41, negated_conjecture, (neq(sk2,nil)|nil!=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_17)).
% 159.29/159.50  cnf(c_0_42, axiom, (X1!=X2|~neq(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause115)).
% 159.29/159.50  cnf(c_0_43, negated_conjecture, (cons(sk5,nil)=sk1|sk1=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_0, c_0_1]), c_0_1])).
% 159.29/159.50  cnf(c_0_44, negated_conjecture, (sk1=nil|ssItem(sk5)), inference(rw,[status(thm)],[c_0_2, c_0_1])).
% 159.29/159.50  cnf(c_0_45, plain, (memberP(cons(X1,X2),X1)|~ssItem(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_3])).
% 159.29/159.50  cnf(c_0_46, negated_conjecture, (sk1=nil|sk5=X1|~memberP(sk1,X1)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_43]), c_0_5])]), c_0_44]), c_0_6])).
% 159.29/159.50  cnf(c_0_47, plain, (nil=X1|memberP(X1,skaf83(X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_7]), c_0_8]), c_0_9])])).
% 159.29/159.50  cnf(c_0_48, negated_conjecture, (skaf83(sk1)=sk5|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_8]), c_0_10])])).
% 159.29/159.50  cnf(c_0_49, negated_conjecture, (sk1=nil|X1=nil|cons(X2,X1)!=sk1|~ssItem(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_43]), c_0_5])]), c_0_44])).
% 159.29/159.50  cnf(c_0_50, negated_conjecture, (cons(sk5,skaf82(sk1))=sk1|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_48]), c_0_10])])).
% 159.29/159.50  cnf(c_0_51, negated_conjecture, (skaf82(sk1)=nil|sk1=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_9])]), c_0_44])).
% 159.29/159.50  cnf(c_0_52, negated_conjecture, (cons(skaf83(sk1),nil)=sk1|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_51]), c_0_10])])).
% 159.29/159.50  cnf(c_0_53, negated_conjecture, (cons(skaf83(sk1),X1)=app(sk1,X1)|sk1=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_52]), c_0_8])])).
% 159.29/159.50  cnf(c_0_54, negated_conjecture, (cons(sk5,nil)=sk1|sk4=nil), inference(rw,[status(thm)],[c_0_13, c_0_1])).
% 159.29/159.50  cnf(c_0_55, negated_conjecture, (app(sk1,skaf82(sk1))=sk1|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_53]), c_0_10]), c_0_9])])).
% 159.29/159.50  cnf(c_0_56, plain, (tl(X1)=nil|~singletonP(X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_5])])).
% 159.29/159.50  cnf(c_0_57, negated_conjecture, (sk1=nil|singletonP(X1)|sk1!=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_52]), c_0_8])])).
% 159.29/159.50  cnf(c_0_58, negated_conjecture, (sk4=nil|~neq(sk1,nil)|~segmentP(sk4,sk1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_19])).
% 159.29/159.50  cnf(c_0_59, negated_conjecture, (sk4=nil|sk1!=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_54]), c_0_5])]), c_0_21])).
% 159.29/159.50  cnf(c_0_60, negated_conjecture, (app(tl(sk1),skaf82(sk1))=tl(sk1)|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_55]), c_0_10]), c_0_9])])).
% 159.29/159.50  cnf(c_0_61, negated_conjecture, (tl(X1)=nil|sk1=nil|sk1!=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
% 159.29/159.50  cnf(c_0_62, negated_conjecture, (sk4=nil|~segmentP(sk4,sk1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_23]), c_0_10]), c_0_5])]), c_0_59])).
% 159.29/159.50  cnf(c_0_63, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_24, c_0_19])).
% 159.29/159.50  cnf(c_0_64, plain, (memberP(X1,X2)|~memberP(X3,X2)|~rearsegP(X1,X3)|~ssItem(X2)|~ssList(X3)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
% 159.29/159.50  cnf(c_0_65, plain, (rearsegP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_28]), c_0_29])).
% 159.29/159.50  cnf(c_0_66, negated_conjecture, (app(nil,skaf82(sk1))=nil|sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_10])])).
% 159.29/159.50  cnf(c_0_67, negated_conjecture, (sk4=nil|app(app(X1,sk1),X2)!=sk4|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_30]), c_0_63]), c_0_10])])).
% 159.29/159.50  cnf(c_0_68, plain, (~memberP(X1,X2)|~rearsegP(nil,X1)|~ssItem(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_64]), c_0_5])])).
% 159.29/159.50  cnf(c_0_69, plain, (rearsegP(X1,skaf45(X1,X2))|~frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_31]), c_0_32])])])).
% 159.29/159.50  cnf(c_0_70, negated_conjecture, (sk1=nil|rearsegP(nil,skaf82(sk1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_9]), c_0_5])])).
% 159.29/159.50  cnf(c_0_71, negated_conjecture, (sk4=nil|app(X1,app(sk1,X2))!=sk4|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_33]), c_0_10])])).
% 159.29/159.50  cnf(c_0_72, negated_conjecture, (app(sk1,X1)=cons(sk5,X1)|sk4=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_54]), c_0_21])).
% 159.29/159.50  cnf(c_0_73, plain, (~memberP(skaf45(nil,X1),X2)|~frontsegP(nil,X1)|~ssItem(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_32]), c_0_5])])).
% 159.29/159.50  cnf(c_0_74, negated_conjecture, (sk1=nil|~memberP(skaf82(sk1),X1)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_70]), c_0_9])])).
% 159.29/159.50  cnf(c_0_75, negated_conjecture, (X1=skaf82(sk1)|sk1=nil|app(nil,X1)!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_66]), c_0_9]), c_0_5])])).
% 159.29/159.50  cnf(c_0_76, negated_conjecture, (sk4=nil|app(X1,cons(sk5,X2))!=sk4|~ssList(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
% 159.29/159.50  cnf(c_0_77, plain, (skaf45(nil,X1)=nil|~frontsegP(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_47]), c_0_8]), c_0_32])])).
% 159.29/159.50  cnf(c_0_78, negated_conjecture, (sk1=nil|app(nil,X1)!=nil|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
% 159.29/159.50  cnf(c_0_79, negated_conjecture, (sk1=nil|memberP(sk4,sk5)), inference(rw,[status(thm)],[c_0_35, c_0_1])).
% 159.29/159.50  cnf(c_0_80, negated_conjecture, (sk4=nil|X1!=sk4|~memberP(X1,sk5)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_36]), c_0_37]), c_0_38])]), c_0_21])).
% 159.29/159.50  cnf(c_0_81, plain, (app(X1,nil)=nil|~frontsegP(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_77]), c_0_5])])).
% 159.29/159.50  cnf(c_0_82, negated_conjecture, (sk1=nil|app(nil,sk4)!=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_63])]), c_0_44])).
% 159.29/159.50  cnf(c_0_83, negated_conjecture, (sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_39]), c_0_63])])).
% 159.29/159.50  cnf(c_0_84, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_40]), c_0_5])])).
% 159.29/159.50  cnf(c_0_85, negated_conjecture, (neq(sk4,nil)|sk1!=nil), inference(rw,[status(thm)],[c_0_41, c_0_19])).
% 159.29/159.50  cnf(c_0_86, negated_conjecture, (sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_83]), c_0_84])])).
% 159.29/159.50  cnf(c_0_87, plain, (~neq(X1,X1)|~ssList(X1)), inference(er,[status(thm)],[c_0_42])).
% 159.29/159.50  cnf(c_0_88, negated_conjecture, (neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_83]), c_0_86])])).
% 159.29/159.50  cnf(c_0_89, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_5])]), ['proof']).
% 159.29/159.50  # SZS output end CNFRefutation
% 159.29/159.50  # Training examples: 0 positive, 0 negative
% 159.40/159.53  
% 159.40/159.53  # -------------------------------------------------
% 159.40/159.53  # User time                : 155.851 s
% 159.40/159.53  # System time              : 3.532 s
% 159.40/159.53  # Total time               : 159.382 s
% 159.40/159.53  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
