<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC329-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC329-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC329-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC329-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC329-1.p">SWC329-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n165.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:31 EDT 2016

% Result     : Unsatisfiable 152.26s
% Output     : CNFRefutation 152.26s
% Verified   : 
% Statistics : Number of clauses        :   50 ( 220 expanded)
%              Number of leaves         :   26 ( 145 expanded)
%              Depth                    :   11
%              Number of atoms          :  116 ( 442 expanded)
%              Number of equality atoms :   40 ( 151 expanded)
%              Maximal clause size      :    6 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,axiom,
    ( memberP(app(X3,X1),X2)
    | ~ memberP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssItem(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause141)).

<A NAME="c_0_1"></A>cnf(c_0_1,axiom,
    ( app(X1,X2) = nil
    | nil != X1
    | nil != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause145)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause71)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( memberP(sk4,sk5)
    | nil = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ssItem(sk5)
    | nil = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( ~ segmentP(sk2,sk1)
    | ~ equalelemsP(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_15)).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( equalelemsP(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause1)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( segmentP(X1,X3)
    | ~ segmentP(X1,X2)
    | ~ segmentP(X2,X3)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause152)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( segmentP(nil,X1)
    | nil != X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause79)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( segmentP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause56)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( equalelemsP(cons(X1,nil))
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause64)).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( cons(sk5,nil) = sk3
    | nil = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( cons(sk5,nil) = sk3
    | nil = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( ssItem(sk5)
    | nil = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
    | ~ memberP(X1,X2)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause169)).

<A NAME="c_0_23"></A>cnf(c_0_23,axiom,
    ( ssList(skaf43(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause52)).

<A NAME="c_0_24"></A>cnf(c_0_24,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( memberP(sk4,sk5)
    | nil = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_26"></A>cnf(c_0_26,plain,
    ( nil != X1
    | nil != X2
    | ~ memberP(X1,X3)
    | ~ ssItem(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,c_0_1]),<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>,<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( sk3 = nil
    | sk4 != nil
    | nil != X1
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_5]),c_0_27])]),<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( ~ segmentP(sk4,sk3)
    | ~ equalelemsP(sk3) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_8]),c_0_4]),<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( sk3 = nil
    | sk4 != nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( sk4 != nil
    | ~ segmentP(sk4,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_30]),c_0_9])])).

<A NAME="c_0_32"></A>cnf(c_0_32,plain,
    ( segmentP(X1,X2)
    | nil != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_11]),c_0_12])]),<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( sk3 = nil
    | equalelemsP(sk3) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_15]),<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( sk4 = nil
    | equalelemsP(sk3) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_16]),<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( sk4 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_32]),c_0_12]),c_0_27])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( sk3 = nil
    | ~ segmentP(sk4,sk3) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,<A HREF="#c_0_33">c_0_33</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( equalelemsP(sk3) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( sk3 = nil
    | app(app(X1,sk3),X2) != sk4
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_19]),c_0_27]),c_0_37])])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( cons(sk5,nil) = sk3 ),
    inference(sr,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( ssItem(sk5) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( ~ segmentP(sk4,sk3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_38])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( sk3 = nil
    | app(X1,app(sk3,X2)) != sk4
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_20]),c_0_37])])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( app(sk3,X1) = cons(sk5,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_40]),c_0_41])])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( sk3 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_32]),c_0_37]),c_0_27])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( app(X1,cons(sk5,X2)) != sk4
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_44]),<A HREF="#c_0_45">c_0_45</A>])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( X1 != sk4
    | ~ memberP(X1,sk5)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_22]),c_0_23]),c_0_24]),c_0_41])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( memberP(sk4,sk5) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_48]),c_0_27])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC329-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n165.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:41:08 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.24  # No SInE strategy applied
% 0.03/0.24  # Trying AutoSched0 for 151 seconds
% 151.07/151.36  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 151.07/151.36  # and selection function SelectCQArNTEqFirst.
% 151.07/151.36  #
% 151.07/151.36  # Preprocessing time       : 0.010 s
% 151.07/151.36  # Presaturation interreduction done
% 151.37/151.69  # No success with AutoSched0
% 151.37/151.69  # Trying AutoSched1 for 73 seconds
% 152.26/152.56  # AutoSched1-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S00
% 152.26/152.56  # and selection function NoSelection.
% 152.26/152.56  #
% 152.26/152.56  # Preprocessing time       : 0.010 s
% 152.26/152.56  # Presaturation interreduction done
% 152.26/152.56  
% 152.26/152.56  # Proof found!
% 152.26/152.56  # SZS status Unsatisfiable
% 152.26/152.56  # SZS output start CNFRefutation
% 152.26/152.56  cnf(c_0_0, axiom, (memberP(app(X3,X1),X2)|~memberP(X1,X2)|~ssList(X1)|~ssList(X3)|~ssItem(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause141)).
% 152.26/152.56  cnf(c_0_1, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause145)).
% 152.26/152.56  cnf(c_0_2, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause71)).
% 152.26/152.56  cnf(c_0_3, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
% 152.26/152.56  cnf(c_0_4, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 152.26/152.56  cnf(c_0_5, negated_conjecture, (memberP(sk4,sk5)|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_13)).
% 152.26/152.56  cnf(c_0_6, negated_conjecture, (ssItem(sk5)|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
% 152.26/152.56  cnf(c_0_7, negated_conjecture, (~segmentP(sk2,sk1)|~equalelemsP(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_15)).
% 152.26/152.56  cnf(c_0_8, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 152.26/152.56  cnf(c_0_9, axiom, (equalelemsP(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause1)).
% 152.26/152.56  cnf(c_0_10, axiom, (segmentP(X1,X3)|~segmentP(X1,X2)|~segmentP(X2,X3)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause152)).
% 152.26/152.56  cnf(c_0_11, axiom, (segmentP(nil,X1)|nil!=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause79)).
% 152.26/152.56  cnf(c_0_12, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 152.26/152.56  cnf(c_0_13, axiom, (segmentP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause56)).
% 152.26/152.56  cnf(c_0_14, axiom, (equalelemsP(cons(X1,nil))|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause64)).
% 152.26/152.56  cnf(c_0_15, negated_conjecture, (cons(sk5,nil)=sk3|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_12)).
% 152.26/152.56  cnf(c_0_16, negated_conjecture, (cons(sk5,nil)=sk3|nil=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
% 152.26/152.56  cnf(c_0_17, negated_conjecture, (ssItem(sk5)|nil=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 152.26/152.56  cnf(c_0_18, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 152.26/152.56  cnf(c_0_19, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause173)).
% 152.26/152.56  cnf(c_0_20, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause149)).
% 152.26/152.56  cnf(c_0_21, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause120)).
% 152.26/152.56  cnf(c_0_22, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause169)).
% 152.26/152.56  cnf(c_0_23, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause52)).
% 152.26/152.56  cnf(c_0_24, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause53)).
% 152.26/152.56  cnf(c_0_25, negated_conjecture, (memberP(sk4,sk5)|nil=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
% 152.26/152.56  cnf(c_0_26, plain, (nil!=X1|nil!=X2|~memberP(X1,X3)|~ssItem(X3)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
% 152.26/152.56  cnf(c_0_27, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_3, c_0_4])).
% 152.26/152.56  cnf(c_0_28, negated_conjecture, (sk3=nil|sk4!=nil|nil!=X1|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_5]), c_0_27])]), c_0_6])).
% 152.26/152.56  cnf(c_0_29, negated_conjecture, (~segmentP(sk4,sk3)|~equalelemsP(sk3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_8]), c_0_4]), c_0_8])).
% 152.26/152.56  cnf(c_0_30, negated_conjecture, (sk3=nil|sk4!=nil), inference(spm,[status(thm)],[c_0_28, c_0_27])).
% 152.26/152.56  cnf(c_0_31, negated_conjecture, (sk4!=nil|~segmentP(sk4,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_9])])).
% 152.26/152.56  cnf(c_0_32, plain, (segmentP(X1,X2)|nil!=X2|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])]), c_0_13])).
% 152.26/152.56  cnf(c_0_33, negated_conjecture, (sk3=nil|equalelemsP(sk3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_6])).
% 152.26/152.56  cnf(c_0_34, negated_conjecture, (sk4=nil|equalelemsP(sk3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_16]), c_0_17])).
% 152.26/152.56  cnf(c_0_35, negated_conjecture, (sk4!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_12]), c_0_27])])).
% 152.26/152.56  cnf(c_0_36, negated_conjecture, (sk3=nil|~segmentP(sk4,sk3)), inference(spm,[status(thm)],[c_0_29, c_0_33])).
% 152.26/152.56  cnf(c_0_37, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_18, c_0_8])).
% 152.26/152.56  cnf(c_0_38, negated_conjecture, (equalelemsP(sk3)), inference(sr,[status(thm)],[c_0_34, c_0_35])).
% 152.26/152.56  cnf(c_0_39, negated_conjecture, (sk3=nil|app(app(X1,sk3),X2)!=sk4|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_19]), c_0_27]), c_0_37])])).
% 152.26/152.56  cnf(c_0_40, negated_conjecture, (cons(sk5,nil)=sk3), inference(sr,[status(thm)],[c_0_16, c_0_35])).
% 152.26/152.56  cnf(c_0_41, negated_conjecture, (ssItem(sk5)), inference(sr,[status(thm)],[c_0_17, c_0_35])).
% 152.26/152.56  cnf(c_0_42, negated_conjecture, (~segmentP(sk4,sk3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_38])])).
% 152.26/152.56  cnf(c_0_43, negated_conjecture, (sk3=nil|app(X1,app(sk3,X2))!=sk4|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_20]), c_0_37])])).
% 152.26/152.56  cnf(c_0_44, negated_conjecture, (app(sk3,X1)=cons(sk5,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_40]), c_0_41])])).
% 152.26/152.56  cnf(c_0_45, negated_conjecture, (sk3!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_32]), c_0_37]), c_0_27])])).
% 152.26/152.56  cnf(c_0_46, negated_conjecture, (app(X1,cons(sk5,X2))!=sk4|~ssList(X2)|~ssList(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
% 152.26/152.56  cnf(c_0_47, negated_conjecture, (X1!=sk4|~memberP(X1,sk5)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_22]), c_0_23]), c_0_24]), c_0_41])])).
% 152.26/152.56  cnf(c_0_48, negated_conjecture, (memberP(sk4,sk5)), inference(sr,[status(thm)],[c_0_25, c_0_35])).
% 152.26/152.56  cnf(c_0_49, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_27])]), ['proof']).
% 152.26/152.56  # SZS output end CNFRefutation
% 152.26/152.56  # Training examples: 0 positive, 0 negative
% 152.26/152.57  
% 152.26/152.57  # -------------------------------------------------
% 152.26/152.57  # User time                : 148.772 s
% 152.26/152.57  # System time              : 3.522 s
% 152.26/152.57  # Total time               : 152.294 s
% 152.26/152.57  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
