<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC409-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC409-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC409-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC409-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC409-1.p">SWC409-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n136.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:50 EDT 2016

% Result     : Unsatisfiable 13.73s
% Output     : CNFRefutation 13.73s
% Verified   : 
% Statistics : Number of clauses        :  183 (17551 expanded)
%              Number of leaves         :   57 (11035 expanded)
%              Depth                    :   26
%              Number of atoms          :  462 (37063 expanded)
%              Number of equality atoms :  166 (6750 expanded)
%              Maximal clause size      :    7 (   3 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( memberP(sk2,sk5) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_3"></A>cnf(c_0_3,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_4"></A>cnf(c_0_4,axiom,
    ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
    | ~ memberP(X1,X2)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause169)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( ssItem(sk5) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( ssList(skaf43(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause52)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause99)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause104)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause75)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause76)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( singletonP(X2)
    | cons(X1,nil) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause116)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( nil = X2
    | tl(app(X2,X1)) = app(tl(X2),X1)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause133)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( nil = X1
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause118)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( cons(skaf44(X1),nil) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause101)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( ssItem(skaf44(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause47)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_23"></A>cnf(c_0_23,axiom,
    ( cons(X1,X2) != nil
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause98)).

<A NAME="c_0_24"></A>cnf(c_0_24,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( app(X2,cons(X1,nil)) = sk3
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | app(cons(X1,nil),X2) != sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_26"></A>cnf(c_0_26,axiom,
    ( rearsegP(X3,X2)
    | app(X1,X2) != X3
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause143)).

<A NAME="c_0_27"></A>cnf(c_0_27,axiom,
    ( cons(skaf83(X1),skaf82(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause109)).

<A NAME="c_0_28"></A>cnf(c_0_28,axiom,
    ( ssList(skaf82(X1)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause13)).

<A NAME="c_0_29"></A>cnf(c_0_29,axiom,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause97)).

<A NAME="c_0_30"></A>cnf(c_0_30,axiom,
    ( rearsegP(nil,X1)
    | nil != X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause81)).

<A NAME="c_0_31"></A>cnf(c_0_31,axiom,
    ( rearsegP(X1,X3)
    | ~ rearsegP(X1,X2)
    | ~ rearsegP(X2,X3)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause153)).

<A NAME="c_0_32"></A>cnf(c_0_32,axiom,
    ( rearsegP(app(X3,X1),X2)
    | ~ rearsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause136)).

<A NAME="c_0_33"></A>cnf(c_0_33,axiom,
    ( rearsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause58)).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_36"></A>cnf(c_0_36,axiom,
    ( app(skaf46(X1,X2),X2) = X1
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause131)).

<A NAME="c_0_37"></A>cnf(c_0_37,axiom,
    ( ssList(skaf46(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause50)).

<A NAME="c_0_38"></A>cnf(c_0_38,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_39"></A>cnf(c_0_39,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_40"></A>cnf(c_0_40,axiom,
    ( frontsegP(app(X1,X3),X2)
    | ~ frontsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause137)).

<A NAME="c_0_41"></A>cnf(c_0_41,axiom,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause60)).

<A NAME="c_0_42"></A>cnf(c_0_42,axiom,
    ( frontsegP(X3,X1)
    | app(X1,X2) != X3
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause144)).

<A NAME="c_0_43"></A>cnf(c_0_43,axiom,
    ( nil = X2
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause119)).

<A NAME="c_0_44"></A>cnf(c_0_44,axiom,
    ( app(X1,X2) = nil
    | nil != X1
    | nil != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause145)).

<A NAME="c_0_45"></A>cnf(c_0_45,axiom,
    ( X4 = X2
    | cons(X1,X2) != cons(X3,X4)
    | ~ ssItem(X3)
    | ~ ssItem(X1)
    | ~ ssList(X4)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause171)).

<A NAME="c_0_46"></A>cnf(c_0_46,axiom,
    ( nil = X2
    | hd(app(X2,X1)) = hd(X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause123)).

<A NAME="c_0_47"></A>cnf(c_0_47,axiom,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause157)).

<A NAME="c_0_48"></A>cnf(c_0_48,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_49"></A>cnf(c_0_49,axiom,
    ( memberP(cons(X2,X3),X1)
    | X1 != X2
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause138)).

<A NAME="c_0_50"></A>cnf(c_0_50,axiom,
    ( X1 = X3
    | app(X1,X2) != app(X3,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause151)).

<A NAME="c_0_51"></A>cnf(c_0_51,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_52"></A>cnf(c_0_52,axiom,
    ( memberP(X4,X2)
    | app(X1,cons(X2,X3)) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause175)).

<A NAME="c_0_53"></A>cnf(c_0_53,axiom,
    ( memberP(X2,X3)
    | X3 = X1
    | ~ memberP(cons(X1,X2),X3)
    | ~ ssList(X2)
    | ~ ssItem(X1)
    | ~ ssItem(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause161)).

<A NAME="c_0_54"></A>cnf(c_0_54,axiom,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause71)).

<A NAME="c_0_55"></A>cnf(c_0_55,axiom,
    ( memberP(app(X1,X3),X2)
    | ~ memberP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause140)).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( ~ memberP(sk1,sk5) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( memberP(sk4,sk5) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( app(skaf42(sk4,sk5),cons(sk5,skaf43(sk5,sk4))) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,c_0_57]),c_0_5]),c_0_58])])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( ~ ssList(app(sk4,cons(sk5,X1)))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(sk4,cons(sk5,X1))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,c_0_60]),c_0_5]),c_0_6]),c_0_7])])).

<A NAME="c_0_62"></A>cnf(c_0_62,plain,
    ( nil = X1
    | tl(X1) != X1
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_9]),c_0_10]),<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_63"></A>cnf(c_0_63,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1)
    | ~ ssList(cons(X1,nil)) ),
    inference(er,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(app(sk4,cons(sk5,X2)))
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( nil = X1
    | app(tl(X1),X2) != app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_14]),c_0_15]),<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,plain,
    ( tl(X1) = nil
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_18]),c_0_19]),c_0_20])])).

<A NAME="c_0_67"></A>cnf(c_0_67,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>,c_0_21]),c_0_20])])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(cons(sk5,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,c_0_15]),c_0_58])])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( nil = X1
    | tl(X1) != app(X1,nil)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_22]),c_0_20])]),<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,plain,
    ( tl(cons(X1,nil)) = nil
    | ~ ssItem(X1)
    | ~ ssList(cons(X1,nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>,<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_21]),c_0_5])])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( app(cons(X1,X2),nil) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>,c_0_17]),c_0_21]),<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,plain,
    ( tl(cons(X1,nil)) = nil
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>,c_0_21]),c_0_20])])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>,<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_75"></A>cnf(c_0_75,plain,
    ( cons(X1,nil) != nil
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_24]),c_0_20])])).

<A NAME="c_0_76"></A>cnf(c_0_76,plain,
    ( tl(cons(X1,nil)) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>,c_0_74])])).

<A NAME="c_0_77"></A>cnf(c_0_77,plain,
    ( cons(X1,nil) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_74])])).

<A NAME="c_0_78"></A>cnf(c_0_78,negated_conjecture,
    ( app(X1,cons(X2,nil)) = sk3
    | app(cons(X2,nil),X1) != sk4
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_74])])).

<A NAME="c_0_79"></A>cnf(c_0_79,plain,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_74])])).

<A NAME="c_0_80"></A>cnf(c_0_80,plain,
    ( cons(hd(cons(X1,nil)),nil) = cons(X1,nil)
    | ~ ssList(cons(X1,nil)) ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_76]),<A HREF="#c_0_77">c_0_77</A>])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( ssList(cons(X1,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_74])])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( app(X1,cons(X2,nil)) = sk3
    | cons(X2,X1) != sk4
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,<A HREF="#c_0_79">c_0_79</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,plain,
    ( cons(hd(cons(X1,nil)),nil) = cons(X1,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,c_0_81]),c_0_20])])).

<A NAME="c_0_84"></A>cnf(c_0_84,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_26]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( app(skaf82(sk4),cons(skaf83(sk4),nil)) = sk3
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,c_0_27]),c_0_28])])]),c_0_58])])).

<A NAME="c_0_86"></A>cnf(c_0_86,plain,
    ( ssList(cons(X1,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_83]),c_0_20])])).

<A NAME="c_0_87"></A>cnf(c_0_87,plain,
    ( hd(cons(X1,X2)) = X1
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_74])])).

<A NAME="c_0_88"></A>cnf(c_0_88,plain,
    ( rearsegP(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_30]),c_0_20])])).

<A NAME="c_0_89"></A>cnf(c_0_89,plain,
    ( rearsegP(X1,X2)
    | ~ rearsegP(X1,app(X3,X2))
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_84]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( sk4 = nil
    | rearsegP(sk3,cons(skaf83(sk4),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,c_0_85]),c_0_86]),c_0_28])])).

<A NAME="c_0_91"></A>cnf(c_0_91,plain,
    ( skaf83(X1) = hd(X1)
    | nil = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,c_0_27]),c_0_28])])).

<A NAME="c_0_92"></A>cnf(c_0_92,plain,
    ( rearsegP(app(X1,X2),nil)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_33]),c_0_20])])).

<A NAME="c_0_93"></A>cnf(c_0_93,plain,
    ( rearsegP(app(X1,nil),nil)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_88]),c_0_20])])).

<A NAME="c_0_94"></A>cnf(c_0_94,plain,
    ( rearsegP(X1,X2)
    | ~ rearsegP(X1,cons(X3,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>,c_0_79]),c_0_86])])).

<A NAME="c_0_95"></A>cnf(c_0_95,negated_conjecture,
    ( sk4 = nil
    | rearsegP(sk3,cons(hd(sk4),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,c_0_91]),c_0_58])])).

<A NAME="c_0_96"></A>cnf(c_0_96,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,negated_conjecture,
    ( rearsegP(sk4,nil)
    | ~ ssList(cons(sk5,skaf43(sk5,sk4))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_60]),c_0_7])])).

<A NAME="c_0_98"></A>cnf(c_0_98,plain,
    ( rearsegP(cons(X1,nil),nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_79]),c_0_86]),c_0_20])])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( app(skaf46(X1,nil),nil) = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_33]),c_0_20])])).

<A NAME="c_0_100"></A>cnf(c_0_100,negated_conjecture,
    ( sk4 = nil
    | rearsegP(sk3,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>,c_0_95]),c_0_20]),c_0_96])])).

<A NAME="c_0_101"></A>cnf(c_0_101,negated_conjecture,
    ( rearsegP(sk4,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_81]),c_0_6])])).

<A NAME="c_0_102"></A>cnf(c_0_102,plain,
    ( app(skaf46(cons(X1,nil),nil),nil) = cons(X1,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_98]),c_0_20]),c_0_86])])).

<A NAME="c_0_103"></A>cnf(c_0_103,plain,
    ( skaf46(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_99]),c_0_37])])).

<A NAME="c_0_104"></A>cnf(c_0_104,negated_conjecture,
    ( app(skaf46(sk3,nil),nil) = sk3
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_100]),c_0_20]),c_0_96])])).

<A NAME="c_0_105"></A>cnf(c_0_105,negated_conjecture,
    ( app(skaf46(sk4,nil),nil) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_101]),c_0_20]),c_0_58])])).

<A NAME="c_0_106"></A>cnf(c_0_106,plain,
    ( app(nil,app(X1,X2)) = app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_39]),c_0_20])])).

<A NAME="c_0_107"></A>cnf(c_0_107,plain,
    ( app(cons(X1,nil),nil) = cons(X1,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>,c_0_103]),c_0_86])])).

<A NAME="c_0_108"></A>cnf(c_0_108,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_41]),c_0_20])])).

<A NAME="c_0_109"></A>cnf(c_0_109,negated_conjecture,
    ( app(sk3,nil) = sk3
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>,c_0_103]),c_0_96])])).

<A NAME="c_0_110"></A>cnf(c_0_110,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_42]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_111"></A>cnf(c_0_111,negated_conjecture,
    ( app(sk4,nil) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_105">c_0_105</A>,c_0_103]),c_0_58])])).

<A NAME="c_0_112"></A>cnf(c_0_112,plain,
    ( nil = X1
    | app(X2,app(X3,X1)) != nil
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_38]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_113"></A>cnf(c_0_113,plain,
    ( app(nil,cons(X1,nil)) = cons(X1,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>,c_0_107]),c_0_86]),c_0_20])])).

<A NAME="c_0_114"></A>cnf(c_0_114,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_44])]),c_0_20])])).

<A NAME="c_0_115"></A>cnf(c_0_115,negated_conjecture,
    ( sk4 = nil
    | frontsegP(sk3,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>,c_0_109]),c_0_20]),c_0_96])])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( sk4 = nil
    | frontsegP(sk3,sk3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_110">c_0_110</A>,c_0_109]),c_0_20]),c_0_96])])).

<A NAME="c_0_117"></A>cnf(c_0_117,plain,
    ( X1 = X2
    | cons(X3,X1) != cons(X4,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_74]),c_0_74])])).

<A NAME="c_0_118"></A>cnf(c_0_118,plain,
    ( cons(X1,X2) != nil
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_74])])).

<A NAME="c_0_119"></A>cnf(c_0_119,plain,
    ( cons(hd(X1),tl(app(X1,X2))) = app(X1,X2)
    | nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_46]),c_0_15]),<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_120"></A>cnf(c_0_120,plain,
    ( nil = X1
    | ssList(app(tl(X1),X2))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_14]),c_0_15]),<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_121"></A>cnf(c_0_121,negated_conjecture,
    ( app(tl(sk4),nil) = tl(sk4)
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_111]),c_0_58]),c_0_20])])).

<A NAME="c_0_122"></A>cnf(c_0_122,plain,
    ( app(X1,X2) = nil
    | app(X1,app(X2,X3)) != nil
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_38]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_123"></A>cnf(c_0_123,plain,
    ( app(X1,cons(X2,nil)) != nil
    | ~ ssList(X1) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_113]),c_0_86]),c_0_20])]),<A HREF="#c_0_77">c_0_77</A>])).

<A NAME="c_0_124"></A>cnf(c_0_124,plain,
    ( app(cons(X1,X2),X3) = cons(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_74])])).

<A NAME="c_0_125"></A>cnf(c_0_125,plain,
    ( app(nil,app(nil,X1)) = app(nil,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_114]),c_0_20])])).

<A NAME="c_0_126"></A>cnf(c_0_126,negated_conjecture,
    ( app(nil,skaf45(sk3,nil)) = sk3
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_115]),c_0_20]),c_0_96])])).

<A NAME="c_0_127"></A>cnf(c_0_127,plain,
    ( ssList(app(X1,app(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_38]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_128"></A>cnf(c_0_128,negated_conjecture,
    ( app(sk3,skaf45(sk3,sk3)) = sk3
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_116]),c_0_96])])).

<A NAME="c_0_129"></A>cnf(c_0_129,plain,
    ( skaf82(cons(X1,X2)) = X2
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_117">c_0_117</A>,c_0_27]),c_0_28])])]),c_0_81]),<A HREF="#c_0_118">c_0_118</A>])).

<A NAME="c_0_130"></A>cnf(c_0_130,negated_conjecture,
    ( cons(hd(sk4),tl(sk4)) = sk4
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>,c_0_111]),c_0_58]),c_0_20])])).

<A NAME="c_0_131"></A>cnf(c_0_131,negated_conjecture,
    ( sk4 = nil
    | ssList(tl(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_120">c_0_120</A>,c_0_121]),c_0_58]),c_0_20])])).

<A NAME="c_0_132"></A>cnf(c_0_132,plain,
    ( app(X1,cons(X2,X3)) != nil
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_122">c_0_122</A>,c_0_79]),c_0_86])]),<A HREF="#c_0_123">c_0_123</A>])).

<A NAME="c_0_133"></A>cnf(c_0_133,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_134"></A>cnf(c_0_134,plain,
    ( X1 = nil
    | app(X1,nil) != nil
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_114]),c_0_20])])).

<A NAME="c_0_135"></A>cnf(c_0_135,plain,
    ( cons(hd(app(X1,X2)),app(tl(X1),X2)) = app(X1,X2)
    | nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_14]),c_0_15]),<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_136"></A>cnf(c_0_136,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_74])])).

<A NAME="c_0_137"></A>cnf(c_0_137,plain,
    ( tl(cons(X1,app(X2,X3))) = app(tl(cons(X1,X2)),X3)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_124]),c_0_81]),<A HREF="#c_0_118">c_0_118</A>])).

<A NAME="c_0_138"></A>cnf(c_0_138,negated_conjecture,
    ( app(nil,sk3) = sk3
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_125">c_0_125</A>,c_0_126]),c_0_51])])).

<A NAME="c_0_139"></A>cnf(c_0_139,negated_conjecture,
    ( sk4 = nil
    | ssList(app(X1,sk3))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_127">c_0_127</A>,c_0_128]),c_0_51]),c_0_96])])).

<A NAME="c_0_140"></A>cnf(c_0_140,negated_conjecture,
    ( sk4 = nil
    | app(X1,sk3) != nil
    | ~ ssList(X1) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_85]),c_0_86]),c_0_28])]),<A HREF="#c_0_77">c_0_77</A>])).

<A NAME="c_0_141"></A>cnf(c_0_141,negated_conjecture,
    ( skaf82(sk4) = tl(sk4)
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_129">c_0_129</A>,c_0_130]),<A HREF="#c_0_131">c_0_131</A>])).

<A NAME="c_0_142"></A>cnf(c_0_142,negated_conjecture,
    ( sk4 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_132">c_0_132</A>,c_0_60]),c_0_6]),c_0_7])])).

<A NAME="c_0_143"></A>cnf(c_0_143,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_133">c_0_133</A>,c_0_74])])).

<A NAME="c_0_144"></A>cnf(c_0_144,plain,
    ( cons(X1,app(X2,nil)) != nil
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_134">c_0_134</A>,c_0_124]),c_0_20])]),c_0_81]),<A HREF="#c_0_118">c_0_118</A>])).

<A NAME="c_0_145"></A>cnf(c_0_145,plain,
    ( cons(hd(app(cons(X1,X2),X3)),app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_135">c_0_135</A>,c_0_136]),c_0_81]),<A HREF="#c_0_118">c_0_118</A>])).

<A NAME="c_0_146"></A>cnf(c_0_146,plain,
    ( cons(X1,tl(cons(X1,X2))) = cons(X1,X2)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_87]),c_0_81]),<A HREF="#c_0_118">c_0_118</A>])).

<A NAME="c_0_147"></A>cnf(c_0_147,negated_conjecture,
    ( tl(cons(X1,sk3)) = app(nil,sk3)
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_137">c_0_137</A>,c_0_138]),c_0_76]),c_0_96]),c_0_20])])).

<A NAME="c_0_148"></A>cnf(c_0_148,negated_conjecture,
    ( sk4 = nil
    | ssList(cons(X1,sk3)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_139">c_0_139</A>,c_0_79]),c_0_86]),c_0_96])])).

<A NAME="c_0_149"></A>cnf(c_0_149,negated_conjecture,
    ( sk4 = nil
    | cons(X1,sk3) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_140">c_0_140</A>,c_0_79]),c_0_86]),c_0_96])])).

<A NAME="c_0_150"></A>cnf(c_0_150,plain,
    ( memberP(app(X1,cons(X2,X3)),X2)
    | ~ ssItem(X2)
    | ~ ssList(app(X1,cons(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>])).

<A NAME="c_0_151"></A>cnf(c_0_151,negated_conjecture,
    ( app(skaf82(sk4),cons(hd(sk4),nil)) = sk3
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_91]),c_0_58])])).

<A NAME="c_0_152"></A>cnf(c_0_152,negated_conjecture,
    ( skaf82(sk4) = tl(sk4) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_141">c_0_141</A>,<A HREF="#c_0_142">c_0_142</A>])).

<A NAME="c_0_153"></A>cnf(c_0_153,plain,
    ( X1 = X2
    | memberP(X3,X2)
    | ~ memberP(cons(X1,X3),X2)
    | ~ ssList(X3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_74]),c_0_74])])).

<A NAME="c_0_154"></A>cnf(c_0_154,plain,
    ( memberP(cons(X1,nil),hd(cons(X1,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_143">c_0_143</A>,c_0_83]),c_0_20])])).

<A NAME="c_0_155"></A>cnf(c_0_155,plain,
    ( ~ memberP(nil,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_74])])).

<A NAME="c_0_156"></A>cnf(c_0_156,plain,
    ( app(cons(X1,X2),nil) != nil
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_144">c_0_144</A>,c_0_145]),c_0_20])])).

<A NAME="c_0_157"></A>cnf(c_0_157,negated_conjecture,
    ( cons(X1,app(nil,sk3)) = cons(X1,sk3)
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_146">c_0_146</A>,c_0_147]),c_0_96])])).

<A NAME="c_0_158"></A>cnf(c_0_158,negated_conjecture,
    ( sk4 = nil
    | ssList(app(nil,sk3)) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_147]),c_0_148]),<A HREF="#c_0_149">c_0_149</A>])).

<A NAME="c_0_159"></A>cnf(c_0_159,plain,
    ( memberP(app(X1,cons(X2,X3)),X2)
    | ~ ssList(app(X1,cons(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_150">c_0_150</A>,c_0_74])])).

<A NAME="c_0_160"></A>cnf(c_0_160,plain,
    ( cons(hd(X1),app(tl(X1),X2)) = app(X1,X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_124">c_0_124</A>,c_0_9]),<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_161"></A>cnf(c_0_161,negated_conjecture,
    ( app(tl(sk4),cons(hd(sk4),nil)) = sk3 ),
    inference(rw,[status(thm)],[inference(sr,[status(thm)],[<A HREF="#c_0_151">c_0_151</A>,c_0_142]),<A HREF="#c_0_152">c_0_152</A>])).

<A NAME="c_0_162"></A>cnf(c_0_162,plain,
    ( hd(cons(X1,app(X2,X3))) = hd(cons(X1,X2))
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_124]),c_0_81]),<A HREF="#c_0_118">c_0_118</A>])).

<A NAME="c_0_163"></A>cnf(c_0_163,plain,
    ( hd(cons(X1,nil)) = X1 ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_153">c_0_153</A>,c_0_154]),c_0_20])]),<A HREF="#c_0_155">c_0_155</A>])).

<A NAME="c_0_164"></A>cnf(c_0_164,negated_conjecture,
    ( sk4 = nil
    | app(cons(X1,sk3),nil) != nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_156">c_0_156</A>,c_0_157]),<A HREF="#c_0_158">c_0_158</A>])).

<A NAME="c_0_165"></A>cnf(c_0_165,negated_conjecture,
    ( sk4 = nil
    | memberP(sk3,skaf83(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_159">c_0_159</A>,c_0_85]),c_0_96]),c_0_20]),c_0_28])])).

<A NAME="c_0_166"></A>cnf(c_0_166,negated_conjecture,
    ( memberP(app(sk4,X1),sk5)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_57]),c_0_5]),c_0_58])])).

<A NAME="c_0_167"></A>cnf(c_0_167,negated_conjecture,
    ( app(sk4,cons(hd(sk4),nil)) = cons(hd(sk4),sk3) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_160">c_0_160</A>,c_0_161]),c_0_86]),c_0_58])]),<A HREF="#c_0_142">c_0_142</A>])).

<A NAME="c_0_168"></A>cnf(c_0_168,negated_conjecture,
    ( hd(cons(X1,sk3)) = X1
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_162">c_0_162</A>,c_0_138]),c_0_163]),c_0_96]),c_0_20])])).

<A NAME="c_0_169"></A>cnf(c_0_169,negated_conjecture,
    ( app(nil,sk3) = sk3 ),
    inference(sr,[status(thm)],[<A HREF="#c_0_138">c_0_138</A>,<A HREF="#c_0_142">c_0_142</A>])).

<A NAME="c_0_170"></A>cnf(c_0_170,negated_conjecture,
    ( sk4 = nil
    | cons(X1,app(sk3,nil)) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_164">c_0_164</A>,c_0_124]),c_0_20]),c_0_96])])).

<A NAME="c_0_171"></A>cnf(c_0_171,negated_conjecture,
    ( app(sk3,nil) = sk3 ),
    inference(sr,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>,<A HREF="#c_0_142">c_0_142</A>])).

<A NAME="c_0_172"></A>cnf(c_0_172,negated_conjecture,
    ( sk4 = nil
    | memberP(sk3,hd(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_165">c_0_165</A>,c_0_91]),c_0_58])])).

<A NAME="c_0_173"></A>cnf(c_0_173,plain,
    ( hd(X1) = X2
    | nil = X1
    | memberP(tl(X1),X2)
    | ~ memberP(X1,X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_153">c_0_153</A>,c_0_9]),<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_174"></A>cnf(c_0_174,negated_conjecture,
    ( memberP(cons(hd(sk4),sk3),sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_166">c_0_166</A>,c_0_167]),c_0_86])])).

<A NAME="c_0_175"></A>cnf(c_0_175,negated_conjecture,
    ( hd(cons(X1,sk3)) = X1 ),
    inference(sr,[status(thm)],[<A HREF="#c_0_168">c_0_168</A>,<A HREF="#c_0_142">c_0_142</A>])).

<A NAME="c_0_176"></A>cnf(c_0_176,negated_conjecture,
    ( tl(cons(X1,sk3)) = sk3 ),
    inference(rw,[status(thm)],[inference(sr,[status(thm)],[<A HREF="#c_0_147">c_0_147</A>,c_0_142]),<A HREF="#c_0_169">c_0_169</A>])).

<A NAME="c_0_177"></A>cnf(c_0_177,negated_conjecture,
    ( ssList(cons(X1,sk3)) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_148">c_0_148</A>,<A HREF="#c_0_142">c_0_142</A>])).

<A NAME="c_0_178"></A>cnf(c_0_178,negated_conjecture,
    ( cons(X1,sk3) != nil ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_170">c_0_170</A>,c_0_171]),<A HREF="#c_0_142">c_0_142</A>])).

<A NAME="c_0_179"></A>cnf(c_0_179,negated_conjecture,
    ( ~ memberP(sk3,sk5) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_180"></A>cnf(c_0_180,negated_conjecture,
    ( memberP(sk3,hd(sk4)) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_172">c_0_172</A>,<A HREF="#c_0_142">c_0_142</A>])).

<A NAME="c_0_181"></A>cnf(c_0_181,negated_conjecture,
    ( hd(sk4) = sk5 ),
    inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_173">c_0_173</A>,c_0_174]),c_0_175]),c_0_176]),c_0_177])]),c_0_178]),<A HREF="#c_0_179">c_0_179</A>])).

<A NAME="c_0_182"></A>cnf(c_0_182,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_180">c_0_180</A>,c_0_181]),<A HREF="#c_0_179">c_0_179</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC409-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n136.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:43:38 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 13.73/13.99  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 13.73/13.99  # and selection function SelectCQArNTEqFirst.
% 13.73/13.99  #
% 13.73/13.99  # Preprocessing time       : 0.009 s
% 13.73/13.99  # Presaturation interreduction done
% 13.73/13.99  
% 13.73/13.99  # Proof found!
% 13.73/13.99  # SZS status Unsatisfiable
% 13.73/13.99  # SZS output start CNFRefutation
% 13.73/13.99  cnf(c_0_0, negated_conjecture, (memberP(sk2,sk5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
% 13.73/13.99  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 13.73/13.99  cnf(c_0_2, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
% 13.73/13.99  cnf(c_0_3, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause179)).
% 13.73/13.99  cnf(c_0_4, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause169)).
% 13.73/13.99  cnf(c_0_5, negated_conjecture, (ssItem(sk5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
% 13.73/13.99  cnf(c_0_6, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause52)).
% 13.73/13.99  cnf(c_0_7, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause53)).
% 13.73/13.99  cnf(c_0_8, axiom, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause99)).
% 13.73/13.99  cnf(c_0_9, axiom, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause104)).
% 13.73/13.99  cnf(c_0_10, axiom, (ssList(tl(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause75)).
% 13.73/13.99  cnf(c_0_11, axiom, (ssItem(hd(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause76)).
% 13.73/13.99  cnf(c_0_12, axiom, (singletonP(X2)|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause116)).
% 13.73/13.99  cnf(c_0_13, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause72)).
% 13.73/13.99  cnf(c_0_14, axiom, (nil=X2|tl(app(X2,X1))=app(tl(X2),X1)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause133)).
% 13.73/13.99  cnf(c_0_15, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
% 13.73/13.99  cnf(c_0_16, axiom, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause118)).
% 13.73/13.99  cnf(c_0_17, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause96)).
% 13.73/13.99  cnf(c_0_18, axiom, (cons(skaf44(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause101)).
% 13.73/13.99  cnf(c_0_19, axiom, (ssItem(skaf44(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause47)).
% 13.73/13.99  cnf(c_0_20, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 13.73/13.99  cnf(c_0_21, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause86)).
% 13.73/13.99  cnf(c_0_22, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause73)).
% 13.73/13.99  cnf(c_0_23, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause98)).
% 13.73/13.99  cnf(c_0_24, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause120)).
% 13.73/13.99  cnf(c_0_25, negated_conjecture, (app(X2,cons(X1,nil))=sk3|~ssItem(X1)|~ssList(X2)|app(cons(X1,nil),X2)!=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 13.73/13.99  cnf(c_0_26, axiom, (rearsegP(X3,X2)|app(X1,X2)!=X3|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause143)).
% 13.73/13.99  cnf(c_0_27, axiom, (cons(skaf83(X1),skaf82(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause109)).
% 13.73/13.99  cnf(c_0_28, axiom, (ssList(skaf82(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause13)).
% 13.73/13.99  cnf(c_0_29, axiom, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause97)).
% 13.73/13.99  cnf(c_0_30, axiom, (rearsegP(nil,X1)|nil!=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause81)).
% 13.73/13.99  cnf(c_0_31, axiom, (rearsegP(X1,X3)|~rearsegP(X1,X2)|~rearsegP(X2,X3)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause153)).
% 13.73/13.99  cnf(c_0_32, axiom, (rearsegP(app(X3,X1),X2)|~rearsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause136)).
% 13.73/13.99  cnf(c_0_33, axiom, (rearsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause58)).
% 13.73/13.99  cnf(c_0_34, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 13.73/13.99  cnf(c_0_35, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 13.73/13.99  cnf(c_0_36, axiom, (app(skaf46(X1,X2),X2)=X1|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause131)).
% 13.73/13.99  cnf(c_0_37, axiom, (ssList(skaf46(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause50)).
% 13.73/13.99  cnf(c_0_38, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause149)).
% 13.73/13.99  cnf(c_0_39, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause74)).
% 13.73/13.99  cnf(c_0_40, axiom, (frontsegP(app(X1,X3),X2)|~frontsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause137)).
% 13.73/13.99  cnf(c_0_41, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause60)).
% 13.73/13.99  cnf(c_0_42, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause144)).
% 13.73/13.99  cnf(c_0_43, axiom, (nil=X2|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause119)).
% 13.73/13.99  cnf(c_0_44, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause145)).
% 13.73/13.99  cnf(c_0_45, axiom, (X4=X2|cons(X1,X2)!=cons(X3,X4)|~ssItem(X3)|~ssItem(X1)|~ssList(X4)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause171)).
% 13.73/13.99  cnf(c_0_46, axiom, (nil=X2|hd(app(X2,X1))=hd(X2)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause123)).
% 13.73/13.99  cnf(c_0_47, axiom, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause157)).
% 13.73/13.99  cnf(c_0_48, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause132)).
% 13.73/13.99  cnf(c_0_49, axiom, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause138)).
% 13.73/13.99  cnf(c_0_50, axiom, (X1=X3|app(X1,X2)!=app(X3,X2)|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause151)).
% 13.73/13.99  cnf(c_0_51, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause51)).
% 13.73/13.99  cnf(c_0_52, axiom, (memberP(X4,X2)|app(X1,cons(X2,X3))!=X4|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause175)).
% 13.73/13.99  cnf(c_0_53, axiom, (memberP(X2,X3)|X3=X1|~memberP(cons(X1,X2),X3)|~ssList(X2)|~ssItem(X1)|~ssItem(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause161)).
% 13.73/13.99  cnf(c_0_54, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause71)).
% 13.73/13.99  cnf(c_0_55, axiom, (memberP(app(X1,X3),X2)|~memberP(X1,X2)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause140)).
% 13.73/13.99  cnf(c_0_56, negated_conjecture, (~memberP(sk1,sk5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
% 13.73/13.99  cnf(c_0_57, negated_conjecture, (memberP(sk4,sk5)), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 13.73/13.99  cnf(c_0_58, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_2, c_0_1])).
% 13.73/13.99  cnf(c_0_59, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_3])).
% 13.73/13.99  cnf(c_0_60, negated_conjecture, (app(skaf42(sk4,sk5),cons(sk5,skaf43(sk5,sk4)))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_57]), c_0_5]), c_0_58])])).
% 13.73/13.99  cnf(c_0_61, negated_conjecture, (~ssList(app(sk4,cons(sk5,X1)))|~ssList(X1)|~duplicatefreeP(app(sk4,cons(sk5,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_5]), c_0_6]), c_0_7])])).
% 13.73/13.99  cnf(c_0_62, plain, (nil=X1|tl(X1)!=X1|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10]), c_0_11])).
% 13.73/13.99  cnf(c_0_63, plain, (singletonP(cons(X1,nil))|~ssItem(X1)|~ssList(cons(X1,nil))), inference(er,[status(thm)],[c_0_12])).
% 13.73/13.99  cnf(c_0_64, negated_conjecture, (ssItem(X1)|~ssList(app(sk4,cons(sk5,X2)))|~ssList(X2)), inference(spm,[status(thm)],[c_0_61, c_0_13])).
% 13.73/13.99  cnf(c_0_65, plain, (nil=X1|app(tl(X1),X2)!=app(X1,X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_14]), c_0_15]), c_0_16])).
% 13.73/13.99  cnf(c_0_66, plain, (tl(X1)=nil|~singletonP(X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])])).
% 13.73/13.99  cnf(c_0_67, plain, (singletonP(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_21]), c_0_20])])).
% 13.73/13.99  cnf(c_0_68, negated_conjecture, (ssItem(X1)|~ssList(cons(sk5,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_15]), c_0_58])])).
% 13.73/13.99  cnf(c_0_69, plain, (nil=X1|tl(X1)!=app(X1,nil)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_22]), c_0_20])]), c_0_10])).
% 13.73/13.99  cnf(c_0_70, plain, (tl(cons(X1,nil))=nil|~ssItem(X1)|~ssList(cons(X1,nil))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
% 13.73/13.99  cnf(c_0_71, negated_conjecture, (ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_21]), c_0_5])])).
% 13.73/13.99  cnf(c_0_72, plain, (app(cons(X1,X2),nil)!=X2|~ssItem(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_17]), c_0_21]), c_0_23])).
% 13.73/13.99  cnf(c_0_73, plain, (tl(cons(X1,nil))=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_21]), c_0_20])])).
% 13.73/13.99  cnf(c_0_74, negated_conjecture, (ssItem(X1)), inference(spm,[status(thm)],[c_0_71, c_0_7])).
% 13.73/13.99  cnf(c_0_75, plain, (cons(X1,nil)!=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_24]), c_0_20])])).
% 13.73/13.99  cnf(c_0_76, plain, (tl(cons(X1,nil))=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_74])])).
% 13.73/13.99  cnf(c_0_77, plain, (cons(X1,nil)!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_74])])).
% 13.73/13.99  cnf(c_0_78, negated_conjecture, (app(X1,cons(X2,nil))=sk3|app(cons(X2,nil),X1)!=sk4|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_74])])).
% 13.73/13.99  cnf(c_0_79, plain, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_74])])).
% 13.73/13.99  cnf(c_0_80, plain, (cons(hd(cons(X1,nil)),nil)=cons(X1,nil)|~ssList(cons(X1,nil))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_76]), c_0_77])).
% 13.73/13.99  cnf(c_0_81, plain, (ssList(cons(X1,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_74])])).
% 13.73/13.99  cnf(c_0_82, negated_conjecture, (app(X1,cons(X2,nil))=sk3|cons(X2,X1)!=sk4|~ssList(X1)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
% 13.73/13.99  cnf(c_0_83, plain, (cons(hd(cons(X1,nil)),nil)=cons(X1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_20])])).
% 13.73/13.99  cnf(c_0_84, plain, (rearsegP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_26]), c_0_15])).
% 13.73/13.99  cnf(c_0_85, negated_conjecture, (app(skaf82(sk4),cons(skaf83(sk4),nil))=sk3|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_27]), c_0_28])])]), c_0_58])])).
% 13.73/13.99  cnf(c_0_86, plain, (ssList(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_83]), c_0_20])])).
% 13.73/13.99  cnf(c_0_87, plain, (hd(cons(X1,X2))=X1|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_74])])).
% 13.73/13.99  cnf(c_0_88, plain, (rearsegP(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_30]), c_0_20])])).
% 13.73/13.99  cnf(c_0_89, plain, (rearsegP(X1,X2)|~rearsegP(X1,app(X3,X2))|~ssList(X2)|~ssList(X1)|~ssList(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_84]), c_0_15])).
% 13.73/13.99  cnf(c_0_90, negated_conjecture, (sk4=nil|rearsegP(sk3,cons(skaf83(sk4),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86]), c_0_28])])).
% 13.73/13.99  cnf(c_0_91, plain, (skaf83(X1)=hd(X1)|nil=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_27]), c_0_28])])).
% 13.73/13.99  cnf(c_0_92, plain, (rearsegP(app(X1,X2),nil)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_20])])).
% 13.73/13.99  cnf(c_0_93, plain, (rearsegP(app(X1,nil),nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_88]), c_0_20])])).
% 13.73/13.99  cnf(c_0_94, plain, (rearsegP(X1,X2)|~rearsegP(X1,cons(X3,X2))|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_79]), c_0_86])])).
% 13.73/13.99  cnf(c_0_95, negated_conjecture, (sk4=nil|rearsegP(sk3,cons(hd(sk4),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_58])])).
% 13.73/13.99  cnf(c_0_96, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_34, c_0_35])).
% 13.73/13.99  cnf(c_0_97, negated_conjecture, (rearsegP(sk4,nil)|~ssList(cons(sk5,skaf43(sk5,sk4)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_60]), c_0_7])])).
% 13.73/13.99  cnf(c_0_98, plain, (rearsegP(cons(X1,nil),nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_79]), c_0_86]), c_0_20])])).
% 13.73/13.99  cnf(c_0_99, plain, (app(skaf46(X1,nil),nil)=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_33]), c_0_20])])).
% 13.73/13.99  cnf(c_0_100, negated_conjecture, (sk4=nil|rearsegP(sk3,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_20]), c_0_96])])).
% 13.73/13.99  cnf(c_0_101, negated_conjecture, (rearsegP(sk4,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_81]), c_0_6])])).
% 13.73/13.99  cnf(c_0_102, plain, (app(skaf46(cons(X1,nil),nil),nil)=cons(X1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_98]), c_0_20]), c_0_86])])).
% 13.73/13.99  cnf(c_0_103, plain, (skaf46(X1,nil)=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_99]), c_0_37])])).
% 13.73/13.99  cnf(c_0_104, negated_conjecture, (app(skaf46(sk3,nil),nil)=sk3|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_100]), c_0_20]), c_0_96])])).
% 13.73/13.99  cnf(c_0_105, negated_conjecture, (app(skaf46(sk4,nil),nil)=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_101]), c_0_20]), c_0_58])])).
% 13.73/13.99  cnf(c_0_106, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_20])])).
% 13.73/13.99  cnf(c_0_107, plain, (app(cons(X1,nil),nil)=cons(X1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_86])])).
% 13.73/13.99  cnf(c_0_108, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_20])])).
% 13.73/13.99  cnf(c_0_109, negated_conjecture, (app(sk3,nil)=sk3|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_103]), c_0_96])])).
% 13.73/13.99  cnf(c_0_110, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_42]), c_0_15])).
% 13.73/13.99  cnf(c_0_111, negated_conjecture, (app(sk4,nil)=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_103]), c_0_58])])).
% 13.73/13.99  cnf(c_0_112, plain, (nil=X1|app(X2,app(X3,X1))!=nil|~ssList(X1)|~ssList(X2)|~ssList(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_38]), c_0_15])).
% 13.73/13.99  cnf(c_0_113, plain, (app(nil,cons(X1,nil))=cons(X1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_86]), c_0_20])])).
% 13.73/13.99  cnf(c_0_114, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_44])]), c_0_20])])).
% 13.73/13.99  cnf(c_0_115, negated_conjecture, (sk4=nil|frontsegP(sk3,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_20]), c_0_96])])).
% 13.73/13.99  cnf(c_0_116, negated_conjecture, (sk4=nil|frontsegP(sk3,sk3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_109]), c_0_20]), c_0_96])])).
% 13.73/13.99  cnf(c_0_117, plain, (X1=X2|cons(X3,X1)!=cons(X4,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_74]), c_0_74])])).
% 13.73/13.99  cnf(c_0_118, plain, (cons(X1,X2)!=nil|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_74])])).
% 13.73/13.99  cnf(c_0_119, plain, (cons(hd(X1),tl(app(X1,X2)))=app(X1,X2)|nil=X1|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_46]), c_0_15]), c_0_16])).
% 13.73/13.99  cnf(c_0_120, plain, (nil=X1|ssList(app(tl(X1),X2))|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_14]), c_0_15]), c_0_16])).
% 13.73/13.99  cnf(c_0_121, negated_conjecture, (app(tl(sk4),nil)=tl(sk4)|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_111]), c_0_58]), c_0_20])])).
% 13.73/13.99  cnf(c_0_122, plain, (app(X1,X2)=nil|app(X1,app(X2,X3))!=nil|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_38]), c_0_15])).
% 13.73/13.99  cnf(c_0_123, plain, (app(X1,cons(X2,nil))!=nil|~ssList(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_86]), c_0_20])]), c_0_77])).
% 13.73/13.99  cnf(c_0_124, plain, (app(cons(X1,X2),X3)=cons(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_74])])).
% 13.73/13.99  cnf(c_0_125, plain, (app(nil,app(nil,X1))=app(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_114]), c_0_20])])).
% 13.73/13.99  cnf(c_0_126, negated_conjecture, (app(nil,skaf45(sk3,nil))=sk3|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_115]), c_0_20]), c_0_96])])).
% 13.73/13.99  cnf(c_0_127, plain, (ssList(app(X1,app(X2,X3)))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_38]), c_0_15])).
% 13.73/13.99  cnf(c_0_128, negated_conjecture, (app(sk3,skaf45(sk3,sk3))=sk3|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_116]), c_0_96])])).
% 13.73/13.99  cnf(c_0_129, plain, (skaf82(cons(X1,X2))=X2|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_27]), c_0_28])])]), c_0_81]), c_0_118])).
% 13.73/13.99  cnf(c_0_130, negated_conjecture, (cons(hd(sk4),tl(sk4))=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_111]), c_0_58]), c_0_20])])).
% 13.73/13.99  cnf(c_0_131, negated_conjecture, (sk4=nil|ssList(tl(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_58]), c_0_20])])).
% 13.73/13.99  cnf(c_0_132, plain, (app(X1,cons(X2,X3))!=nil|~ssList(X3)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_79]), c_0_86])]), c_0_123])).
% 13.73/13.99  cnf(c_0_133, plain, (memberP(cons(X1,X2),X1)|~ssItem(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_49])).
% 13.73/13.99  cnf(c_0_134, plain, (X1=nil|app(X1,nil)!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_114]), c_0_20])])).
% 13.73/13.99  cnf(c_0_135, plain, (cons(hd(app(X1,X2)),app(tl(X1),X2))=app(X1,X2)|nil=X1|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_14]), c_0_15]), c_0_16])).
% 13.73/13.99  cnf(c_0_136, plain, (tl(cons(X1,X2))=X2|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_74])])).
% 13.73/13.99  cnf(c_0_137, plain, (tl(cons(X1,app(X2,X3)))=app(tl(cons(X1,X2)),X3)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_124]), c_0_81]), c_0_118])).
% 13.73/13.99  cnf(c_0_138, negated_conjecture, (app(nil,sk3)=sk3|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_51])])).
% 13.73/13.99  cnf(c_0_139, negated_conjecture, (sk4=nil|ssList(app(X1,sk3))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_51]), c_0_96])])).
% 13.73/13.99  cnf(c_0_140, negated_conjecture, (sk4=nil|app(X1,sk3)!=nil|~ssList(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_85]), c_0_86]), c_0_28])]), c_0_77])).
% 13.73/13.99  cnf(c_0_141, negated_conjecture, (skaf82(sk4)=tl(sk4)|sk4=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_131])).
% 13.73/13.99  cnf(c_0_142, negated_conjecture, (sk4!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_60]), c_0_6]), c_0_7])])).
% 13.73/13.99  cnf(c_0_143, plain, (memberP(cons(X1,X2),X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_74])])).
% 13.73/13.99  cnf(c_0_144, plain, (cons(X1,app(X2,nil))!=nil|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_124]), c_0_20])]), c_0_81]), c_0_118])).
% 13.73/13.99  cnf(c_0_145, plain, (cons(hd(app(cons(X1,X2),X3)),app(X2,X3))=app(cons(X1,X2),X3)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_136]), c_0_81]), c_0_118])).
% 13.73/13.99  cnf(c_0_146, plain, (cons(X1,tl(cons(X1,X2)))=cons(X1,X2)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_87]), c_0_81]), c_0_118])).
% 13.73/13.99  cnf(c_0_147, negated_conjecture, (tl(cons(X1,sk3))=app(nil,sk3)|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_76]), c_0_96]), c_0_20])])).
% 13.73/13.99  cnf(c_0_148, negated_conjecture, (sk4=nil|ssList(cons(X1,sk3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_79]), c_0_86]), c_0_96])])).
% 13.73/13.99  cnf(c_0_149, negated_conjecture, (sk4=nil|cons(X1,sk3)!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_79]), c_0_86]), c_0_96])])).
% 13.73/13.99  cnf(c_0_150, plain, (memberP(app(X1,cons(X2,X3)),X2)|~ssItem(X2)|~ssList(app(X1,cons(X2,X3)))|~ssList(X3)|~ssList(X1)), inference(er,[status(thm)],[c_0_52])).
% 13.73/13.99  cnf(c_0_151, negated_conjecture, (app(skaf82(sk4),cons(hd(sk4),nil))=sk3|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_91]), c_0_58])])).
% 13.73/13.99  cnf(c_0_152, negated_conjecture, (skaf82(sk4)=tl(sk4)), inference(sr,[status(thm)],[c_0_141, c_0_142])).
% 13.73/13.99  cnf(c_0_153, plain, (X1=X2|memberP(X3,X2)|~memberP(cons(X1,X3),X2)|~ssList(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_74]), c_0_74])])).
% 13.73/13.99  cnf(c_0_154, plain, (memberP(cons(X1,nil),hd(cons(X1,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_83]), c_0_20])])).
% 13.73/13.99  cnf(c_0_155, plain, (~memberP(nil,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_74])])).
% 13.73/13.99  cnf(c_0_156, plain, (app(cons(X1,X2),nil)!=nil|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_20])])).
% 13.73/13.99  cnf(c_0_157, negated_conjecture, (cons(X1,app(nil,sk3))=cons(X1,sk3)|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_96])])).
% 13.73/13.99  cnf(c_0_158, negated_conjecture, (sk4=nil|ssList(app(nil,sk3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_147]), c_0_148]), c_0_149])).
% 13.73/13.99  cnf(c_0_159, plain, (memberP(app(X1,cons(X2,X3)),X2)|~ssList(app(X1,cons(X2,X3)))|~ssList(X3)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_150, c_0_74])])).
% 13.73/13.99  cnf(c_0_160, plain, (cons(hd(X1),app(tl(X1),X2))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_9]), c_0_10])).
% 13.73/13.99  cnf(c_0_161, negated_conjecture, (app(tl(sk4),cons(hd(sk4),nil))=sk3), inference(rw,[status(thm)],[inference(sr,[status(thm)],[c_0_151, c_0_142]), c_0_152])).
% 13.73/13.99  cnf(c_0_162, plain, (hd(cons(X1,app(X2,X3)))=hd(cons(X1,X2))|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_124]), c_0_81]), c_0_118])).
% 13.73/13.99  cnf(c_0_163, plain, (hd(cons(X1,nil))=X1), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154]), c_0_20])]), c_0_155])).
% 13.73/13.99  cnf(c_0_164, negated_conjecture, (sk4=nil|app(cons(X1,sk3),nil)!=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_158])).
% 13.73/13.99  cnf(c_0_165, negated_conjecture, (sk4=nil|memberP(sk3,skaf83(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_85]), c_0_96]), c_0_20]), c_0_28])])).
% 13.73/13.99  cnf(c_0_166, negated_conjecture, (memberP(app(sk4,X1),sk5)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_57]), c_0_5]), c_0_58])])).
% 13.73/13.99  cnf(c_0_167, negated_conjecture, (app(sk4,cons(hd(sk4),nil))=cons(hd(sk4),sk3)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_161]), c_0_86]), c_0_58])]), c_0_142])).
% 13.73/13.99  cnf(c_0_168, negated_conjecture, (hd(cons(X1,sk3))=X1|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_138]), c_0_163]), c_0_96]), c_0_20])])).
% 13.73/13.99  cnf(c_0_169, negated_conjecture, (app(nil,sk3)=sk3), inference(sr,[status(thm)],[c_0_138, c_0_142])).
% 13.73/13.99  cnf(c_0_170, negated_conjecture, (sk4=nil|cons(X1,app(sk3,nil))!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_124]), c_0_20]), c_0_96])])).
% 13.73/13.99  cnf(c_0_171, negated_conjecture, (app(sk3,nil)=sk3), inference(sr,[status(thm)],[c_0_109, c_0_142])).
% 13.73/13.99  cnf(c_0_172, negated_conjecture, (sk4=nil|memberP(sk3,hd(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_91]), c_0_58])])).
% 13.73/13.99  cnf(c_0_173, plain, (hd(X1)=X2|nil=X1|memberP(tl(X1),X2)|~memberP(X1,X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_9]), c_0_10])).
% 13.73/13.99  cnf(c_0_174, negated_conjecture, (memberP(cons(hd(sk4),sk3),sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_167]), c_0_86])])).
% 13.73/13.99  cnf(c_0_175, negated_conjecture, (hd(cons(X1,sk3))=X1), inference(sr,[status(thm)],[c_0_168, c_0_142])).
% 13.73/13.99  cnf(c_0_176, negated_conjecture, (tl(cons(X1,sk3))=sk3), inference(rw,[status(thm)],[inference(sr,[status(thm)],[c_0_147, c_0_142]), c_0_169])).
% 13.73/13.99  cnf(c_0_177, negated_conjecture, (ssList(cons(X1,sk3))), inference(sr,[status(thm)],[c_0_148, c_0_142])).
% 13.73/13.99  cnf(c_0_178, negated_conjecture, (cons(X1,sk3)!=nil), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_170, c_0_171]), c_0_142])).
% 13.73/13.99  cnf(c_0_179, negated_conjecture, (~memberP(sk3,sk5)), inference(rw,[status(thm)],[c_0_56, c_0_35])).
% 13.73/13.99  cnf(c_0_180, negated_conjecture, (memberP(sk3,hd(sk4))), inference(sr,[status(thm)],[c_0_172, c_0_142])).
% 13.73/13.99  cnf(c_0_181, negated_conjecture, (hd(sk4)=sk5), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174]), c_0_175]), c_0_176]), c_0_177])]), c_0_178]), c_0_179])).
% 13.73/13.99  cnf(c_0_182, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_180, c_0_181]), c_0_179]), ['proof']).
% 13.73/13.99  # SZS output end CNFRefutation
% 13.73/13.99  # Training examples: 0 positive, 0 negative
% 13.82/14.02  
% 13.82/14.02  # -------------------------------------------------
% 13.82/14.02  # User time                : 13.431 s
% 13.82/14.02  # System time              : 0.375 s
% 13.82/14.02  # Total time               : 13.806 s
% 13.82/14.02  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
