<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC355-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC355-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC355-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC355-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC355-1.p">SWC355-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n161.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:37 EDT 2016

% Result     : Unsatisfiable 0.07s
% Output     : CNFRefutation 0.07s
% Verified   : 
% Statistics : Number of clauses        :   42 (  93 expanded)
%              Number of leaves         :   20 (  59 expanded)
%              Depth                    :   11
%              Number of atoms          :  101 ( 172 expanded)
%              Number of equality atoms :   16 (  36 expanded)
%              Maximal clause size      :    7 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( neq(sk2,nil)
    | neq(sk2,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_1"></A>cnf(c_0_1,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( app(sk3,cons(sk5,nil)) = sk4
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( ssItem(sk5)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( frontsegP(X3,X1)
    | app(X1,X2) != X3
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause144)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( X2 = X3
    | app(X1,X2) != app(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause150)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( ~ segmentP(sk2,sk1)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_14)).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( neq(sk2,nil) ),
    inference(cn,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( app(sk1,cons(sk5,nil)) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_3]),c_0_4]),c_0_4]),c_0_20])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( ssItem(sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,c_0_4]),c_0_20])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( ~ ssList(app(sk2,cons(sk5,X1)))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(sk2,cons(sk5,X1))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_22]),c_0_23]),c_0_6]),c_0_7])])).

<A NAME="c_0_25"></A>cnf(c_0_25,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_8]),<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(app(sk2,cons(sk5,X2)))
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( frontsegP(sk2,sk1)
    | ~ ssList(cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_22]),c_0_7])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(cons(sk5,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_9]),c_0_11])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( frontsegP(sk2,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_12]),c_0_23]),c_0_6])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_12]),c_0_23])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( X1 = cons(sk5,nil)
    | app(sk1,X1) != sk2
    | ~ ssList(cons(sk5,nil))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_22]),c_0_7])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( app(sk1,skaf45(sk2,sk1)) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_29]),c_0_7]),c_0_11])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( skaf45(sk2,sk1) = cons(sk5,nil)
    | ~ ssList(cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_32]),c_0_16])])).

<A NAME="c_0_35"></A>cnf(c_0_35,plain,
    ( ssList(cons(X1,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_33])])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( skaf45(sk2,sk1) = cons(sk5,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_35]),c_0_6])])).

<A NAME="c_0_38"></A>cnf(c_0_38,plain,
    ( segmentP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_18]),c_0_6])]),<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( ssList(cons(sk5,nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,<A HREF="#c_0_37">c_0_37</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( ~ segmentP(sk2,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_4]),c_0_20])])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_22]),c_0_39]),c_0_7])]),<A HREF="#c_0_40">c_0_40</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC355-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n161.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:42:37 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.07/0.34  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.07/0.34  # and selection function SelectCQArNTEqFirst.
% 0.07/0.34  #
% 0.07/0.34  # Preprocessing time       : 0.009 s
% 0.07/0.34  # Presaturation interreduction done
% 0.07/0.34  
% 0.07/0.34  # Proof found!
% 0.07/0.34  # SZS status Unsatisfiable
% 0.07/0.34  # SZS output start CNFRefutation
% 0.07/0.34  cnf(c_0_0, negated_conjecture, (neq(sk2,nil)|neq(sk2,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 0.07/0.34  cnf(c_0_1, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause179)).
% 0.07/0.34  cnf(c_0_2, negated_conjecture, (app(sk3,cons(sk5,nil))=sk4|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_13)).
% 0.07/0.34  cnf(c_0_3, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 0.07/0.34  cnf(c_0_4, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 0.07/0.34  cnf(c_0_5, negated_conjecture, (ssItem(sk5)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_12)).
% 0.07/0.34  cnf(c_0_6, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.07/0.34  cnf(c_0_7, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 0.07/0.34  cnf(c_0_8, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause144)).
% 0.07/0.34  cnf(c_0_9, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.07/0.34  cnf(c_0_10, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause72)).
% 0.07/0.34  cnf(c_0_11, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
% 0.07/0.34  cnf(c_0_12, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.07/0.34  cnf(c_0_13, axiom, (X2=X3|app(X1,X2)!=app(X1,X3)|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause150)).
% 0.07/0.34  cnf(c_0_14, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause132)).
% 0.07/0.34  cnf(c_0_15, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause53)).
% 0.07/0.34  cnf(c_0_16, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause51)).
% 0.07/0.34  cnf(c_0_17, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause173)).
% 0.07/0.34  cnf(c_0_18, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause74)).
% 0.07/0.34  cnf(c_0_19, negated_conjecture, (~segmentP(sk2,sk1)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_14)).
% 0.07/0.34  cnf(c_0_20, negated_conjecture, (neq(sk2,nil)), inference(cn,[status(thm)],[c_0_0])).
% 0.07/0.34  cnf(c_0_21, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_1])).
% 0.07/0.34  cnf(c_0_22, negated_conjecture, (app(sk1,cons(sk5,nil))=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_2, c_0_3]), c_0_4]), c_0_4]), c_0_20])])).
% 0.07/0.34  cnf(c_0_23, negated_conjecture, (ssItem(sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_4]), c_0_20])])).
% 0.07/0.34  cnf(c_0_24, negated_conjecture, (~ssList(app(sk2,cons(sk5,X1)))|~ssList(X1)|~duplicatefreeP(app(sk2,cons(sk5,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_6]), c_0_7])])).
% 0.07/0.34  cnf(c_0_25, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_8]), c_0_9])).
% 0.07/0.34  cnf(c_0_26, negated_conjecture, (ssItem(X1)|~ssList(app(sk2,cons(sk5,X2)))|~ssList(X2)), inference(spm,[status(thm)],[c_0_24, c_0_10])).
% 0.07/0.34  cnf(c_0_27, negated_conjecture, (frontsegP(sk2,sk1)|~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_22]), c_0_7])])).
% 0.07/0.34  cnf(c_0_28, negated_conjecture, (ssItem(X1)|~ssList(cons(sk5,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_9]), c_0_11])])).
% 0.07/0.34  cnf(c_0_29, negated_conjecture, (frontsegP(sk2,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_12]), c_0_23]), c_0_6])])).
% 0.07/0.34  cnf(c_0_30, negated_conjecture, (ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_12]), c_0_23])])).
% 0.07/0.34  cnf(c_0_31, negated_conjecture, (X1=cons(sk5,nil)|app(sk1,X1)!=sk2|~ssList(cons(sk5,nil))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_22]), c_0_7])])).
% 0.07/0.34  cnf(c_0_32, negated_conjecture, (app(sk1,skaf45(sk2,sk1))=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_29]), c_0_7]), c_0_11])])).
% 0.07/0.34  cnf(c_0_33, negated_conjecture, (ssItem(X1)), inference(spm,[status(thm)],[c_0_30, c_0_15])).
% 0.07/0.34  cnf(c_0_34, negated_conjecture, (skaf45(sk2,sk1)=cons(sk5,nil)|~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_16])])).
% 0.07/0.34  cnf(c_0_35, plain, (ssList(cons(X1,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_33])])).
% 0.07/0.34  cnf(c_0_36, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(er,[status(thm)],[c_0_17])).
% 0.07/0.34  cnf(c_0_37, negated_conjecture, (skaf45(sk2,sk1)=cons(sk5,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_6])])).
% 0.07/0.34  cnf(c_0_38, plain, (segmentP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_18]), c_0_6])]), c_0_9])).
% 0.07/0.34  cnf(c_0_39, negated_conjecture, (ssList(cons(sk5,nil))), inference(spm,[status(thm)],[c_0_16, c_0_37])).
% 0.07/0.34  cnf(c_0_40, negated_conjecture, (~segmentP(sk2,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_4]), c_0_20])])).
% 0.07/0.34  cnf(c_0_41, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_22]), c_0_39]), c_0_7])]), c_0_40]), ['proof']).
% 0.07/0.34  # SZS output end CNFRefutation
% 0.07/0.34  # Training examples: 0 positive, 0 negative
% 0.07/0.34  
% 0.07/0.34  # -------------------------------------------------
% 0.07/0.34  # User time                : 0.111 s
% 0.07/0.34  # System time              : 0.003 s
% 0.07/0.34  # Total time               : 0.114 s
% 0.07/0.34  # Maximum resident set size: 1864 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
