<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC099-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC099-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC099-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC099-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC099-1.p">SWC099-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n137.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:39 EDT 2016

% Result     : Unsatisfiable 1.62s
% Output     : CNFRefutation 1.62s
% Verified   : 
% Statistics : Number of clauses        :  130 (1175 expanded)
%              Number of leaves         :   57 ( 782 expanded)
%              Depth                    :   28
%              Number of atoms          :  339 (2426 expanded)
%              Number of equality atoms :   98 ( 704 expanded)
%              Maximal clause size      :    7 (   3 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( frontsegP(sk4,sk3) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( nil = sk2
    | ~ neq(sk1,nil)
    | ~ frontsegP(sk2,sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_4"></A>cnf(c_0_4,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( nil = X1
    | ~ frontsegP(nil,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause84)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( rearsegP(X3,X2)
    | app(X1,X2) != X3
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause143)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( app(app(skaf47(X1,X2),X2),skaf48(X2,X1)) = X1
    | ~ segmentP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause168)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( ssList(skaf48(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause48)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( ssList(skaf47(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause49)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause99)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause104)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause75)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause76)).

<A NAME="c_0_23"></A>cnf(c_0_23,axiom,
    ( rearsegP(app(X3,X1),X2)
    | ~ rearsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause136)).

<A NAME="c_0_24"></A>cnf(c_0_24,axiom,
    ( rearsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause58)).

<A NAME="c_0_25"></A>cnf(c_0_25,axiom,
    ( app(skaf46(X1,X2),X2) = X1
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause131)).

<A NAME="c_0_26"></A>cnf(c_0_26,axiom,
    ( nil = X2
    | tl(app(X2,X1)) = app(tl(X2),X1)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause133)).

<A NAME="c_0_27"></A>cnf(c_0_27,axiom,
    ( nil = X1
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause118)).

<A NAME="c_0_28"></A>cnf(c_0_28,axiom,
    ( ssList(skaf46(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause50)).

<A NAME="c_0_29"></A>cnf(c_0_29,axiom,
    ( memberP(cons(X2,X3),X1)
    | X1 != X2
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause138)).

<A NAME="c_0_30"></A>cnf(c_0_30,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_31"></A>cnf(c_0_31,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_32"></A>cnf(c_0_32,axiom,
    ( cons(X1,X2) != nil
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause98)).

<A NAME="c_0_33"></A>cnf(c_0_33,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_34"></A>cnf(c_0_34,axiom,
    ( cons(skaf83(X1),skaf82(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause109)).

<A NAME="c_0_35"></A>cnf(c_0_35,axiom,
    ( ssItem(skaf83(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause12)).

<A NAME="c_0_36"></A>cnf(c_0_36,axiom,
    ( ssList(skaf82(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause13)).

<A NAME="c_0_37"></A>cnf(c_0_37,axiom,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause157)).

<A NAME="c_0_38"></A>cnf(c_0_38,axiom,
    ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
    | ~ memberP(X1,X2)
    | ~ ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause169)).

<A NAME="c_0_39"></A>cnf(c_0_39,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_40"></A>cnf(c_0_40,axiom,
    ( ssList(skaf43(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause52)).

<A NAME="c_0_41"></A>cnf(c_0_41,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_42"></A>cnf(c_0_42,axiom,
    ( nil = X2
    | hd(app(X2,X1)) = hd(X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause123)).

<A NAME="c_0_43"></A>cnf(c_0_43,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_44"></A>cnf(c_0_44,axiom,
    ( X4 = X2
    | cons(X1,X2) != cons(X3,X4)
    | ~ ssItem(X3)
    | ~ ssItem(X1)
    | ~ ssList(X4)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause171)).

<A NAME="c_0_45"></A>cnf(c_0_45,axiom,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause60)).

<A NAME="c_0_46"></A>cnf(c_0_46,axiom,
    ( app(X1,X2) = nil
    | nil != X1
    | nil != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause145)).

<A NAME="c_0_47"></A>cnf(c_0_47,axiom,
    ( frontsegP(cons(X3,X1),cons(X4,X2))
    | ~ frontsegP(X1,X2)
    | X3 != X4
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssItem(X4)
    | ~ ssItem(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause178)).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( ~ ssList(X1)
    | ~ neq(sk3,X1)
    | ~ frontsegP(sk4,X1)
    | ~ segmentP(X1,sk3)
    | ~ equalelemsP(X1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_49"></A>cnf(c_0_49,axiom,
    ( ssItem(skaf44(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause47)).

<A NAME="c_0_50"></A>cnf(c_0_50,axiom,
    ( segmentP(X1,nil)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause56)).

<A NAME="c_0_51"></A>cnf(c_0_51,axiom,
    ( equalelemsP(cons(X1,nil))
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause64)).

<A NAME="c_0_52"></A>cnf(c_0_52,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( equalelemsP(sk3) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_54"></A>cnf(c_0_54,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssItem(X1)
    | ~ ssItem(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause102)).

<A NAME="c_0_55"></A>cnf(c_0_55,axiom,
    ( segmentP(nil,X1)
    | nil != X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause79)).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( neq(sk2,nil)
    | nil != sk1 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( frontsegP(sk2,sk1) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,c_0_1]),<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( sk2 = nil
    | ~ neq(sk1,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>,c_0_57])])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( sk1 = nil
    | sk2 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_4]),c_0_5]),c_0_6])])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( sk1 = nil
    | frontsegP(nil,sk1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,<A HREF="#c_0_59">c_0_59</A>])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_60]),c_0_5])])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( frontsegP(sk2,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,<A HREF="#c_0_61">c_0_61</A>])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( app(nil,skaf45(sk2,nil)) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_62]),c_0_6]),c_0_9])])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( skaf45(sk2,nil) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_63]),c_0_11])])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,negated_conjecture,
    ( app(nil,sk2) = sk2 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>,<A HREF="#c_0_64">c_0_64</A>])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( segmentP(app(sk2,X1),sk2)
    | ~ ssList(app(sk2,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_66]),c_0_9]),c_0_6])])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( segmentP(sk2,sk2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,c_0_13]),c_0_9]),c_0_6])])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_14]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( app(app(skaf47(sk2,sk2),sk2),skaf48(sk2,sk2)) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_68]),c_0_9])])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( rearsegP(sk2,skaf48(sk2,sk2))
    | ~ ssList(app(skaf47(sk2,sk2),sk2)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>,c_0_70]),c_0_17])])).

<A NAME="c_0_72"></A>cnf(c_0_72,negated_conjecture,
    ( rearsegP(sk2,skaf48(sk2,sk2)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>,c_0_15]),c_0_18]),c_0_9])])).

<A NAME="c_0_73"></A>cnf(c_0_73,plain,
    ( nil = X1
    | tl(X1) != X1
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_20]),c_0_21]),<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_74"></A>cnf(c_0_74,plain,
    ( rearsegP(app(X1,X2),nil)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_24]),c_0_6])])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( app(skaf46(sk2,skaf48(sk2,sk2)),skaf48(sk2,sk2)) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_72]),c_0_17]),c_0_9])])).

<A NAME="c_0_76"></A>cnf(c_0_76,plain,
    ( nil = X1
    | app(tl(X1),X2) != app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>,c_0_26]),c_0_15]),<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,negated_conjecture,
    ( rearsegP(sk2,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,c_0_75]),c_0_28]),c_0_17])])).

<A NAME="c_0_78"></A>cnf(c_0_78,plain,
    ( nil = X1
    | tl(X1) != app(X1,nil)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>,c_0_13]),c_0_6])]),<A HREF="#c_0_21">c_0_21</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( app(skaf46(sk2,nil),nil) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_77]),c_0_6]),c_0_9])])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( app(cons(X1,X2),nil) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_30]),c_0_31]),<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,plain,
    ( nil = X1
    | memberP(X1,skaf83(X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_79">c_0_79</A>,c_0_34]),c_0_35]),c_0_36])])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( skaf46(sk2,nil) = sk2 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_80]),c_0_28])])).

<A NAME="c_0_85"></A>cnf(c_0_85,plain,
    ( cons(X1,app(X2,nil)) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_37]),c_0_6])])).

<A NAME="c_0_86"></A>cnf(c_0_86,plain,
    ( nil = X1
    | ~ ssList(app(app(X2,cons(skaf83(X1),X3)),X1))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(app(X2,cons(skaf83(X1),X3)),X1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,c_0_34]),c_0_35]),c_0_36])])).

<A NAME="c_0_87"></A>cnf(c_0_87,plain,
    ( app(skaf42(X1,skaf83(X1)),cons(skaf83(X1),skaf43(skaf83(X1),X1))) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_83]),c_0_35])])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( app(sk2,nil) = sk2 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,plain,
    ( cons(X1,app(X2,app(X3,nil))) != app(X2,X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_39]),c_0_6])]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_90"></A>cnf(c_0_90,plain,
    ( nil = X1
    | ~ ssList(app(X1,X1))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(X1,X1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>,c_0_87]),c_0_40]),c_0_41])])).

<A NAME="c_0_91"></A>cnf(c_0_91,plain,
    ( cons(hd(X1),tl(app(X1,X2))) = app(X1,X2)
    | nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_42]),c_0_15]),<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_92"></A>cnf(c_0_92,plain,
    ( nil = X1
    | ssList(app(tl(X1),X2))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_26]),c_0_15]),<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_93"></A>cnf(c_0_93,negated_conjecture,
    ( app(tl(sk2),nil) = tl(sk2)
    | sk2 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_88]),c_0_9]),c_0_6])])).

<A NAME="c_0_94"></A>cnf(c_0_94,plain,
    ( cons(X1,app(X2,nil)) != app(nil,X2)
    | ~ ssItem(X1)
    | ~ ssList(app(X2,nil))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>,c_0_10]),c_0_6])])).

<A NAME="c_0_95"></A>cnf(c_0_95,plain,
    ( nil = X1
    | ssItem(X2)
    | ~ ssList(app(X1,X1))
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,<A HREF="#c_0_43">c_0_43</A>])).

<A NAME="c_0_96"></A>cnf(c_0_96,plain,
    ( skaf82(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_34]),c_0_35]),c_0_36])])]),c_0_31]),<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,negated_conjecture,
    ( cons(hd(sk2),tl(sk2)) = sk2
    | sk2 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_88]),c_0_9]),c_0_6])])).

<A NAME="c_0_98"></A>cnf(c_0_98,negated_conjecture,
    ( sk2 = nil
    | ssList(tl(sk2)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_93]),c_0_9]),c_0_6])])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( cons(X1,X2) != app(nil,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,plain,
    ( app(nil,skaf45(X1,nil)) = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,c_0_45]),c_0_6])])).

<A NAME="c_0_101"></A>cnf(c_0_101,plain,
    ( nil = X1
    | ssItem(X2)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_102"></A>cnf(c_0_102,plain,
    ( ssList(cons(X1,app(X2,X3)))
    | ~ ssItem(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_37]),<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_103"></A>cnf(c_0_103,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_46])]),c_0_6])])).

<A NAME="c_0_104"></A>cnf(c_0_104,negated_conjecture,
    ( skaf82(sk2) = tl(sk2)
    | sk2 = nil
    | ~ ssItem(hd(sk2)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,c_0_97]),<A HREF="#c_0_98">c_0_98</A>])).

<A NAME="c_0_105"></A>cnf(c_0_105,plain,
    ( cons(X1,skaf45(X2,nil)) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_100]),c_0_11])])).

<A NAME="c_0_106"></A>cnf(c_0_106,plain,
    ( skaf45(X1,X2) = nil
    | ssItem(X3) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_101">c_0_101</A>,<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_107"></A>cnf(c_0_107,plain,
    ( ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>,c_0_103]),c_0_6])])).

<A NAME="c_0_108"></A>cnf(c_0_108,plain,
    ( frontsegP(cons(X1,X2),cons(X1,X3))
    | ~ frontsegP(X2,X3)
    | ~ ssItem(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_109"></A>cnf(c_0_109,negated_conjecture,
    ( skaf82(sk2) = tl(sk2)
    | sk2 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>,c_0_22]),c_0_9])])).

<A NAME="c_0_110"></A>cnf(c_0_110,plain,
    ( ssItem(X1)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_105">c_0_105</A>,c_0_106])]),<A HREF="#c_0_107">c_0_107</A>])).

<A NAME="c_0_111"></A>cnf(c_0_111,negated_conjecture,
    ( ~ neq(sk1,X1)
    | ~ frontsegP(sk2,X1)
    | ~ segmentP(X1,sk1)
    | ~ ssList(X1)
    | ~ equalelemsP(X1) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_2]),c_0_1]),<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_112"></A>cnf(c_0_112,plain,
    ( frontsegP(cons(X1,X2),cons(X1,nil))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>,c_0_45]),c_0_6])])).

<A NAME="c_0_113"></A>cnf(c_0_113,negated_conjecture,
    ( cons(skaf83(sk2),tl(sk2)) = sk2
    | sk2 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_109]),c_0_9])])).

<A NAME="c_0_114"></A>cnf(c_0_114,plain,
    ( ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_110">c_0_110</A>,<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,negated_conjecture,
    ( ~ neq(nil,X1)
    | ~ frontsegP(sk2,X1)
    | ~ ssList(X1)
    | ~ equalelemsP(X1) ),
    inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_111">c_0_111</A>,c_0_61]),c_0_61]),<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( sk2 = nil
    | frontsegP(sk2,cons(skaf83(sk2),nil)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_113]),c_0_35])]),<A HREF="#c_0_98">c_0_98</A>])).

<A NAME="c_0_117"></A>cnf(c_0_117,plain,
    ( ssList(cons(X1,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>,c_0_114])])).

<A NAME="c_0_118"></A>cnf(c_0_118,plain,
    ( equalelemsP(cons(X1,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,c_0_114])])).

<A NAME="c_0_119"></A>cnf(c_0_119,plain,
    ( cons(X1,nil) != nil
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_52]),c_0_6])])).

<A NAME="c_0_120"></A>cnf(c_0_120,negated_conjecture,
    ( equalelemsP(sk1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_121"></A>cnf(c_0_121,negated_conjecture,
    ( sk2 = nil
    | ~ neq(nil,cons(skaf83(sk2),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_115">c_0_115</A>,c_0_116]),c_0_117]),c_0_118])])).

<A NAME="c_0_122"></A>cnf(c_0_122,plain,
    ( X1 = X2
    | neq(X2,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_114]),c_0_114])])).

<A NAME="c_0_123"></A>cnf(c_0_123,plain,
    ( cons(X1,nil) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>,c_0_114])])).

<A NAME="c_0_124"></A>cnf(c_0_124,negated_conjecture,
    ( ~ neq(sk1,sk1)
    | ~ segmentP(sk1,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_111">c_0_111</A>,c_0_57]),c_0_5]),c_0_120])])).

<A NAME="c_0_125"></A>cnf(c_0_125,plain,
    ( segmentP(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_55]),c_0_6])])).

<A NAME="c_0_126"></A>cnf(c_0_126,negated_conjecture,
    ( neq(sk2,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_61])])).

<A NAME="c_0_127"></A>cnf(c_0_127,negated_conjecture,
    ( sk2 = nil ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_121">c_0_121</A>,c_0_122]),<A HREF="#c_0_123">c_0_123</A>])).

<A NAME="c_0_128"></A>cnf(c_0_128,negated_conjecture,
    ( ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_124">c_0_124</A>,c_0_61]),c_0_61]),c_0_61]),c_0_61]),c_0_125])])).

<A NAME="c_0_129"></A>cnf(c_0_129,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_126">c_0_126</A>,c_0_127]),<A HREF="#c_0_128">c_0_128</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC099-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.22  % Computer   : n137.star.cs.uiowa.edu
% 0.03/0.22  % Model      : x86_64 x86_64
% 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.22  % Memory     : 32218.75MB
% 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.22  % CPULimit   : 300
% 0.03/0.22  % DateTime   : Tue May 10 20:35:38 CDT 2016
% 0.03/0.22  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 1.62/1.87  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 1.62/1.87  # and selection function SelectCQArNTEqFirst.
% 1.62/1.87  #
% 1.62/1.87  # Preprocessing time       : 0.009 s
% 1.62/1.87  # Presaturation interreduction done
% 1.62/1.87  
% 1.62/1.87  # Proof found!
% 1.62/1.87  # SZS status Unsatisfiable
% 1.62/1.87  # SZS output start CNFRefutation
% 1.62/1.87  cnf(c_0_0, negated_conjecture, (frontsegP(sk4,sk3)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 1.62/1.87  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 1.62/1.87  cnf(c_0_2, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 1.62/1.87  cnf(c_0_3, negated_conjecture, (nil=sk2|~neq(sk1,nil)|~frontsegP(sk2,sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_11)).
% 1.62/1.87  cnf(c_0_4, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause100)).
% 1.62/1.87  cnf(c_0_5, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 1.62/1.87  cnf(c_0_6, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 1.62/1.87  cnf(c_0_7, axiom, (nil=X1|~frontsegP(nil,X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause84)).
% 1.62/1.87  cnf(c_0_8, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause132)).
% 1.62/1.87  cnf(c_0_9, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 1.62/1.87  cnf(c_0_10, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause74)).
% 1.62/1.87  cnf(c_0_11, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause51)).
% 1.62/1.87  cnf(c_0_12, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause173)).
% 1.62/1.87  cnf(c_0_13, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause73)).
% 1.62/1.87  cnf(c_0_14, axiom, (rearsegP(X3,X2)|app(X1,X2)!=X3|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause143)).
% 1.62/1.87  cnf(c_0_15, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 1.62/1.87  cnf(c_0_16, axiom, (app(app(skaf47(X1,X2),X2),skaf48(X2,X1))=X1|~segmentP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause168)).
% 1.62/1.87  cnf(c_0_17, axiom, (ssList(skaf48(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause48)).
% 1.62/1.87  cnf(c_0_18, axiom, (ssList(skaf47(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause49)).
% 1.62/1.87  cnf(c_0_19, axiom, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause99)).
% 1.62/1.87  cnf(c_0_20, axiom, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause104)).
% 1.62/1.87  cnf(c_0_21, axiom, (ssList(tl(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause75)).
% 1.62/1.87  cnf(c_0_22, axiom, (ssItem(hd(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause76)).
% 1.62/1.87  cnf(c_0_23, axiom, (rearsegP(app(X3,X1),X2)|~rearsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause136)).
% 1.62/1.87  cnf(c_0_24, axiom, (rearsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause58)).
% 1.62/1.87  cnf(c_0_25, axiom, (app(skaf46(X1,X2),X2)=X1|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause131)).
% 1.62/1.87  cnf(c_0_26, axiom, (nil=X2|tl(app(X2,X1))=app(tl(X2),X1)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause133)).
% 1.62/1.87  cnf(c_0_27, axiom, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause118)).
% 1.62/1.87  cnf(c_0_28, axiom, (ssList(skaf46(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause50)).
% 1.62/1.87  cnf(c_0_29, axiom, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause138)).
% 1.62/1.87  cnf(c_0_30, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause96)).
% 1.62/1.87  cnf(c_0_31, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
% 1.62/1.87  cnf(c_0_32, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause98)).
% 1.62/1.87  cnf(c_0_33, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
% 1.62/1.87  cnf(c_0_34, axiom, (cons(skaf83(X1),skaf82(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause109)).
% 1.62/1.87  cnf(c_0_35, axiom, (ssItem(skaf83(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause12)).
% 1.62/1.87  cnf(c_0_36, axiom, (ssList(skaf82(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause13)).
% 1.62/1.87  cnf(c_0_37, axiom, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause157)).
% 1.62/1.87  cnf(c_0_38, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause169)).
% 1.62/1.87  cnf(c_0_39, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
% 1.62/1.87  cnf(c_0_40, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause52)).
% 1.62/1.87  cnf(c_0_41, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause53)).
% 1.62/1.87  cnf(c_0_42, axiom, (nil=X2|hd(app(X2,X1))=hd(X2)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause123)).
% 1.62/1.87  cnf(c_0_43, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
% 1.62/1.87  cnf(c_0_44, axiom, (X4=X2|cons(X1,X2)!=cons(X3,X4)|~ssItem(X3)|~ssItem(X1)|~ssList(X4)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause171)).
% 1.62/1.87  cnf(c_0_45, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause60)).
% 1.62/1.87  cnf(c_0_46, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause145)).
% 1.62/1.87  cnf(c_0_47, axiom, (frontsegP(cons(X3,X1),cons(X4,X2))|~frontsegP(X1,X2)|X3!=X4|~ssList(X2)|~ssList(X1)|~ssItem(X4)|~ssItem(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause178)).
% 1.62/1.87  cnf(c_0_48, negated_conjecture, (~ssList(X1)|~neq(sk3,X1)|~frontsegP(sk4,X1)|~segmentP(X1,sk3)|~equalelemsP(X1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
% 1.62/1.87  cnf(c_0_49, axiom, (ssItem(skaf44(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause47)).
% 1.62/1.87  cnf(c_0_50, axiom, (segmentP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause56)).
% 1.62/1.87  cnf(c_0_51, axiom, (equalelemsP(cons(X1,nil))|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause64)).
% 1.62/1.87  cnf(c_0_52, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
% 1.62/1.87  cnf(c_0_53, negated_conjecture, (equalelemsP(sk3)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
% 1.62/1.87  cnf(c_0_54, axiom, (neq(X2,X1)|X2=X1|~ssItem(X1)|~ssItem(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause102)).
% 1.62/1.87  cnf(c_0_55, axiom, (segmentP(nil,X1)|nil!=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause79)).
% 1.62/1.87  cnf(c_0_56, negated_conjecture, (neq(sk2,nil)|nil!=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_12)).
% 1.62/1.87  cnf(c_0_57, negated_conjecture, (frontsegP(sk2,sk1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
% 1.62/1.87  cnf(c_0_58, negated_conjecture, (sk2=nil|~neq(sk1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_57])])).
% 1.62/1.87  cnf(c_0_59, negated_conjecture, (sk1=nil|sk2=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_4]), c_0_5]), c_0_6])])).
% 1.62/1.87  cnf(c_0_60, negated_conjecture, (sk1=nil|frontsegP(nil,sk1)), inference(spm,[status(thm)],[c_0_57, c_0_59])).
% 1.62/1.87  cnf(c_0_61, negated_conjecture, (sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_60]), c_0_5])])).
% 1.62/1.87  cnf(c_0_62, negated_conjecture, (frontsegP(sk2,nil)), inference(rw,[status(thm)],[c_0_57, c_0_61])).
% 1.62/1.87  cnf(c_0_63, negated_conjecture, (app(nil,skaf45(sk2,nil))=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_62]), c_0_6]), c_0_9])])).
% 1.62/1.87  cnf(c_0_64, negated_conjecture, (skaf45(sk2,nil)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_63]), c_0_11])])).
% 1.62/1.87  cnf(c_0_65, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(er,[status(thm)],[c_0_12])).
% 1.62/1.87  cnf(c_0_66, negated_conjecture, (app(nil,sk2)=sk2), inference(rw,[status(thm)],[c_0_63, c_0_64])).
% 1.62/1.87  cnf(c_0_67, negated_conjecture, (segmentP(app(sk2,X1),sk2)|~ssList(app(sk2,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_9]), c_0_6])])).
% 1.62/1.87  cnf(c_0_68, negated_conjecture, (segmentP(sk2,sk2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_13]), c_0_9]), c_0_6])])).
% 1.62/1.87  cnf(c_0_69, plain, (rearsegP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_14]), c_0_15])).
% 1.62/1.87  cnf(c_0_70, negated_conjecture, (app(app(skaf47(sk2,sk2),sk2),skaf48(sk2,sk2))=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_68]), c_0_9])])).
% 1.62/1.87  cnf(c_0_71, negated_conjecture, (rearsegP(sk2,skaf48(sk2,sk2))|~ssList(app(skaf47(sk2,sk2),sk2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_17])])).
% 1.62/1.87  cnf(c_0_72, negated_conjecture, (rearsegP(sk2,skaf48(sk2,sk2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_15]), c_0_18]), c_0_9])])).
% 1.62/1.87  cnf(c_0_73, plain, (nil=X1|tl(X1)!=X1|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22])).
% 1.62/1.87  cnf(c_0_74, plain, (rearsegP(app(X1,X2),nil)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_6])])).
% 1.62/1.87  cnf(c_0_75, negated_conjecture, (app(skaf46(sk2,skaf48(sk2,sk2)),skaf48(sk2,sk2))=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_72]), c_0_17]), c_0_9])])).
% 1.62/1.87  cnf(c_0_76, plain, (nil=X1|app(tl(X1),X2)!=app(X1,X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_26]), c_0_15]), c_0_27])).
% 1.62/1.87  cnf(c_0_77, negated_conjecture, (rearsegP(sk2,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_28]), c_0_17])])).
% 1.62/1.87  cnf(c_0_78, plain, (nil=X1|tl(X1)!=app(X1,nil)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_13]), c_0_6])]), c_0_21])).
% 1.62/1.87  cnf(c_0_79, plain, (memberP(cons(X1,X2),X1)|~ssItem(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_29])).
% 1.62/1.87  cnf(c_0_80, negated_conjecture, (app(skaf46(sk2,nil),nil)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_77]), c_0_6]), c_0_9])])).
% 1.62/1.87  cnf(c_0_81, plain, (app(cons(X1,X2),nil)!=X2|~ssItem(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_30]), c_0_31]), c_0_32])).
% 1.62/1.87  cnf(c_0_82, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_33])).
% 1.62/1.87  cnf(c_0_83, plain, (nil=X1|memberP(X1,skaf83(X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_34]), c_0_35]), c_0_36])])).
% 1.62/1.87  cnf(c_0_84, negated_conjecture, (skaf46(sk2,nil)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_80]), c_0_28])])).
% 1.62/1.87  cnf(c_0_85, plain, (cons(X1,app(X2,nil))!=X2|~ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_37]), c_0_6])])).
% 1.62/1.87  cnf(c_0_86, plain, (nil=X1|~ssList(app(app(X2,cons(skaf83(X1),X3)),X1))|~ssList(X3)|~ssList(X2)|~ssList(X1)|~duplicatefreeP(app(app(X2,cons(skaf83(X1),X3)),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_34]), c_0_35]), c_0_36])])).
% 1.62/1.87  cnf(c_0_87, plain, (app(skaf42(X1,skaf83(X1)),cons(skaf83(X1),skaf43(skaf83(X1),X1)))=X1|nil=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_83]), c_0_35])])).
% 1.62/1.87  cnf(c_0_88, negated_conjecture, (app(sk2,nil)=sk2), inference(rw,[status(thm)],[c_0_80, c_0_84])).
% 1.62/1.87  cnf(c_0_89, plain, (cons(X1,app(X2,app(X3,nil)))!=app(X2,X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_39]), c_0_6])]), c_0_15])).
% 1.62/1.87  cnf(c_0_90, plain, (nil=X1|~ssList(app(X1,X1))|~ssList(X1)|~duplicatefreeP(app(X1,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_40]), c_0_41])])).
% 1.62/1.87  cnf(c_0_91, plain, (cons(hd(X1),tl(app(X1,X2)))=app(X1,X2)|nil=X1|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_42]), c_0_15]), c_0_27])).
% 1.62/1.87  cnf(c_0_92, plain, (nil=X1|ssList(app(tl(X1),X2))|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_26]), c_0_15]), c_0_27])).
% 1.62/1.87  cnf(c_0_93, negated_conjecture, (app(tl(sk2),nil)=tl(sk2)|sk2=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_88]), c_0_9]), c_0_6])])).
% 1.62/1.87  cnf(c_0_94, plain, (cons(X1,app(X2,nil))!=app(nil,X2)|~ssItem(X1)|~ssList(app(X2,nil))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_10]), c_0_6])])).
% 1.62/1.87  cnf(c_0_95, plain, (nil=X1|ssItem(X2)|~ssList(app(X1,X1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_90, c_0_43])).
% 1.62/1.87  cnf(c_0_96, plain, (skaf82(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_34]), c_0_35]), c_0_36])])]), c_0_31]), c_0_32])).
% 1.62/1.87  cnf(c_0_97, negated_conjecture, (cons(hd(sk2),tl(sk2))=sk2|sk2=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_88]), c_0_9]), c_0_6])])).
% 1.62/1.87  cnf(c_0_98, negated_conjecture, (sk2=nil|ssList(tl(sk2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_9]), c_0_6])])).
% 1.62/1.87  cnf(c_0_99, plain, (cons(X1,X2)!=app(nil,X2)|~ssItem(X1)|~ssList(X2)), inference(spm,[status(thm)],[c_0_94, c_0_13])).
% 1.62/1.87  cnf(c_0_100, plain, (app(nil,skaf45(X1,nil))=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_45]), c_0_6])])).
% 1.62/1.87  cnf(c_0_101, plain, (nil=X1|ssItem(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_95, c_0_15])).
% 1.62/1.87  cnf(c_0_102, plain, (ssList(cons(X1,app(X2,X3)))|~ssItem(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_37]), c_0_31])).
% 1.62/1.87  cnf(c_0_103, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_46])]), c_0_6])])).
% 1.62/1.87  cnf(c_0_104, negated_conjecture, (skaf82(sk2)=tl(sk2)|sk2=nil|~ssItem(hd(sk2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98])).
% 1.62/1.87  cnf(c_0_105, plain, (cons(X1,skaf45(X2,nil))!=X2|~ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_11])])).
% 1.62/1.87  cnf(c_0_106, plain, (skaf45(X1,X2)=nil|ssItem(X3)), inference(spm,[status(thm)],[c_0_101, c_0_11])).
% 1.62/1.87  cnf(c_0_107, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_6])])).
% 1.62/1.87  cnf(c_0_108, plain, (frontsegP(cons(X1,X2),cons(X1,X3))|~frontsegP(X2,X3)|~ssItem(X1)|~ssList(X3)|~ssList(X2)), inference(er,[status(thm)],[c_0_47])).
% 1.62/1.87  cnf(c_0_109, negated_conjecture, (skaf82(sk2)=tl(sk2)|sk2=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_22]), c_0_9])])).
% 1.62/1.87  cnf(c_0_110, plain, (ssItem(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106])]), c_0_107])).
% 1.62/1.87  cnf(c_0_111, negated_conjecture, (~neq(sk1,X1)|~frontsegP(sk2,X1)|~segmentP(X1,sk1)|~ssList(X1)|~equalelemsP(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_2]), c_0_1]), c_0_2])).
% 1.62/1.87  cnf(c_0_112, plain, (frontsegP(cons(X1,X2),cons(X1,nil))|~ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_45]), c_0_6])])).
% 1.62/1.87  cnf(c_0_113, negated_conjecture, (cons(skaf83(sk2),tl(sk2))=sk2|sk2=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_109]), c_0_9])])).
% 1.62/1.87  cnf(c_0_114, plain, (ssItem(X1)), inference(spm,[status(thm)],[c_0_110, c_0_49])).
% 1.62/1.87  cnf(c_0_115, negated_conjecture, (~neq(nil,X1)|~frontsegP(sk2,X1)|~ssList(X1)|~equalelemsP(X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_61]), c_0_61]), c_0_50])).
% 1.62/1.87  cnf(c_0_116, negated_conjecture, (sk2=nil|frontsegP(sk2,cons(skaf83(sk2),nil))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_35])]), c_0_98])).
% 1.62/1.87  cnf(c_0_117, plain, (ssList(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_114])])).
% 1.62/1.87  cnf(c_0_118, plain, (equalelemsP(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_114])])).
% 1.62/1.87  cnf(c_0_119, plain, (cons(X1,nil)!=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_52]), c_0_6])])).
% 1.62/1.87  cnf(c_0_120, negated_conjecture, (equalelemsP(sk1)), inference(rw,[status(thm)],[c_0_53, c_0_2])).
% 1.62/1.87  cnf(c_0_121, negated_conjecture, (sk2=nil|~neq(nil,cons(skaf83(sk2),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_117]), c_0_118])])).
% 1.62/1.87  cnf(c_0_122, plain, (X1=X2|neq(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_114]), c_0_114])])).
% 1.62/1.87  cnf(c_0_123, plain, (cons(X1,nil)!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_114])])).
% 1.62/1.87  cnf(c_0_124, negated_conjecture, (~neq(sk1,sk1)|~segmentP(sk1,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_57]), c_0_5]), c_0_120])])).
% 1.62/1.87  cnf(c_0_125, plain, (segmentP(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_55]), c_0_6])])).
% 1.62/1.87  cnf(c_0_126, negated_conjecture, (neq(sk2,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_61])])).
% 1.62/1.87  cnf(c_0_127, negated_conjecture, (sk2=nil), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_123])).
% 1.62/1.87  cnf(c_0_128, negated_conjecture, (~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_61]), c_0_61]), c_0_61]), c_0_61]), c_0_125])])).
% 1.62/1.87  cnf(c_0_129, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_127]), c_0_128]), ['proof']).
% 1.62/1.87  # SZS output end CNFRefutation
% 1.62/1.87  # Training examples: 0 positive, 0 negative
% 1.62/1.88  
% 1.62/1.88  # -------------------------------------------------
% 1.62/1.88  # User time                : 1.616 s
% 1.62/1.88  # System time              : 0.038 s
% 1.62/1.88  # Total time               : 1.655 s
% 1.62/1.88  # Maximum resident set size: 1864 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
