<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC367+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC367+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC367+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC367%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC367+1.p">SWC367+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n174.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:40 EDT 2016

% Result     : Theorem 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of formulae       :   19 (  94 expanded)
%              Number of clauses        :   13 (  34 expanded)
%              Number of leaves         :    3 (  21 expanded)
%              Depth                    :    9
%              Number of atoms          :   64 ( 753 expanded)
%              Number of equality atoms :   24 ( 288 expanded)
%              Maximal formula depth    :   16 (   4 average)
%              Maximal clause size      :   15 (   1 average)
%              Maximal term depth       :    1 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | rearsegP(X2,X1)
                    | ( ( nil != X4
                        | nil != X3 )
                      & ( ~ neq(X3,nil)
                        | ~ rearsegP(X4,X3) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( rearsegP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax52)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_3"></A>fof(c_0_3,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | rearsegP(X2,X1)
                      | ( ( nil != X4
                          | nil != X3 )
                        & ( ~ neq(X3,nil)
                          | ~ rearsegP(X4,X3) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_4"></A>fof(c_0_4,negated_conjecture,
    ( ssList(esk48_0)
    & ssList(esk49_0)
    & ssList(esk50_0)
    & ssList(esk51_0)
    & esk49_0 = esk51_0
    & esk48_0 = esk50_0
    & ~ rearsegP(esk49_0,esk48_0)
    & ( neq(esk50_0,nil)
      | nil = esk51_0 )
    & ( rearsegP(esk51_0,esk50_0)
      | nil = esk51_0 )
    & ( neq(esk50_0,nil)
      | nil = esk50_0 )
    & ( rearsegP(esk51_0,esk50_0)
      | nil = esk50_0 ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( ~ rearsegP(esk49_0,esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( nil = esk50_0
    | rearsegP(esk51_0,esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ~ rearsegP(esk51_0,esk50_0) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,c_0_6]),<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( esk50_0 = nil ),
    inference(sr,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( nil = esk51_0
    | rearsegP(esk51_0,esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( ~ rearsegP(esk51_0,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_13"></A>fof(c_0_13,plain,(
    ! [X2] :
      ( ( ~ rearsegP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | rearsegP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( esk51_0 = nil ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_10]),<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,plain,
    ( rearsegP(nil,X1)
    | ~ ssList(X1)
    | nil != X1 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( ~ rearsegP(nil,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,plain,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_15]),c_0_16])]),<A HREF="#c_0_17">c_0_17</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC367+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n174.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:42:38 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.03/0.25  # and selection function SelectCQPrecWNTNp.
% 0.03/0.25  #
% 0.03/0.25  # Preprocessing time       : 0.014 s
% 0.03/0.25  # Presaturation interreduction done
% 0.03/0.25  
% 0.03/0.25  # Proof found!
% 0.03/0.25  # SZS status Theorem
% 0.03/0.25  # SZS output start CNFRefutation
% 0.03/0.25  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|rearsegP(X2,X1))|((nil!=X4|nil!=X3)&(~(neq(X3,nil))|~(rearsegP(X4,X3)))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.03/0.25  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>(rearsegP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax52)).
% 0.03/0.25  fof(c_0_2, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.03/0.25  fof(c_0_3, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|rearsegP(X2,X1))|((nil!=X4|nil!=X3)&(~neq(X3,nil)|~rearsegP(X4,X3)))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.03/0.25  fof(c_0_4, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((esk49_0=esk51_0&esk48_0=esk50_0)&~rearsegP(esk49_0,esk48_0))&(((neq(esk50_0,nil)|nil=esk51_0)&(rearsegP(esk51_0,esk50_0)|nil=esk51_0))&((neq(esk50_0,nil)|nil=esk50_0)&(rearsegP(esk51_0,esk50_0)|nil=esk50_0))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
% 0.03/0.25  cnf(c_0_5, negated_conjecture, (~rearsegP(esk49_0,esk48_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.03/0.25  cnf(c_0_6, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.03/0.25  cnf(c_0_7, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.03/0.25  cnf(c_0_8, negated_conjecture, (nil=esk50_0|rearsegP(esk51_0,esk50_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.03/0.25  cnf(c_0_9, negated_conjecture, (~rearsegP(esk51_0,esk50_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_6]), c_0_7])).
% 0.03/0.25  cnf(c_0_10, negated_conjecture, (esk50_0=nil), inference(sr,[status(thm)],[c_0_8, c_0_9])).
% 0.03/0.25  cnf(c_0_11, negated_conjecture, (nil=esk51_0|rearsegP(esk51_0,esk50_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.03/0.25  cnf(c_0_12, negated_conjecture, (~rearsegP(esk51_0,nil)), inference(rw,[status(thm)],[c_0_9, c_0_10])).
% 0.03/0.25  fof(c_0_13, plain, (![X2]:(((~rearsegP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|rearsegP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.03/0.25  cnf(c_0_14, negated_conjecture, (esk51_0=nil), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_10]), c_0_12])).
% 0.03/0.25  cnf(c_0_15, plain, (rearsegP(nil,X1)|~ssList(X1)|nil!=X1), inference(split_conjunct,[status(thm)],[c_0_13])).
% 0.03/0.25  cnf(c_0_16, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.03/0.25  cnf(c_0_17, negated_conjecture, (~rearsegP(nil,nil)), inference(rw,[status(thm)],[c_0_12, c_0_14])).
% 0.03/0.25  cnf(c_0_18, plain, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_16])]), c_0_17]), ['proof']).
% 0.03/0.25  # SZS output end CNFRefutation
% 0.03/0.25  # Training examples: 0 positive, 0 negative
% 0.03/0.25  
% 0.03/0.25  # -------------------------------------------------
% 0.03/0.25  # User time                : 0.017 s
% 0.03/0.25  # System time              : 0.001 s
% 0.03/0.25  # Total time               : 0.018 s
% 0.03/0.25  # Maximum resident set size: 1868 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
