<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC184-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC184-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC184-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC184-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC184-1.p">SWC184-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n131.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:58 EDT 2016

% Result     : Unsatisfiable 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of clauses        :   68 ( 926 expanded)
%              Number of leaves         :   29 ( 609 expanded)
%              Depth                    :   23
%              Number of atoms          :  154 (1795 expanded)
%              Number of equality atoms :   53 ( 376 expanded)
%              Maximal clause size      :    7 (   2 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( app(app(sk6,cons(sk5,nil)),sk7) = sk1 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( ssList(sk6) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssList(sk7) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ssItem(sk5) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( frontsegP(X3,X1)
    | app(X1,X2) != X3
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause144)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( ssList(skaf42(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause53)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( nil = X2
    | tl(app(X2,X1)) = app(tl(X2),X1)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause133)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( nil = X2
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause119)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( cons(sk8,nil) = sk3
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_17)).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( ssItem(sk8)
    | nil = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( nil = X1
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause118)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_23"></A>cnf(c_0_23,axiom,
    ( memberP(cons(X2,X3),X1)
    | X1 != X2
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause138)).

<A NAME="c_0_24"></A>cnf(c_0_24,axiom,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause75)).

<A NAME="c_0_25"></A>cnf(c_0_25,axiom,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause71)).

<A NAME="c_0_26"></A>cnf(c_0_26,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_27"></A>cnf(c_0_27,axiom,
    ( cons(X1,X2) != nil
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause98)).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( memberP(sk6,sk5)
    | memberP(sk7,sk5) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( app(app(sk6,cons(sk5,nil)),sk7) = sk3 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( app(sk6,app(cons(sk5,nil),sk7)) = sk3
    | ~ ssList(cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_2]),c_0_3]),c_0_4])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( app(sk6,app(cons(sk5,nil),sk7)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_5]),c_0_6]),c_0_7])])).

<A NAME="c_0_32"></A>cnf(c_0_32,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( app(sk6,cons(sk5,sk7)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_9]),c_0_6]),c_0_4])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( ~ ssList(app(sk3,cons(sk5,X1)))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(sk3,cons(sk5,X1))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_33]),c_0_6]),c_0_4]),c_0_3])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(app(sk3,cons(sk5,X2)))
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(cons(sk5,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_12]),c_0_36])])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_5]),c_0_6])])).

<A NAME="c_0_39"></A>cnf(c_0_39,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]),<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( frontsegP(sk3,sk6)
    | ~ ssList(cons(sk5,sk7)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_33]),c_0_3])])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( app(tl(sk6),cons(sk5,sk7)) = tl(sk3)
    | sk6 = nil
    | ~ ssList(cons(sk5,sk7)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_33]),c_0_3])])).

<A NAME="c_0_43"></A>cnf(c_0_43,plain,
    ( ssList(cons(X1,X2))
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,c_0_40])])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( frontsegP(sk3,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_5]),c_0_6]),c_0_4])])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( app(tl(sk6),cons(sk5,sk7)) = tl(sk3)
    | sk6 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_43]),c_0_4])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( app(sk6,skaf45(sk3,sk6)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_44]),c_0_3]),c_0_36])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( cons(sk5,sk7) = nil
    | sk6 = nil
    | tl(sk3) != nil
    | ~ ssList(cons(sk5,sk7))
    | ~ ssList(tl(sk6)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,<A HREF="#c_0_45">c_0_45</A>])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( tl(sk3) = nil
    | sk3 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_19]),c_0_7])]),<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( sk6 = nil
    | sk3 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_46]),c_0_22]),c_0_3])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( cons(sk5,sk7) = nil
    | sk6 = nil
    | ~ ssList(cons(sk5,sk7))
    | ~ ssList(tl(sk6)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_48]),<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( cons(sk5,sk7) = nil
    | sk6 = nil
    | ~ ssList(cons(sk5,sk7)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_24]),c_0_3])])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,c_0_40])])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( cons(sk5,sk7) = nil
    | sk6 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_43]),c_0_4])])).

<A NAME="c_0_55"></A>cnf(c_0_55,plain,
    ( ~ memberP(nil,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_40])])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( sk6 = nil ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_54]),c_0_4])]),<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( app(nil,cons(sk5,sk7)) = sk3 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( cons(sk5,sk7) = sk3
    | ~ ssList(cons(sk5,sk7)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,plain,
    ( cons(X1,X2) != nil
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_40])])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( cons(sk5,sk7) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_43]),c_0_4])])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( sk3 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,c_0_60]),c_0_4])])).

<A NAME="c_0_62"></A>cnf(c_0_62,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_40])])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( cons(sk8,nil) = sk3 ),
    inference(sr,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,<A HREF="#c_0_61">c_0_61</A>])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( tl(sk3) = sk7 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_60]),c_0_4])])).

<A NAME="c_0_65"></A>cnf(c_0_65,negated_conjecture,
    ( memberP(sk7,sk5) ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_56]),<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,negated_conjecture,
    ( sk7 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_63]),c_0_64]),c_0_7])])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_66]),<A HREF="#c_0_55">c_0_55</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC184-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n131.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:36:38 CDT 2016
% 0.02/0.22  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.28  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.02/0.28  # and selection function SelectCQArNTEqFirst.
% 0.02/0.28  #
% 0.02/0.28  # Preprocessing time       : 0.009 s
% 0.02/0.28  # Presaturation interreduction done
% 0.02/0.28  
% 0.02/0.28  # Proof found!
% 0.02/0.28  # SZS status Unsatisfiable
% 0.02/0.28  # SZS output start CNFRefutation
% 0.02/0.28  cnf(c_0_0, negated_conjecture, (app(app(sk6,cons(sk5,nil)),sk7)=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_10)).
% 0.02/0.28  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.02/0.28  cnf(c_0_2, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
% 0.02/0.28  cnf(c_0_3, negated_conjecture, (ssList(sk6)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
% 0.02/0.28  cnf(c_0_4, negated_conjecture, (ssList(sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
% 0.02/0.28  cnf(c_0_5, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.02/0.28  cnf(c_0_6, negated_conjecture, (ssItem(sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 0.02/0.28  cnf(c_0_7, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.02/0.28  cnf(c_0_8, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
% 0.02/0.28  cnf(c_0_9, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
% 0.02/0.28  cnf(c_0_10, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
% 0.02/0.28  cnf(c_0_11, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.02/0.28  cnf(c_0_12, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.02/0.28  cnf(c_0_13, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause144)).
% 0.02/0.28  cnf(c_0_14, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause53)).
% 0.02/0.28  cnf(c_0_15, axiom, (nil=X2|tl(app(X2,X1))=app(tl(X2),X1)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause133)).
% 0.02/0.28  cnf(c_0_16, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause132)).
% 0.02/0.28  cnf(c_0_17, axiom, (nil=X2|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause119)).
% 0.02/0.28  cnf(c_0_18, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause96)).
% 0.02/0.28  cnf(c_0_19, negated_conjecture, (cons(sk8,nil)=sk3|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_17)).
% 0.02/0.28  cnf(c_0_20, negated_conjecture, (ssItem(sk8)|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
% 0.02/0.28  cnf(c_0_21, axiom, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause118)).
% 0.02/0.28  cnf(c_0_22, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause51)).
% 0.02/0.28  cnf(c_0_23, axiom, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause138)).
% 0.02/0.28  cnf(c_0_24, axiom, (ssList(tl(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause75)).
% 0.02/0.28  cnf(c_0_25, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause71)).
% 0.02/0.28  cnf(c_0_26, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause74)).
% 0.02/0.28  cnf(c_0_27, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause98)).
% 0.02/0.28  cnf(c_0_28, negated_conjecture, (memberP(sk6,sk5)|memberP(sk7,sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_11)).
% 0.02/0.28  cnf(c_0_29, negated_conjecture, (app(app(sk6,cons(sk5,nil)),sk7)=sk3), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 0.02/0.28  cnf(c_0_30, negated_conjecture, (app(sk6,app(cons(sk5,nil),sk7))=sk3|~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_2]), c_0_3]), c_0_4])])).
% 0.02/0.28  cnf(c_0_31, negated_conjecture, (app(sk6,app(cons(sk5,nil),sk7))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_5]), c_0_6]), c_0_7])])).
% 0.02/0.28  cnf(c_0_32, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_8])).
% 0.02/0.28  cnf(c_0_33, negated_conjecture, (app(sk6,cons(sk5,sk7))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_9]), c_0_6]), c_0_4])])).
% 0.02/0.28  cnf(c_0_34, negated_conjecture, (~ssList(app(sk3,cons(sk5,X1)))|~ssList(X1)|~duplicatefreeP(app(sk3,cons(sk5,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_6]), c_0_4]), c_0_3])])).
% 0.02/0.28  cnf(c_0_35, negated_conjecture, (ssItem(X1)|~ssList(app(sk3,cons(sk5,X2)))|~ssList(X2)), inference(spm,[status(thm)],[c_0_34, c_0_10])).
% 0.02/0.28  cnf(c_0_36, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_11, c_0_1])).
% 0.02/0.28  cnf(c_0_37, negated_conjecture, (ssItem(X1)|~ssList(cons(sk5,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_12]), c_0_36])])).
% 0.02/0.28  cnf(c_0_38, negated_conjecture, (ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_5]), c_0_6])])).
% 0.02/0.28  cnf(c_0_39, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_12])).
% 0.02/0.28  cnf(c_0_40, negated_conjecture, (ssItem(X1)), inference(spm,[status(thm)],[c_0_38, c_0_14])).
% 0.02/0.28  cnf(c_0_41, negated_conjecture, (frontsegP(sk3,sk6)|~ssList(cons(sk5,sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_33]), c_0_3])])).
% 0.02/0.28  cnf(c_0_42, negated_conjecture, (app(tl(sk6),cons(sk5,sk7))=tl(sk3)|sk6=nil|~ssList(cons(sk5,sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_33]), c_0_3])])).
% 0.02/0.28  cnf(c_0_43, plain, (ssList(cons(X1,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_40])])).
% 0.02/0.28  cnf(c_0_44, negated_conjecture, (frontsegP(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_5]), c_0_6]), c_0_4])])).
% 0.02/0.28  cnf(c_0_45, negated_conjecture, (app(tl(sk6),cons(sk5,sk7))=tl(sk3)|sk6=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_4])])).
% 0.02/0.28  cnf(c_0_46, negated_conjecture, (app(sk6,skaf45(sk3,sk6))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_44]), c_0_3]), c_0_36])])).
% 0.02/0.28  cnf(c_0_47, negated_conjecture, (cons(sk5,sk7)=nil|sk6=nil|tl(sk3)!=nil|~ssList(cons(sk5,sk7))|~ssList(tl(sk6))), inference(spm,[status(thm)],[c_0_17, c_0_45])).
% 0.02/0.28  cnf(c_0_48, negated_conjecture, (tl(sk3)=nil|sk3=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_7])]), c_0_20])).
% 0.02/0.28  cnf(c_0_49, negated_conjecture, (sk6=nil|sk3!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_46]), c_0_22]), c_0_3])])).
% 0.02/0.28  cnf(c_0_50, negated_conjecture, (cons(sk5,sk7)=nil|sk6=nil|~ssList(cons(sk5,sk7))|~ssList(tl(sk6))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
% 0.02/0.28  cnf(c_0_51, plain, (memberP(cons(X1,X2),X1)|~ssItem(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_23])).
% 0.02/0.28  cnf(c_0_52, negated_conjecture, (cons(sk5,sk7)=nil|sk6=nil|~ssList(cons(sk5,sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_24]), c_0_3])])).
% 0.02/0.28  cnf(c_0_53, plain, (memberP(cons(X1,X2),X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_40])])).
% 0.02/0.28  cnf(c_0_54, negated_conjecture, (cons(sk5,sk7)=nil|sk6=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_43]), c_0_4])])).
% 0.02/0.28  cnf(c_0_55, plain, (~memberP(nil,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_40])])).
% 0.02/0.28  cnf(c_0_56, negated_conjecture, (sk6=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_4])]), c_0_55])).
% 0.02/0.28  cnf(c_0_57, negated_conjecture, (app(nil,cons(sk5,sk7))=sk3), inference(rw,[status(thm)],[c_0_33, c_0_56])).
% 0.02/0.28  cnf(c_0_58, negated_conjecture, (cons(sk5,sk7)=sk3|~ssList(cons(sk5,sk7))), inference(spm,[status(thm)],[c_0_26, c_0_57])).
% 0.02/0.28  cnf(c_0_59, plain, (cons(X1,X2)!=nil|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_40])])).
% 0.02/0.28  cnf(c_0_60, negated_conjecture, (cons(sk5,sk7)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_43]), c_0_4])])).
% 0.02/0.28  cnf(c_0_61, negated_conjecture, (sk3!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_4])])).
% 0.02/0.28  cnf(c_0_62, plain, (tl(cons(X1,X2))=X2|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_40])])).
% 0.02/0.28  cnf(c_0_63, negated_conjecture, (cons(sk8,nil)=sk3), inference(sr,[status(thm)],[c_0_19, c_0_61])).
% 0.02/0.28  cnf(c_0_64, negated_conjecture, (tl(sk3)=sk7), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_60]), c_0_4])])).
% 0.02/0.28  cnf(c_0_65, negated_conjecture, (memberP(sk7,sk5)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_56]), c_0_55])).
% 0.02/0.28  cnf(c_0_66, negated_conjecture, (sk7=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_7])])).
% 0.02/0.28  cnf(c_0_67, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_66]), c_0_55]), ['proof']).
% 0.02/0.28  # SZS output end CNFRefutation
% 0.02/0.28  # Training examples: 0 positive, 0 negative
% 0.02/0.28  
% 0.02/0.28  # -------------------------------------------------
% 0.02/0.28  # User time                : 0.053 s
% 0.02/0.28  # System time              : 0.005 s
% 0.02/0.28  # Total time               : 0.058 s
% 0.02/0.28  # Maximum resident set size: 1868 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
