<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC378+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC378+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC378+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC378%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC378+1.p">SWC378+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n148.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:43 EDT 2016

% Result     : Theorem 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of formulae       :   39 ( 398 expanded)
%              Number of clauses        :   32 ( 135 expanded)
%              Number of leaves         :    3 (  93 expanded)
%              Depth                    :   15
%              Number of atoms          :  135 (3989 expanded)
%              Number of equality atoms :   25 ( 980 expanded)
%              Maximal formula depth    :   19 (   4 average)
%              Maximal clause size      :   16 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ~ ssList(X4)
                  | X2 != X4
                  | X1 != X3
                  | ! [X5] :
                      ( ssList(X5)
                     => ! [X6] :
                          ( ~ ssList(X6)
                          | app(X5,X6) != X4
                          | app(X6,X5) != X3 ) )
                  | ! [X7] :
                      ( ~ ssItem(X7)
                      | ( ~ memberP(X2,X7)
                        & ~ memberP(X1,X7) )
                      | ( memberP(X2,X7)
                        & memberP(X1,X7) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(app(X2,X3),X1)
              &lt;=> ( memberP(X2,X1)
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax36)).

<A NAME="c_0_3"></A>fof(c_0_3,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ~ ssList(X4)
                    | X2 != X4
                    | X1 != X3
                    | ! [X5] :
                        ( ssList(X5)
                       => ! [X6] :
                            ( ~ ssList(X6)
                            | app(X5,X6) != X4
                            | app(X6,X5) != X3 ) )
                    | ! [X7] :
                        ( ~ ssItem(X7)
                        | ( ~ memberP(X2,X7)
                          & ~ memberP(X1,X7) )
                        | ( memberP(X2,X7)
                          & memberP(X1,X7) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_4"></A>fof(c_0_4,negated_conjecture,
    ( ssList(esk48_0)
    & ssList(esk49_0)
    & ssList(esk50_0)
    & ssList(esk51_0)
    & esk49_0 = esk51_0
    & esk48_0 = esk50_0
    & ssList(esk52_0)
    & ssList(esk53_0)
    & app(esk52_0,esk53_0) = esk51_0
    & app(esk53_0,esk52_0) = esk50_0
    & ssItem(esk54_0)
    & ( memberP(esk49_0,esk54_0)
      | memberP(esk48_0,esk54_0) )
    & ( ~ memberP(esk49_0,esk54_0)
      | ~ memberP(esk48_0,esk54_0) ) ),
    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).

<A NAME="c_0_5"></A>fof(c_0_5,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( app(esk53_0,esk52_0) = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_8"></A>fof(c_0_8,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(app(X5,X6),X4)
        | memberP(X5,X4)
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X5,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).

<A NAME="c_0_9"></A>cnf(c_0_9,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( app(esk53_0,esk52_0) = esk48_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>,<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( ssList(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( ssList(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( app(esk52_0,esk53_0) = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,plain,
    ( memberP(app(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(X3,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( memberP(esk48_0,esk54_0)
    | memberP(esk49_0,esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( ssItem(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( app(esk53_0,app(esk52_0,X1)) = app(esk48_0,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_10]),c_0_11]),c_0_12])])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( app(esk52_0,esk53_0) = esk49_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( memberP(app(X1,esk49_0),esk54_0)
    | memberP(esk48_0,esk54_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_16]),c_0_17]),c_0_18])])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( app(esk53_0,esk49_0) = app(esk48_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_20]),c_0_12])])).

<A NAME="c_0_23"></A>cnf(c_0_23,plain,
    ( memberP(X3,X1)
    | memberP(X2,X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(app(X2,X3),X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( memberP(app(esk48_0,esk53_0),esk54_0)
    | memberP(esk48_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_22]),c_0_12])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,plain,
    ( memberP(app(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(X2,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( memberP(esk48_0,esk54_0)
    | memberP(esk53_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_24]),c_0_12]),c_0_25]),c_0_18])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( memberP(app(esk53_0,X1),esk54_0)
    | memberP(esk48_0,esk54_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_27]),c_0_12]),c_0_18])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( memberP(esk48_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_10]),c_0_11])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( app(esk52_0,app(esk53_0,X1)) = app(esk49_0,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_20]),c_0_12]),c_0_11])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( memberP(app(X1,esk48_0),esk54_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_29]),c_0_25]),c_0_18])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( app(esk52_0,esk48_0) = app(esk49_0,esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_10]),c_0_11])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( ~ memberP(esk48_0,esk54_0)
    | ~ memberP(esk49_0,esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( memberP(app(esk49_0,esk52_0),esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_32]),c_0_11])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( ~ memberP(esk49_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_29])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( memberP(esk52_0,esk54_0) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_34]),c_0_11]),c_0_17]),c_0_18])]),<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( memberP(app(esk52_0,X1),esk54_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_36]),c_0_11]),c_0_18])])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_20]),c_0_12])]),<A HREF="#c_0_35">c_0_35</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC378+1 : TPTP v6.4.0. Released v2.4.0.
% 0.02/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n148.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:42:53 CDT 2016
% 0.02/0.22  % CPUTime    : 
% 0.02/0.22  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.28  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.02/0.28  # and selection function SelectCQPrecWNTNp.
% 0.02/0.28  #
% 0.02/0.28  # Preprocessing time       : 0.014 s
% 0.02/0.28  # Presaturation interreduction done
% 0.02/0.28  
% 0.02/0.28  # Proof found!
% 0.02/0.28  # SZS status Theorem
% 0.02/0.28  # SZS output start CNFRefutation
% 0.02/0.28  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~(ssList(X4))|X2!=X4)|X1!=X3)|![X5]:(ssList(X5)=>![X6]:((~(ssList(X6))|app(X5,X6)!=X4)|app(X6,X5)!=X3)))|![X7]:((~(ssItem(X7))|(~(memberP(X2,X7))&~(memberP(X1,X7))))|(memberP(X2,X7)&memberP(X1,X7)))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 0.02/0.28  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.02/0.28  fof(c_0_2, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(memberP(app(X2,X3),X1)&lt;=>(memberP(X2,X1)|memberP(X3,X1))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax36)).
% 0.02/0.28  fof(c_0_3, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~ssList(X4)|X2!=X4)|X1!=X3)|![X5]:(ssList(X5)=>![X6]:((~ssList(X6)|app(X5,X6)!=X4)|app(X6,X5)!=X3)))|![X7]:((~ssItem(X7)|(~memberP(X2,X7)&~memberP(X1,X7)))|(memberP(X2,X7)&memberP(X1,X7))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.02/0.28  fof(c_0_4, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(ssList(esk52_0)&((ssList(esk53_0)&app(esk52_0,esk53_0)=esk51_0)&app(esk53_0,esk52_0)=esk50_0)))&((ssItem(esk54_0)&(memberP(esk49_0,esk54_0)|memberP(esk48_0,esk54_0)))&(~memberP(esk49_0,esk54_0)|~memberP(esk48_0,esk54_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
% 0.02/0.28  fof(c_0_5, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.02/0.28  cnf(c_0_6, negated_conjecture, (app(esk53_0,esk52_0)=esk50_0), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.28  cnf(c_0_7, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.28  fof(c_0_8, plain, (![X4]:![X5]:![X6]:(((((~memberP(app(X5,X6),X4)|(memberP(X5,X4)|memberP(X6,X4)))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&(((((~memberP(X5,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
% 0.02/0.28  cnf(c_0_9, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.02/0.28  cnf(c_0_10, negated_conjecture, (app(esk53_0,esk52_0)=esk48_0), inference(rw,[status(thm)],[c_0_6, c_0_7])).
% 0.02/0.28  cnf(c_0_11, negated_conjecture, (ssList(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.28  cnf(c_0_12, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.28  cnf(c_0_13, negated_conjecture, (app(esk52_0,esk53_0)=esk51_0), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.28  cnf(c_0_14, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.28  cnf(c_0_15, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.02/0.28  cnf(c_0_16, negated_conjecture, (memberP(esk48_0,esk54_0)|memberP(esk49_0,esk54_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.28  cnf(c_0_17, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.28  cnf(c_0_18, negated_conjecture, (ssItem(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.28  cnf(c_0_19, negated_conjecture, (app(esk53_0,app(esk52_0,X1))=app(esk48_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12])])).
% 0.02/0.28  cnf(c_0_20, negated_conjecture, (app(esk52_0,esk53_0)=esk49_0), inference(rw,[status(thm)],[c_0_13, c_0_14])).
% 0.02/0.28  cnf(c_0_21, negated_conjecture, (memberP(app(X1,esk49_0),esk54_0)|memberP(esk48_0,esk54_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])])).
% 0.02/0.28  cnf(c_0_22, negated_conjecture, (app(esk53_0,esk49_0)=app(esk48_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_12])])).
% 0.02/0.28  cnf(c_0_23, plain, (memberP(X3,X1)|memberP(X2,X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(app(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.02/0.28  cnf(c_0_24, negated_conjecture, (memberP(app(esk48_0,esk53_0),esk54_0)|memberP(esk48_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_12])])).
% 0.02/0.28  cnf(c_0_25, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.28  cnf(c_0_26, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.02/0.28  cnf(c_0_27, negated_conjecture, (memberP(esk48_0,esk54_0)|memberP(esk53_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_12]), c_0_25]), c_0_18])])).
% 0.02/0.28  cnf(c_0_28, negated_conjecture, (memberP(app(esk53_0,X1),esk54_0)|memberP(esk48_0,esk54_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_12]), c_0_18])])).
% 0.02/0.28  cnf(c_0_29, negated_conjecture, (memberP(esk48_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_10]), c_0_11])])).
% 0.02/0.28  cnf(c_0_30, negated_conjecture, (app(esk52_0,app(esk53_0,X1))=app(esk49_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_20]), c_0_12]), c_0_11])])).
% 0.02/0.28  cnf(c_0_31, negated_conjecture, (memberP(app(X1,esk48_0),esk54_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_29]), c_0_25]), c_0_18])])).
% 0.02/0.28  cnf(c_0_32, negated_conjecture, (app(esk52_0,esk48_0)=app(esk49_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_10]), c_0_11])])).
% 0.02/0.28  cnf(c_0_33, negated_conjecture, (~memberP(esk48_0,esk54_0)|~memberP(esk49_0,esk54_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
% 0.02/0.28  cnf(c_0_34, negated_conjecture, (memberP(app(esk49_0,esk52_0),esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_11])])).
% 0.02/0.28  cnf(c_0_35, negated_conjecture, (~memberP(esk49_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_29])])).
% 0.02/0.28  cnf(c_0_36, negated_conjecture, (memberP(esk52_0,esk54_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_34]), c_0_11]), c_0_17]), c_0_18])]), c_0_35])).
% 0.02/0.28  cnf(c_0_37, negated_conjecture, (memberP(app(esk52_0,X1),esk54_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_36]), c_0_11]), c_0_18])])).
% 0.02/0.28  cnf(c_0_38, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_20]), c_0_12])]), c_0_35]), ['proof']).
% 0.02/0.28  # SZS output end CNFRefutation
% 0.02/0.28  # Training examples: 0 positive, 0 negative
% 0.02/0.28  
% 0.02/0.28  # -------------------------------------------------
% 0.02/0.28  # User time                : 0.054 s
% 0.02/0.28  # System time              : 0.004 s
% 0.02/0.28  # Total time               : 0.058 s
% 0.02/0.28  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
