<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC198+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC198+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC198+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC198%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC198+1.p">SWC198+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n112.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:01 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :   28 (  64 expanded)
%              Number of clauses        :   17 (  23 expanded)
%              Number of leaves         :    5 (  15 expanded)
%              Depth                    :    9
%              Number of atoms          :  142 ( 807 expanded)
%              Number of equality atoms :   45 ( 308 expanded)
%              Maximal formula depth    :   24 (   6 average)
%              Maximal clause size      :   34 (   2 average)
%              Maximal term depth       :    2 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ? [X5] :
                        ( ssItem(X5)
                        & ! [X6] :
                            ( ssItem(X6)
                           => ( ~ memberP(X1,X6)
                              | X5 = X6 ) ) )
                    | ( ! [X7] :
                          ( ssItem(X7)
                         => ( cons(X7,nil) != X3
                            | ~ memberP(X4,X7)
                            | ? [X8] :
                                ( ssItem(X8)
                                & X7 != X8
                                & memberP(X4,X8)
                                & leq(X7,X8) ) ) )
                      & ( nil != X4
                        | nil != X3 ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax38)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(cons(X2,X3),X1)
              &lt;=> ( X1 = X2
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax37)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ? [X1] :
      ( ssItem(X1)
      & ? [X2] :
          ( ssItem(X2)
          & X1 != X2 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax2)).

<A NAME="c_0_5"></A>fof(c_0_5,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ? [X5] :
                          ( ssItem(X5)
                          & ! [X6] :
                              ( ssItem(X6)
                             => ( ~ memberP(X1,X6)
                                | X5 = X6 ) ) )
                      | ( ! [X7] :
                            ( ssItem(X7)
                           => ( cons(X7,nil) != X3
                              | ~ memberP(X4,X7)
                              | ? [X8] :
                                  ( ssItem(X8)
                                  & X7 != X8
                                  & memberP(X4,X8)
                                  & leq(X7,X8) ) ) )
                        & ( nil != X4
                          | nil != X3 ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_6"></A>fof(c_0_6,plain,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_7"></A>fof(c_0_7,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(cons(X5,X6),X4)
        | X4 = X5
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( X4 != X5
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).

<A NAME="c_0_8"></A>fof(c_0_8,negated_conjecture,(
    ! [X13,X16] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ( ssItem(esk52_1(X13))
        | ~ ssItem(X13) )
      & ( memberP(esk48_0,esk52_1(X13))
        | ~ ssItem(X13) )
      & ( X13 != esk52_1(X13)
        | ~ ssItem(X13) )
      & ( nil = esk51_0
        | ssItem(esk53_0) )
      & ( nil = esk50_0
        | ssItem(esk53_0) )
      & ( nil = esk51_0
        | cons(esk53_0,nil) = esk50_0 )
      & ( nil = esk50_0
        | cons(esk53_0,nil) = esk50_0 )
      & ( nil = esk51_0
        | memberP(esk51_0,esk53_0) )
      & ( nil = esk50_0
        | memberP(esk51_0,esk53_0) )
      & ( nil = esk51_0
        | ~ ssItem(X16)
        | esk53_0 = X16
        | ~ memberP(esk51_0,X16)
        | ~ leq(esk53_0,X16) )
      & ( nil = esk50_0
        | ~ ssItem(X16)
        | esk53_0 = X16
        | ~ memberP(esk51_0,X16)
        | ~ leq(esk53_0,X16) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).

<A NAME="c_0_9"></A>fof(c_0_9,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | ~ memberP(nil,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).

<A NAME="c_0_10"></A>cnf(c_0_10,plain,
    ( memberP(X3,X1)
    | X1 = X2
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ memberP(cons(X2,X3),X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( cons(esk53_0,nil) = esk50_0
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_12"></A>cnf(c_0_12,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( ssItem(esk53_0)
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,plain,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( memberP(esk48_0,esk52_1(X1))
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( esk50_0 = nil
    | X1 = esk53_0
    | ~ memberP(esk50_0,X1)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_11]),c_0_12])]),c_0_13]),<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( memberP(esk50_0,esk52_1(X1))
    | ~ ssItem(X1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( ssItem(esk52_1(X1))
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( ~ ssItem(X1)
    | X1 != esk52_1(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( esk52_1(X1) = esk53_0
    | esk50_0 = nil
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_18]),<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_21])]),<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_23"></A>fof(c_0_23,plain,
    ( ssItem(esk1_0)
    & ssItem(esk2_0)
    & esk1_0 != esk2_0 ),
    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_4])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( memberP(nil,esk52_1(X1))
    | ~ ssItem(X1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,plain,
    ( ssItem(esk1_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_24]),<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,plain,
    ( $false ),
    inference(sr,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,<A HREF="#c_0_26">c_0_26</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC198+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n112.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:37:22 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.36  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.36  # and selection function SelectCQPrecWNTNp.
% 0.06/0.36  #
% 0.06/0.36  # Preprocessing time       : 0.014 s
% 0.06/0.36  # Presaturation interreduction done
% 0.06/0.36  
% 0.06/0.36  # Proof found!
% 0.06/0.36  # SZS status Theorem
% 0.06/0.36  # SZS output start CNFRefutation
% 0.06/0.36  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|?[X5]:(ssItem(X5)&![X6]:(ssItem(X6)=>(~(memberP(X1,X6))|X5=X6))))|(![X7]:(ssItem(X7)=>((cons(X7,nil)!=X3|~(memberP(X4,X7)))|?[X8]:(((ssItem(X8)&X7!=X8)&memberP(X4,X8))&leq(X7,X8))))&(nil!=X4|nil!=X3)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.06/0.36  fof(c_0_1, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax38)).
% 0.06/0.36  fof(c_0_2, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)&lt;=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax37)).
% 0.06/0.36  fof(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.36  fof(c_0_4, axiom, (?[X1]:(ssItem(X1)&?[X2]:(ssItem(X2)&X1!=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax2)).
% 0.06/0.36  fof(c_0_5, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|?[X5]:(ssItem(X5)&![X6]:(ssItem(X6)=>(~memberP(X1,X6)|X5=X6))))|(![X7]:(ssItem(X7)=>((cons(X7,nil)!=X3|~memberP(X4,X7))|?[X8]:(((ssItem(X8)&X7!=X8)&memberP(X4,X8))&leq(X7,X8))))&(nil!=X4|nil!=X3))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.36  fof(c_0_6, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_1])).
% 0.06/0.36  fof(c_0_7, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
% 0.06/0.36  fof(c_0_8, negated_conjecture, (![X13]:![X16]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((esk49_0=esk51_0&esk48_0=esk50_0)&((ssItem(esk52_1(X13))|~ssItem(X13))&((memberP(esk48_0,esk52_1(X13))|~ssItem(X13))&(X13!=esk52_1(X13)|~ssItem(X13)))))&(((nil=esk51_0|ssItem(esk53_0))&(nil=esk50_0|ssItem(esk53_0)))&((((nil=esk51_0|cons(esk53_0,nil)=esk50_0)&(nil=esk50_0|cons(esk53_0,nil)=esk50_0))&((nil=esk51_0|memberP(esk51_0,esk53_0))&(nil=esk50_0|memberP(esk51_0,esk53_0))))&((nil=esk51_0|(((~ssItem(X16)|esk53_0=X16)|~memberP(esk51_0,X16))|~leq(esk53_0,X16)))&(nil=esk50_0|(((~ssItem(X16)|esk53_0=X16)|~memberP(esk51_0,X16))|~leq(esk53_0,X16)))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
% 0.06/0.36  fof(c_0_9, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
% 0.06/0.36  cnf(c_0_10, plain, (memberP(X3,X1)|X1=X2|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~memberP(cons(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
% 0.06/0.36  cnf(c_0_11, negated_conjecture, (cons(esk53_0,nil)=esk50_0|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.36  cnf(c_0_12, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.36  cnf(c_0_13, negated_conjecture, (ssItem(esk53_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.36  cnf(c_0_14, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.06/0.36  cnf(c_0_15, negated_conjecture, (memberP(esk48_0,esk52_1(X1))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.36  cnf(c_0_16, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.36  cnf(c_0_17, negated_conjecture, (esk50_0=nil|X1=esk53_0|~memberP(esk50_0,X1)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])]), c_0_13]), c_0_14])).
% 0.06/0.36  cnf(c_0_18, negated_conjecture, (memberP(esk50_0,esk52_1(X1))|~ssItem(X1)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
% 0.06/0.36  cnf(c_0_19, negated_conjecture, (ssItem(esk52_1(X1))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.36  cnf(c_0_20, negated_conjecture, (~ssItem(X1)|X1!=esk52_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.06/0.36  cnf(c_0_21, negated_conjecture, (esk52_1(X1)=esk53_0|esk50_0=nil|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
% 0.06/0.36  cnf(c_0_22, negated_conjecture, (esk50_0=nil), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21])]), c_0_13])).
% 0.06/0.36  fof(c_0_23, plain, ((ssItem(esk1_0)&(ssItem(esk2_0)&esk1_0!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[c_0_4])])).
% 0.06/0.36  cnf(c_0_24, negated_conjecture, (memberP(nil,esk52_1(X1))|~ssItem(X1)), inference(rw,[status(thm)],[c_0_18, c_0_22])).
% 0.06/0.36  cnf(c_0_25, plain, (ssItem(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 0.06/0.36  cnf(c_0_26, negated_conjecture, (~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_24]), c_0_19])).
% 0.06/0.36  cnf(c_0_27, plain, ($false), inference(sr,[status(thm)],[c_0_25, c_0_26]), ['proof']).
% 0.06/0.36  # SZS output end CNFRefutation
% 0.06/0.36  # Training examples: 0 positive, 0 negative
% 0.06/0.36  
% 0.06/0.36  # -------------------------------------------------
% 0.06/0.36  # User time                : 0.128 s
% 0.06/0.36  # System time              : 0.004 s
% 0.06/0.36  # Total time               : 0.131 s
% 0.06/0.36  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
