<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC297+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC297+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC297+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC297%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC297+1.p">SWC297+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n139.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:23 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :  184 (6352 expanded)
%              Number of clauses        :  120 (2362 expanded)
%              Number of leaves         :   31 (1622 expanded)
%              Depth                    :   22
%              Number of atoms          :  721 (73676 expanded)
%              Number of equality atoms :  212 (19014 expanded)
%              Maximal formula depth    :   25 (   5 average)
%              Maximal clause size      :   30 (   3 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ! [X5] :
                        ( ssItem(X5)
                       => ! [X6] :
                            ( ssList(X6)
                           => ! [X7] :
                                ( ssList(X7)
                               => ( app(app(X6,cons(X5,nil)),X7) != X1
                                  | ! [X8] :
                                      ( ssItem(X8)
                                     => ( ( ~ memberP(X6,X8)
                                          | lt(X8,X5) )
                                        & ( ~ memberP(X7,X8)
                                          | lt(X5,X8) ) ) ) ) ) ) )
                    | ( nil != X3
                      & nil = X4 )
                    | ( ! [X9] :
                          ( ssItem(X9)
                         => ( cons(X9,nil) != X3
                            | ~ memberP(X4,X9) ) )
                      & neq(X4,nil) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( frontsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X2,X3) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax5)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => hd(cons(X2,X1)) = X2 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax23)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => cons(hd(X1),tl(X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax78)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( strictorderedP(X1)
      &lt;=> ! [X2] :
            ( ssItem(X2)
           => ! [X3] :
                ( ssItem(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ! [X5] :
                        ( ssList(X5)
                       => ! [X6] :
                            ( ssList(X6)
                           => ( app(app(X4,cons(X2,X5)),cons(X3,X6)) = X1
                             => lt(X2,X3) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax12)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( equalelemsP(X1)
      &lt;=> ! [X2] :
            ( ssItem(X2)
           => ! [X3] :
                ( ssItem(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ! [X5] :
                        ( ssList(X5)
                       => ( app(X4,cons(X2,cons(X3,X5))) = X1
                         => X2 = X3 ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax14)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X3,X2) = app(X1,X2)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax79)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssItem(hd(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax22)).

<A NAME="c_0_17"></A>fof(c_0_17,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax38)).

<A NAME="c_0_18"></A>fof(c_0_18,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( memberP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(X3,cons(X2,X4)) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax3)).

<A NAME="c_0_19"></A>fof(c_0_19,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil = X1
        | ? [X2] :
            ( ssList(X2)
            & ? [X3] :
                ( ssItem(X3)
                & cons(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax20)).

<A NAME="c_0_20"></A>fof(c_0_20,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( rearsegP(X1,X2)
               => rearsegP(app(X3,X1),X2) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax50)).

<A NAME="c_0_21"></A>fof(c_0_21,axiom,(
    ! [X1] :
      ( ssList(X1)
     => rearsegP(X1,nil) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax51)).

<A NAME="c_0_22"></A>fof(c_0_22,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(cons(X2,X3),X1)
              &lt;=> ( X1 = X2
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax37)).

<A NAME="c_0_23"></A>fof(c_0_23,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => strictorderedP(cons(X1,nil)) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax68)).

<A NAME="c_0_24"></A>fof(c_0_24,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => equalelemsP(cons(X1,nil)) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax73)).

<A NAME="c_0_25"></A>fof(c_0_25,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( frontsegP(X1,X2)
               => frontsegP(app(X1,X3),X2) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax43)).

<A NAME="c_0_26"></A>fof(c_0_26,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,nil) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax45)).

<A NAME="c_0_27"></A>fof(c_0_27,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( rearsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax6)).

<A NAME="c_0_28"></A>fof(c_0_28,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ~ lt(X1,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax90)).

<A NAME="c_0_29"></A>fof(c_0_29,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_30"></A>fof(c_0_30,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => ! [X4] :
                  ( ssItem(X4)
                 => ( cons(X3,X1) = cons(X4,X2)
                   => ( X3 = X4
                      & X2 = X1 ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax19)).

<A NAME="c_0_31"></A>fof(c_0_31,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ! [X5] :
                          ( ssItem(X5)
                         => ! [X6] :
                              ( ssList(X6)
                             => ! [X7] :
                                  ( ssList(X7)
                                 => ( app(app(X6,cons(X5,nil)),X7) != X1
                                    | ! [X8] :
                                        ( ssItem(X8)
                                       => ( ( ~ memberP(X6,X8)
                                            | lt(X8,X5) )
                                          & ( ~ memberP(X7,X8)
                                            | lt(X5,X8) ) ) ) ) ) ) )
                      | ( nil != X3
                        & nil = X4 )
                      | ( ! [X9] :
                            ( ssItem(X9)
                           => ( cons(X9,nil) != X3
                              | ~ memberP(X4,X9) ) )
                        & neq(X4,nil) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_32"></A>fof(c_0_32,negated_conjecture,
    ( ssList(esk48_0)
    & ssList(esk49_0)
    & ssList(esk50_0)
    & ssList(esk51_0)
    & esk49_0 = esk51_0
    & esk48_0 = esk50_0
    & ssItem(esk52_0)
    & ssList(esk53_0)
    & ssList(esk54_0)
    & app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk48_0
    & ssItem(esk55_0)
    & ( memberP(esk54_0,esk55_0)
      | memberP(esk53_0,esk55_0) )
    & ( ~ lt(esk52_0,esk55_0)
      | memberP(esk53_0,esk55_0) )
    & ( memberP(esk54_0,esk55_0)
      | ~ lt(esk55_0,esk52_0) )
    & ( ~ lt(esk52_0,esk55_0)
      | ~ lt(esk55_0,esk52_0) )
    & ( nil = esk50_0
      | nil != esk51_0 )
    & ( ssItem(esk56_0)
      | ~ neq(esk51_0,nil) )
    & ( cons(esk56_0,nil) = esk50_0
      | ~ neq(esk51_0,nil) )
    & ( memberP(esk51_0,esk56_0)
      | ~ neq(esk51_0,nil) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_35"></A>fof(c_0_35,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( ssList(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( ssList(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_40"></A>fof(c_0_40,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_41"></A>fof(c_0_41,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_44"></A>fof(c_0_44,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk6_2(X4,X5))
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(X5,esk6_2(X4,X5)) = X4
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X5,X7) != X4
        | frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).

<A NAME="c_0_45"></A>fof(c_0_45,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( app(esk53_0,app(cons(esk52_0,nil),esk54_0)) = esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_37]),c_0_38]),c_0_39])])).

<A NAME="c_0_47"></A>cnf(c_0_47,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_48"></A>cnf(c_0_48,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( ssItem(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_50"></A>fof(c_0_50,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).

<A NAME="c_0_51"></A>fof(c_0_51,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | hd(cons(X4,X3)) = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( cons(esk56_0,nil) = esk50_0
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,<A HREF="#c_0_43">c_0_43</A>])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( ssItem(esk56_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( frontsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X2,X3) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( app(esk53_0,app(cons(esk52_0,nil),esk54_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_47]),c_0_48]),c_0_49])])).

<A NAME="c_0_59"></A>cnf(c_0_59,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_60"></A>fof(c_0_60,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | cons(hd(X2),tl(X2)) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).

<A NAME="c_0_61"></A>cnf(c_0_61,plain,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( cons(esk56_0,nil) = esk50_0
    | esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_53]),c_0_48]),c_0_54])])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( esk51_0 = nil
    | ssItem(esk56_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_53]),c_0_48]),c_0_54])])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_65"></A>fof(c_0_65,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).

<A NAME="c_0_66"></A>fof(c_0_66,plain,(
    ! [X7,X8,X9,X10,X11,X12] :
      ( ( ~ strictorderedP(X7)
        | ~ ssItem(X8)
        | ~ ssItem(X9)
        | ~ ssList(X10)
        | ~ ssList(X11)
        | ~ ssList(X12)
        | app(app(X10,cons(X8,X11)),cons(X9,X12)) != X7
        | lt(X8,X9)
        | ~ ssList(X7) )
      & ( ssItem(esk30_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( ssItem(esk31_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk32_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk33_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk34_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( app(app(esk32_1(X7),cons(esk30_1(X7),esk33_1(X7))),cons(esk31_1(X7),esk34_1(X7))) = X7
        | strictorderedP(X7)
        | ~ ssList(X7) )
      & ( ~ lt(esk30_1(X7),esk31_1(X7))
        | strictorderedP(X7)
        | ~ ssList(X7) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).

<A NAME="c_0_67"></A>fof(c_0_67,plain,(
    ! [X6,X7,X8,X9,X10] :
      ( ( ~ equalelemsP(X6)
        | ~ ssItem(X7)
        | ~ ssItem(X8)
        | ~ ssList(X9)
        | ~ ssList(X10)
        | app(X9,cons(X7,cons(X8,X10))) != X6
        | X7 = X8
        | ~ ssList(X6) )
      & ( ssItem(esk40_1(X6))
        | equalelemsP(X6)
        | ~ ssList(X6) )
      & ( ssItem(esk41_1(X6))
        | equalelemsP(X6)
        | ~ ssList(X6) )
      & ( ssList(esk42_1(X6))
        | equalelemsP(X6)
        | ~ ssList(X6) )
      & ( ssList(esk43_1(X6))
        | equalelemsP(X6)
        | ~ ssList(X6) )
      & ( app(esk42_1(X6),cons(esk40_1(X6),cons(esk41_1(X6),esk43_1(X6)))) = X6
        | equalelemsP(X6)
        | ~ ssList(X6) )
      & ( esk40_1(X6) != esk41_1(X6)
        | equalelemsP(X6)
        | ~ ssList(X6) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).

<A NAME="c_0_68"></A>fof(c_0_68,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_56]),<A HREF="#c_0_57">c_0_57</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( app(esk53_0,cons(esk52_0,esk54_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_59]),c_0_38]),c_0_49])])).

<A NAME="c_0_71"></A>fof(c_0_71,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X6,X5) != app(X4,X5)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).

<A NAME="c_0_72"></A>fof(c_0_72,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).

<A NAME="c_0_73"></A>cnf(c_0_73,plain,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( hd(esk50_0) = esk56_0
    | esk51_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_62]),c_0_48])]),<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,negated_conjecture,
    ( nil = esk50_0
    | nil != esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>])).

<A NAME="c_0_78"></A>fof(c_0_78,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssItem(hd(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).

<A NAME="c_0_79"></A>fof(c_0_79,plain,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>])).

<A NAME="c_0_80"></A>fof(c_0_80,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk3_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ssList(esk4_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6))) = X5
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(X9,cons(X6,X10)) != X5
        | memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( lt(X2,X3)
    | ~ ssList(X1)
    | app(app(X4,cons(X2,X5)),cons(X3,X6)) != X1
    | ~ ssList(X6)
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssItem(X3)
    | ~ ssItem(X2)
    | ~ strictorderedP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>])).

<A NAME="c_0_82"></A>fof(c_0_82,plain,(
    ! [X4] :
      ( ( ssList(esk44_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( ssItem(esk45_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( cons(esk45_1(X4),esk44_1(X4)) = X4
        | nil = X4
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).

<A NAME="c_0_83"></A>cnf(c_0_83,plain,
    ( X2 = X3
    | ~ ssList(X1)
    | app(X4,cons(X2,cons(X3,X5))) != X1
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssItem(X3)
    | ~ ssItem(X2)
    | ~ equalelemsP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,plain,
    ( nil = X2
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_85"></A>fof(c_0_85,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ rearsegP(X4,X5)
      | rearsegP(app(X6,X4),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).

<A NAME="c_0_86"></A>fof(c_0_86,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | rearsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( frontsegP(esk50_0,esk53_0)
    | ~ ssList(cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>,c_0_70]),c_0_39])])).

<A NAME="c_0_88"></A>cnf(c_0_88,plain,
    ( X1 = X2
    | app(X1,X3) != app(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>])).

<A NAME="c_0_90"></A>fof(c_0_90,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(cons(X5,X6),X4)
        | X4 = X5
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( X4 != X5
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).

<A NAME="c_0_91"></A>cnf(c_0_91,negated_conjecture,
    ( cons(esk56_0,tl(esk50_0)) = esk50_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>,c_0_74]),c_0_75])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_92"></A>cnf(c_0_92,negated_conjecture,
    ( tl(esk50_0) = nil
    | esk51_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_62]),c_0_48])]),<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_93"></A>cnf(c_0_93,plain,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>])).

<A NAME="c_0_94"></A>fof(c_0_94,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | ~ memberP(nil,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_79])])).

<A NAME="c_0_95"></A>cnf(c_0_95,plain,
    ( memberP(X1,X2)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | app(X3,cons(X2,X4)) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>])).

<A NAME="c_0_96"></A>cnf(c_0_96,plain,
    ( lt(X1,X2)
    | ~ strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,plain,
    ( nil = X1
    | cons(esk45_1(X1),esk44_1(X1)) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>])).

<A NAME="c_0_98"></A>cnf(c_0_98,plain,
    ( nil = X1
    | ssItem(esk45_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( nil = X1
    | ssList(esk44_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>])).

<A NAME="c_0_100"></A>fof(c_0_100,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | strictorderedP(cons(X2,nil)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])).

<A NAME="c_0_101"></A>cnf(c_0_101,plain,
    ( X1 = X2
    | ~ equalelemsP(app(X3,cons(X1,cons(X2,X4))))
    | ~ ssList(app(X3,cons(X1,cons(X2,X4))))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>])).

<A NAME="c_0_102"></A>fof(c_0_102,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | equalelemsP(cons(X2,nil)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])).

<A NAME="c_0_103"></A>cnf(c_0_103,negated_conjecture,
    ( esk54_0 = nil
    | esk50_0 != nil
    | ~ ssList(app(esk53_0,cons(esk52_0,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,c_0_36]),c_0_38])])).

<A NAME="c_0_104"></A>cnf(c_0_104,plain,
    ( rearsegP(app(X1,X2),X3)
    | ~ rearsegP(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>])).

<A NAME="c_0_105"></A>cnf(c_0_105,plain,
    ( rearsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>])).

<A NAME="c_0_106"></A>cnf(c_0_106,plain,
    ( app(X2,esk6_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_107"></A>cnf(c_0_107,negated_conjecture,
    ( frontsegP(esk50_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,c_0_47]),c_0_38]),c_0_49])])).

<A NAME="c_0_108"></A>cnf(c_0_108,plain,
    ( ssList(esk6_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_109"></A>fof(c_0_109,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ frontsegP(X4,X5)
      | frontsegP(app(X4,X6),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).

<A NAME="c_0_110"></A>fof(c_0_110,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])).

<A NAME="c_0_111"></A>cnf(c_0_111,plain,
    ( X1 = nil
    | app(X1,X2) != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>,c_0_89]),c_0_48])])).

<A NAME="c_0_112"></A>cnf(c_0_112,plain,
    ( memberP(X3,X1)
    | X1 = X2
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ memberP(cons(X2,X3),X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>])).

<A NAME="c_0_113"></A>cnf(c_0_113,negated_conjecture,
    ( cons(esk56_0,nil) = esk50_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_92]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_114"></A>cnf(c_0_114,negated_conjecture,
    ( esk50_0 = nil
    | ssItem(esk56_0) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_74]),c_0_75])]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,plain,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>])).

<A NAME="c_0_116"></A>cnf(c_0_116,plain,
    ( memberP(app(X1,cons(X2,X3)),X2)
    | ~ ssList(app(X1,cons(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>])).

<A NAME="c_0_117"></A>cnf(c_0_117,plain,
    ( nil = X1
    | lt(X2,esk45_1(X1))
    | ~ strictorderedP(app(app(X3,cons(X2,X4)),X1))
    | ~ ssList(app(app(X3,cons(X2,X4)),X1))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,c_0_97]),c_0_98]),<A HREF="#c_0_99">c_0_99</A>])).

<A NAME="c_0_118"></A>cnf(c_0_118,plain,
    ( strictorderedP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>])).

<A NAME="c_0_119"></A>cnf(c_0_119,plain,
    ( X1 = esk45_1(X2)
    | nil = X2
    | ~ equalelemsP(app(X3,cons(X1,X2)))
    | ~ ssList(app(X3,cons(X1,X2)))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_101">c_0_101</A>,c_0_97]),c_0_98]),<A HREF="#c_0_99">c_0_99</A>])).

<A NAME="c_0_120"></A>cnf(c_0_120,plain,
    ( equalelemsP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>])).

<A NAME="c_0_121"></A>cnf(c_0_121,negated_conjecture,
    ( esk54_0 = nil
    | esk50_0 != nil
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_103">c_0_103</A>,c_0_57]),c_0_39])])).

<A NAME="c_0_122"></A>fof(c_0_122,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk7_2(X4,X5))
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(esk7_2(X4,X5),X5) = X4
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X7,X5) != X4
        | rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])])).

<A NAME="c_0_123"></A>cnf(c_0_123,plain,
    ( rearsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>,c_0_105]),c_0_48])])).

<A NAME="c_0_124"></A>cnf(c_0_124,negated_conjecture,
    ( app(esk53_0,esk6_2(esk50_0,esk53_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>,c_0_107]),c_0_39]),c_0_75])])).

<A NAME="c_0_125"></A>cnf(c_0_125,negated_conjecture,
    ( ssList(esk6_2(esk50_0,esk53_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>,c_0_107]),c_0_39]),c_0_75])])).

<A NAME="c_0_126"></A>cnf(c_0_126,plain,
    ( frontsegP(app(X1,X2),X3)
    | ~ frontsegP(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>])).

<A NAME="c_0_127"></A>cnf(c_0_127,plain,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_128"></A>cnf(c_0_128,negated_conjecture,
    ( esk53_0 = nil
    | app(cons(esk52_0,nil),esk54_0) != esk50_0
    | ~ ssList(app(cons(esk52_0,nil),esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_111">c_0_111</A>,c_0_58]),c_0_39])])).

<A NAME="c_0_129"></A>cnf(c_0_129,negated_conjecture,
    ( esk50_0 = nil
    | X1 = esk56_0
    | ~ memberP(esk50_0,X1)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_113]),c_0_48])]),c_0_114]),<A HREF="#c_0_115">c_0_115</A>])).

<A NAME="c_0_130"></A>cnf(c_0_130,negated_conjecture,
    ( memberP(esk50_0,esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_116">c_0_116</A>,c_0_70]),c_0_75]),c_0_38]),c_0_39]),c_0_49])])).

<A NAME="c_0_131"></A>cnf(c_0_131,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_132"></A>fof(c_0_132,plain,(
    ! [X1] :
      ( ssItem(X1)
     => ~ lt(X1,X1) ) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_133"></A>cnf(c_0_133,negated_conjecture,
    ( esk54_0 = nil
    | lt(esk52_0,esk45_1(esk54_0))
    | ~ strictorderedP(esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_117">c_0_117</A>,c_0_36]),c_0_75]),c_0_48]),c_0_39]),c_0_38]),c_0_49])])).

<A NAME="c_0_134"></A>cnf(c_0_134,negated_conjecture,
    ( esk51_0 = nil
    | strictorderedP(esk50_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_118">c_0_118</A>,c_0_62]),<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_135"></A>cnf(c_0_135,negated_conjecture,
    ( esk45_1(esk54_0) = esk52_0
    | esk54_0 = nil
    | ~ equalelemsP(esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>,c_0_70]),c_0_75]),c_0_39]),c_0_38]),c_0_49])])).

<A NAME="c_0_136"></A>cnf(c_0_136,negated_conjecture,
    ( esk50_0 = nil
    | equalelemsP(esk50_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_120">c_0_120</A>,c_0_113]),<A HREF="#c_0_114">c_0_114</A>])).

<A NAME="c_0_137"></A>cnf(c_0_137,negated_conjecture,
    ( esk54_0 = nil
    | esk50_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_121">c_0_121</A>,c_0_47]),c_0_48]),c_0_49])])).

<A NAME="c_0_138"></A>fof(c_0_138,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])).

<A NAME="c_0_139"></A>cnf(c_0_139,plain,
    ( app(esk7_2(X1,X2),X2) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_122">c_0_122</A>])).

<A NAME="c_0_140"></A>cnf(c_0_140,negated_conjecture,
    ( rearsegP(esk50_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>,c_0_124]),c_0_125]),c_0_39])])).

<A NAME="c_0_141"></A>cnf(c_0_141,plain,
    ( ssList(esk7_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_122">c_0_122</A>])).

<A NAME="c_0_142"></A>cnf(c_0_142,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_126">c_0_126</A>,c_0_127]),c_0_48])])).

<A NAME="c_0_143"></A>cnf(c_0_143,negated_conjecture,
    ( esk51_0 = nil
    | X1 = esk56_0
    | ~ memberP(esk50_0,X1)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_62]),c_0_48])]),c_0_63]),<A HREF="#c_0_115">c_0_115</A>])).

<A NAME="c_0_144"></A>cnf(c_0_144,negated_conjecture,
    ( esk53_0 = nil
    | app(cons(esk52_0,nil),esk54_0) != esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_128">c_0_128</A>,c_0_57]),c_0_38])])).

<A NAME="c_0_145"></A>cnf(c_0_145,negated_conjecture,
    ( esk56_0 = esk52_0
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_129">c_0_129</A>,c_0_130]),c_0_49])])).

<A NAME="c_0_146"></A>cnf(c_0_146,negated_conjecture,
    ( esk53_0 = nil
    | esk50_0 != nil
    | ~ ssList(cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>,c_0_70]),c_0_39])])).

<A NAME="c_0_147"></A>fof(c_0_147,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | ~ lt(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_132])])).

<A NAME="c_0_148"></A>cnf(c_0_148,negated_conjecture,
    ( esk51_0 = nil
    | esk54_0 = nil
    | lt(esk52_0,esk45_1(esk54_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_133">c_0_133</A>,<A HREF="#c_0_134">c_0_134</A>])).

<A NAME="c_0_149"></A>cnf(c_0_149,negated_conjecture,
    ( esk45_1(esk54_0) = esk52_0
    | esk54_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_135">c_0_135</A>,c_0_136]),<A HREF="#c_0_137">c_0_137</A>])).

<A NAME="c_0_150"></A>cnf(c_0_150,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_138">c_0_138</A>])).

<A NAME="c_0_151"></A>cnf(c_0_151,negated_conjecture,
    ( app(esk7_2(esk50_0,nil),nil) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_139">c_0_139</A>,c_0_140]),c_0_48]),c_0_75])])).

<A NAME="c_0_152"></A>cnf(c_0_152,negated_conjecture,
    ( ssList(esk7_2(esk50_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_141">c_0_141</A>,c_0_140]),c_0_48]),c_0_75])])).

<A NAME="c_0_153"></A>cnf(c_0_153,negated_conjecture,
    ( frontsegP(esk50_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_142">c_0_142</A>,c_0_124]),c_0_125]),c_0_39])])).

<A NAME="c_0_154"></A>cnf(c_0_154,negated_conjecture,
    ( esk56_0 = esk52_0
    | esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_143">c_0_143</A>,c_0_130]),c_0_49])])).

<A NAME="c_0_155"></A>cnf(c_0_155,negated_conjecture,
    ( esk53_0 = nil
    | app(cons(esk52_0,nil),esk54_0) != esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_144">c_0_144</A>,c_0_47]),c_0_48]),c_0_49])])).

<A NAME="c_0_156"></A>cnf(c_0_156,negated_conjecture,
    ( cons(esk52_0,nil) = esk50_0
    | esk50_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_113">c_0_113</A>,<A HREF="#c_0_145">c_0_145</A>])).

<A NAME="c_0_157"></A>cnf(c_0_157,negated_conjecture,
    ( esk53_0 = nil
    | esk50_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_146">c_0_146</A>,c_0_47]),c_0_38]),c_0_49])])).

<A NAME="c_0_158"></A>cnf(c_0_158,plain,
    ( ~ lt(X1,X1)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_147">c_0_147</A>])).

<A NAME="c_0_159"></A>cnf(c_0_159,negated_conjecture,
    ( esk54_0 = nil
    | esk51_0 = nil
    | lt(esk52_0,esk52_0) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_148">c_0_148</A>,<A HREF="#c_0_149">c_0_149</A>])).

<A NAME="c_0_160"></A>cnf(c_0_160,negated_conjecture,
    ( esk7_2(esk50_0,nil) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_150">c_0_150</A>,c_0_151]),c_0_152])])).

<A NAME="c_0_161"></A>cnf(c_0_161,negated_conjecture,
    ( app(nil,esk6_2(esk50_0,nil)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>,c_0_153]),c_0_48]),c_0_75])])).

<A NAME="c_0_162"></A>cnf(c_0_162,negated_conjecture,
    ( ssList(esk6_2(esk50_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>,c_0_153]),c_0_48]),c_0_75])])).

<A NAME="c_0_163"></A>cnf(c_0_163,negated_conjecture,
    ( cons(esk52_0,nil) = esk50_0
    | esk51_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,<A HREF="#c_0_154">c_0_154</A>])).

<A NAME="c_0_164"></A>cnf(c_0_164,negated_conjecture,
    ( esk53_0 = nil
    | app(esk50_0,esk54_0) != esk50_0 ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_155">c_0_155</A>,c_0_156]),<A HREF="#c_0_157">c_0_157</A>])).

<A NAME="c_0_165"></A>cnf(c_0_165,negated_conjecture,
    ( esk51_0 = nil
    | esk54_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_158">c_0_158</A>,c_0_159]),c_0_49])])).

<A NAME="c_0_166"></A>cnf(c_0_166,negated_conjecture,
    ( app(esk50_0,nil) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_151">c_0_151</A>,<A HREF="#c_0_160">c_0_160</A>])).

<A NAME="c_0_167"></A>cnf(c_0_167,negated_conjecture,
    ( esk6_2(esk50_0,nil) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>,c_0_161]),c_0_162])])).

<A NAME="c_0_168"></A>cnf(c_0_168,negated_conjecture,
    ( app(app(esk53_0,esk50_0),esk54_0) = esk50_0
    | esk51_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,<A HREF="#c_0_163">c_0_163</A>])).

<A NAME="c_0_169"></A>cnf(c_0_169,negated_conjecture,
    ( esk51_0 = nil
    | esk53_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_164">c_0_164</A>,c_0_165]),c_0_166])])).

<A NAME="c_0_170"></A>cnf(c_0_170,negated_conjecture,
    ( app(nil,esk50_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_161">c_0_161</A>,<A HREF="#c_0_167">c_0_167</A>])).

<A NAME="c_0_171"></A>fof(c_0_171,plain,(
    ! [X5,X6,X7,X8] :
      ( ( X7 = X8
        | cons(X7,X5) != cons(X8,X6)
        | ~ ssItem(X8)
        | ~ ssItem(X7)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( X6 = X5
        | cons(X7,X5) != cons(X8,X6)
        | ~ ssItem(X8)
        | ~ ssItem(X7)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])).

<A NAME="c_0_172"></A>cnf(c_0_172,negated_conjecture,
    ( app(esk50_0,X1) = cons(esk56_0,X1)
    | esk51_0 = nil
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,c_0_62]),<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_173"></A>cnf(c_0_173,negated_conjecture,
    ( app(esk50_0,esk54_0) = esk50_0
    | esk51_0 = nil ),
    inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_168">c_0_168</A>,c_0_169]),<A HREF="#c_0_170">c_0_170</A>])).

<A NAME="c_0_174"></A>cnf(c_0_174,plain,
    ( X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssItem(X3)
    | ~ ssItem(X4)
    | cons(X3,X1) != cons(X4,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_171">c_0_171</A>])).

<A NAME="c_0_175"></A>cnf(c_0_175,negated_conjecture,
    ( cons(esk56_0,esk54_0) = esk50_0
    | esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_172">c_0_172</A>,c_0_173]),c_0_38])])).

<A NAME="c_0_176"></A>cnf(c_0_176,negated_conjecture,
    ( esk50_0 = nil
    | nil = X1
    | cons(X2,X1) != esk50_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_174">c_0_174</A>,c_0_113]),c_0_48])]),<A HREF="#c_0_114">c_0_114</A>])).

<A NAME="c_0_177"></A>cnf(c_0_177,negated_conjecture,
    ( cons(esk52_0,esk54_0) = esk50_0
    | esk51_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_175">c_0_175</A>,<A HREF="#c_0_154">c_0_154</A>])).

<A NAME="c_0_178"></A>cnf(c_0_178,negated_conjecture,
    ( esk54_0 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_176">c_0_176</A>,c_0_177]),c_0_38]),c_0_49])]),c_0_76]),<A HREF="#c_0_137">c_0_137</A>])).

<A NAME="c_0_179"></A>cnf(c_0_179,negated_conjecture,
    ( memberP(esk53_0,esk55_0)
    | memberP(esk54_0,esk55_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_180"></A>cnf(c_0_180,negated_conjecture,
    ( esk53_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_164">c_0_164</A>,c_0_178]),c_0_166])])).

<A NAME="c_0_181"></A>cnf(c_0_181,negated_conjecture,
    ( memberP(nil,esk55_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_179">c_0_179</A>,c_0_178]),c_0_180])])).

<A NAME="c_0_182"></A>cnf(c_0_182,negated_conjecture,
    ( ssItem(esk55_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_183"></A>cnf(c_0_183,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_115">c_0_115</A>,c_0_181]),c_0_182])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC297+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n139.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:39:52 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.98  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.98  # and selection function SelectCQPrecWNTNp.
% 0.06/0.98  #
% 0.06/0.98  # Preprocessing time       : 0.014 s
% 0.06/0.98  # Presaturation interreduction done
% 0.06/0.98  
% 0.06/0.98  # Proof found!
% 0.06/0.98  # SZS status Theorem
% 0.06/0.98  # SZS output start CNFRefutation
% 0.06/0.98  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>(app(app(X6,cons(X5,nil)),X7)!=X1|![X8]:(ssItem(X8)=>((~(memberP(X6,X8))|lt(X8,X5))&(~(memberP(X7,X8))|lt(X5,X8)))))))))|(nil!=X3&nil=X4))|(![X9]:(ssItem(X9)=>(cons(X9,nil)!=X3|~(memberP(X4,X9))))&neq(X4,nil)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.06/0.98  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.06/0.98  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.06/0.98  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.06/0.98  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax5)).
% 0.06/0.98  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.06/0.98  fof(c_0_6, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.98  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.06/0.98  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax23)).
% 0.06/0.98  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax78)).
% 0.06/0.98  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
% 0.06/0.98  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>(strictorderedP(X1)&lt;=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(app(X4,cons(X2,X5)),cons(X3,X6))=X1=>lt(X2,X3)))))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax12)).
% 0.06/0.98  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>(equalelemsP(X1)&lt;=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>(app(X4,cons(X2,cons(X3,X5)))=X1=>X2=X3)))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax14)).
% 0.06/0.98  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.06/0.98  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax79)).
% 0.06/0.98  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
% 0.06/0.98  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax22)).
% 0.06/0.98  fof(c_0_17, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax38)).
% 0.06/0.98  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>(memberP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(X3,cons(X2,X4))=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax3)).
% 0.06/0.98  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax20)).
% 0.06/0.98  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(rearsegP(X1,X2)=>rearsegP(app(X3,X1),X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax50)).
% 0.06/0.98  fof(c_0_21, axiom, (![X1]:(ssList(X1)=>rearsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax51)).
% 0.06/0.98  fof(c_0_22, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)&lt;=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax37)).
% 0.06/0.98  fof(c_0_23, axiom, (![X1]:(ssItem(X1)=>strictorderedP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax68)).
% 0.06/0.98  fof(c_0_24, axiom, (![X1]:(ssItem(X1)=>equalelemsP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax73)).
% 0.06/0.98  fof(c_0_25, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(frontsegP(X1,X2)=>frontsegP(app(X1,X3),X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax43)).
% 0.06/0.98  fof(c_0_26, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax45)).
% 0.06/0.98  fof(c_0_27, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax6)).
% 0.06/0.98  fof(c_0_28, axiom, (![X1]:(ssItem(X1)=>~(lt(X1,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax90)).
% 0.06/0.98  fof(c_0_29, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
% 0.06/0.98  fof(c_0_30, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssItem(X4)=>(cons(X3,X1)=cons(X4,X2)=>(X3=X4&X2=X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax19)).
% 0.06/0.98  fof(c_0_31, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>(app(app(X6,cons(X5,nil)),X7)!=X1|![X8]:(ssItem(X8)=>((~memberP(X6,X8)|lt(X8,X5))&(~memberP(X7,X8)|lt(X5,X8)))))))))|(nil!=X3&nil=X4))|(![X9]:(ssItem(X9)=>(cons(X9,nil)!=X3|~memberP(X4,X9)))&neq(X4,nil))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.98  fof(c_0_32, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(ssItem(esk52_0)&(ssList(esk53_0)&(ssList(esk54_0)&(app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk48_0&(ssItem(esk55_0)&(((memberP(esk54_0,esk55_0)|memberP(esk53_0,esk55_0))&(~lt(esk52_0,esk55_0)|memberP(esk53_0,esk55_0)))&((memberP(esk54_0,esk55_0)|~lt(esk55_0,esk52_0))&(~lt(esk52_0,esk55_0)|~lt(esk55_0,esk52_0))))))))))&(nil=esk50_0|nil!=esk51_0))&((ssItem(esk56_0)|~neq(esk51_0,nil))&((cons(esk56_0,nil)=esk50_0|~neq(esk51_0,nil))&(memberP(esk51_0,esk56_0)|~neq(esk51_0,nil)))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])])).
% 0.06/0.98  cnf(c_0_33, negated_conjecture, (app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  cnf(c_0_34, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  fof(c_0_35, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.06/0.98  cnf(c_0_36, negated_conjecture, (app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk50_0), inference(rw,[status(thm)],[c_0_33, c_0_34])).
% 0.06/0.98  cnf(c_0_37, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
% 0.06/0.98  cnf(c_0_38, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  cnf(c_0_39, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  fof(c_0_40, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.06/0.98  fof(c_0_41, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
% 0.06/0.98  cnf(c_0_42, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  cnf(c_0_43, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  fof(c_0_44, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
% 0.06/0.98  fof(c_0_45, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 0.06/0.98  cnf(c_0_46, negated_conjecture, (app(esk53_0,app(cons(esk52_0,nil),esk54_0))=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39])])).
% 0.06/0.98  cnf(c_0_47, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
% 0.06/0.98  cnf(c_0_48, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
% 0.06/0.98  cnf(c_0_49, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  fof(c_0_50, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
% 0.06/0.98  fof(c_0_51, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
% 0.06/0.98  cnf(c_0_52, negated_conjecture, (cons(esk56_0,nil)=esk50_0|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  cnf(c_0_53, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 0.06/0.98  cnf(c_0_54, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_42, c_0_43])).
% 0.06/0.98  cnf(c_0_55, negated_conjecture, (ssItem(esk56_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  cnf(c_0_56, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
% 0.06/0.98  cnf(c_0_57, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
% 0.06/0.98  cnf(c_0_58, negated_conjecture, (app(esk53_0,app(cons(esk52_0,nil),esk54_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49])])).
% 0.06/0.98  cnf(c_0_59, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
% 0.06/0.98  fof(c_0_60, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).
% 0.06/0.98  cnf(c_0_61, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
% 0.06/0.98  cnf(c_0_62, negated_conjecture, (cons(esk56_0,nil)=esk50_0|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_48]), c_0_54])])).
% 0.06/0.98  cnf(c_0_63, negated_conjecture, (esk51_0=nil|ssItem(esk56_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_53]), c_0_48]), c_0_54])])).
% 0.06/0.98  cnf(c_0_64, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  fof(c_0_65, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
% 0.06/0.98  fof(c_0_66, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((~strictorderedP(X7)|(~ssItem(X8)|(~ssItem(X9)|(~ssList(X10)|(~ssList(X11)|(~ssList(X12)|(app(app(X10,cons(X8,X11)),cons(X9,X12))!=X7|lt(X8,X9))))))))|~ssList(X7))&(((ssItem(esk30_1(X7))|strictorderedP(X7))|~ssList(X7))&(((ssItem(esk31_1(X7))|strictorderedP(X7))|~ssList(X7))&(((ssList(esk32_1(X7))|strictorderedP(X7))|~ssList(X7))&(((ssList(esk33_1(X7))|strictorderedP(X7))|~ssList(X7))&(((ssList(esk34_1(X7))|strictorderedP(X7))|~ssList(X7))&(((app(app(esk32_1(X7),cons(esk30_1(X7),esk33_1(X7))),cons(esk31_1(X7),esk34_1(X7)))=X7|strictorderedP(X7))|~ssList(X7))&((~lt(esk30_1(X7),esk31_1(X7))|strictorderedP(X7))|~ssList(X7)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
% 0.06/0.98  fof(c_0_67, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((~equalelemsP(X6)|(~ssItem(X7)|(~ssItem(X8)|(~ssList(X9)|(~ssList(X10)|(app(X9,cons(X7,cons(X8,X10)))!=X6|X7=X8))))))|~ssList(X6))&(((ssItem(esk40_1(X6))|equalelemsP(X6))|~ssList(X6))&(((ssItem(esk41_1(X6))|equalelemsP(X6))|~ssList(X6))&(((ssList(esk42_1(X6))|equalelemsP(X6))|~ssList(X6))&(((ssList(esk43_1(X6))|equalelemsP(X6))|~ssList(X6))&(((app(esk42_1(X6),cons(esk40_1(X6),cons(esk41_1(X6),esk43_1(X6))))=X6|equalelemsP(X6))|~ssList(X6))&((esk40_1(X6)!=esk41_1(X6)|equalelemsP(X6))|~ssList(X6))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).
% 0.06/0.98  fof(c_0_68, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).
% 0.06/0.98  cnf(c_0_69, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_56]), c_0_57])).
% 0.06/0.98  cnf(c_0_70, negated_conjecture, (app(esk53_0,cons(esk52_0,esk54_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_38]), c_0_49])])).
% 0.06/0.98  fof(c_0_71, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
% 0.06/0.98  fof(c_0_72, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
% 0.06/0.98  cnf(c_0_73, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
% 0.06/0.98  cnf(c_0_74, negated_conjecture, (hd(esk50_0)=esk56_0|esk51_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_48])]), c_0_63])).
% 0.06/0.98  cnf(c_0_75, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_64, c_0_34])).
% 0.06/0.98  cnf(c_0_76, negated_conjecture, (nil=esk50_0|nil!=esk51_0), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  cnf(c_0_77, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
% 0.06/0.98  fof(c_0_78, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).
% 0.06/0.98  fof(c_0_79, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_17])).
% 0.06/0.98  fof(c_0_80, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk3_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&((((ssList(esk4_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&(((app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6)))=X5|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(X9,cons(X6,X10))!=X5))|memberP(X5,X6))|~ssItem(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])])).
% 0.06/0.98  cnf(c_0_81, plain, (lt(X2,X3)|~ssList(X1)|app(app(X4,cons(X2,X5)),cons(X3,X6))!=X1|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssItem(X3)|~ssItem(X2)|~strictorderedP(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
% 0.06/0.98  fof(c_0_82, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
% 0.06/0.98  cnf(c_0_83, plain, (X2=X3|~ssList(X1)|app(X4,cons(X2,cons(X3,X5)))!=X1|~ssList(X5)|~ssList(X4)|~ssItem(X3)|~ssItem(X2)|~equalelemsP(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
% 0.06/0.98  cnf(c_0_84, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
% 0.06/0.98  fof(c_0_85, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~rearsegP(X4,X5)|rearsegP(app(X6,X4),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
% 0.06/0.98  fof(c_0_86, plain, (![X2]:(~ssList(X2)|rearsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).
% 0.06/0.98  cnf(c_0_87, negated_conjecture, (frontsegP(esk50_0,esk53_0)|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_39])])).
% 0.06/0.98  cnf(c_0_88, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
% 0.06/0.98  cnf(c_0_89, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
% 0.06/0.98  fof(c_0_90, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).
% 0.06/0.98  cnf(c_0_91, negated_conjecture, (cons(esk56_0,tl(esk50_0))=esk50_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75])]), c_0_76])).
% 0.06/0.98  cnf(c_0_92, negated_conjecture, (tl(esk50_0)=nil|esk51_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_62]), c_0_48])]), c_0_63])).
% 0.06/0.98  cnf(c_0_93, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
% 0.06/0.98  fof(c_0_94, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_79])])).
% 0.06/0.98  cnf(c_0_95, plain, (memberP(X1,X2)|~ssList(X1)|~ssItem(X2)|app(X3,cons(X2,X4))!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_80])).
% 0.06/0.98  cnf(c_0_96, plain, (lt(X1,X2)|~strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_81])).
% 0.06/0.98  cnf(c_0_97, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
% 0.06/0.98  cnf(c_0_98, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
% 0.06/0.98  cnf(c_0_99, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
% 0.06/0.98  fof(c_0_100, plain, (![X2]:(~ssItem(X2)|strictorderedP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])).
% 0.06/0.98  cnf(c_0_101, plain, (X1=X2|~equalelemsP(app(X3,cons(X1,cons(X2,X4))))|~ssList(app(X3,cons(X1,cons(X2,X4))))|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_83])).
% 0.06/0.98  fof(c_0_102, plain, (![X2]:(~ssItem(X2)|equalelemsP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])).
% 0.06/0.98  cnf(c_0_103, negated_conjecture, (esk54_0=nil|esk50_0!=nil|~ssList(app(esk53_0,cons(esk52_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_36]), c_0_38])])).
% 0.06/0.98  cnf(c_0_104, plain, (rearsegP(app(X1,X2),X3)|~rearsegP(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
% 0.06/0.98  cnf(c_0_105, plain, (rearsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
% 0.06/0.98  cnf(c_0_106, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
% 0.06/0.98  cnf(c_0_107, negated_conjecture, (frontsegP(esk50_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_47]), c_0_38]), c_0_49])])).
% 0.06/0.98  cnf(c_0_108, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
% 0.06/0.98  fof(c_0_109, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~frontsegP(X4,X5)|frontsegP(app(X4,X6),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
% 0.06/0.98  fof(c_0_110, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])).
% 0.06/0.98  cnf(c_0_111, plain, (X1=nil|app(X1,X2)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_48])])).
% 0.06/0.98  cnf(c_0_112, plain, (memberP(X3,X1)|X1=X2|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~memberP(cons(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
% 0.06/0.98  cnf(c_0_113, negated_conjecture, (cons(esk56_0,nil)=esk50_0|esk50_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_76])).
% 0.06/0.98  cnf(c_0_114, negated_conjecture, (esk50_0=nil|ssItem(esk56_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_74]), c_0_75])]), c_0_76])).
% 0.06/0.98  cnf(c_0_115, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
% 0.06/0.98  cnf(c_0_116, plain, (memberP(app(X1,cons(X2,X3)),X2)|~ssList(app(X1,cons(X2,X3)))|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(er,[status(thm)],[c_0_95])).
% 0.06/0.98  cnf(c_0_117, plain, (nil=X1|lt(X2,esk45_1(X1))|~strictorderedP(app(app(X3,cons(X2,X4)),X1))|~ssList(app(app(X3,cons(X2,X4)),X1))|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98]), c_0_99])).
% 0.06/0.98  cnf(c_0_118, plain, (strictorderedP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
% 0.06/0.98  cnf(c_0_119, plain, (X1=esk45_1(X2)|nil=X2|~equalelemsP(app(X3,cons(X1,X2)))|~ssList(app(X3,cons(X1,X2)))|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_97]), c_0_98]), c_0_99])).
% 0.06/0.98  cnf(c_0_120, plain, (equalelemsP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
% 0.06/0.98  cnf(c_0_121, negated_conjecture, (esk54_0=nil|esk50_0!=nil|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_57]), c_0_39])])).
% 0.06/0.98  fof(c_0_122, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])])).
% 0.06/0.98  cnf(c_0_123, plain, (rearsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_48])])).
% 0.06/0.98  cnf(c_0_124, negated_conjecture, (app(esk53_0,esk6_2(esk50_0,esk53_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_39]), c_0_75])])).
% 0.06/0.98  cnf(c_0_125, negated_conjecture, (ssList(esk6_2(esk50_0,esk53_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_107]), c_0_39]), c_0_75])])).
% 0.06/0.98  cnf(c_0_126, plain, (frontsegP(app(X1,X2),X3)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
% 0.06/0.98  cnf(c_0_127, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
% 0.06/0.98  cnf(c_0_128, negated_conjecture, (esk53_0=nil|app(cons(esk52_0,nil),esk54_0)!=esk50_0|~ssList(app(cons(esk52_0,nil),esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_58]), c_0_39])])).
% 0.06/0.98  cnf(c_0_129, negated_conjecture, (esk50_0=nil|X1=esk56_0|~memberP(esk50_0,X1)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_48])]), c_0_114]), c_0_115])).
% 0.06/0.98  cnf(c_0_130, negated_conjecture, (memberP(esk50_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_70]), c_0_75]), c_0_38]), c_0_39]), c_0_49])])).
% 0.06/0.98  cnf(c_0_131, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
% 0.06/0.98  fof(c_0_132, plain, (![X1]:(ssItem(X1)=>~lt(X1,X1))), inference(fof_simplification,[status(thm)],[c_0_28])).
% 0.06/0.98  cnf(c_0_133, negated_conjecture, (esk54_0=nil|lt(esk52_0,esk45_1(esk54_0))|~strictorderedP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_36]), c_0_75]), c_0_48]), c_0_39]), c_0_38]), c_0_49])])).
% 0.06/0.98  cnf(c_0_134, negated_conjecture, (esk51_0=nil|strictorderedP(esk50_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_62]), c_0_63])).
% 0.06/0.98  cnf(c_0_135, negated_conjecture, (esk45_1(esk54_0)=esk52_0|esk54_0=nil|~equalelemsP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_70]), c_0_75]), c_0_39]), c_0_38]), c_0_49])])).
% 0.06/0.98  cnf(c_0_136, negated_conjecture, (esk50_0=nil|equalelemsP(esk50_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_113]), c_0_114])).
% 0.06/0.98  cnf(c_0_137, negated_conjecture, (esk54_0=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_47]), c_0_48]), c_0_49])])).
% 0.06/0.98  fof(c_0_138, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])).
% 0.06/0.98  cnf(c_0_139, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_122])).
% 0.06/0.98  cnf(c_0_140, negated_conjecture, (rearsegP(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_125]), c_0_39])])).
% 0.06/0.98  cnf(c_0_141, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_122])).
% 0.06/0.98  cnf(c_0_142, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_48])])).
% 0.06/0.98  cnf(c_0_143, negated_conjecture, (esk51_0=nil|X1=esk56_0|~memberP(esk50_0,X1)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_62]), c_0_48])]), c_0_63]), c_0_115])).
% 0.06/0.98  cnf(c_0_144, negated_conjecture, (esk53_0=nil|app(cons(esk52_0,nil),esk54_0)!=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_57]), c_0_38])])).
% 0.06/0.98  cnf(c_0_145, negated_conjecture, (esk56_0=esk52_0|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_49])])).
% 0.06/0.98  cnf(c_0_146, negated_conjecture, (esk53_0=nil|esk50_0!=nil|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_70]), c_0_39])])).
% 0.06/0.98  fof(c_0_147, plain, (![X2]:(~ssItem(X2)|~lt(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_132])])).
% 0.06/0.98  cnf(c_0_148, negated_conjecture, (esk51_0=nil|esk54_0=nil|lt(esk52_0,esk45_1(esk54_0))), inference(spm,[status(thm)],[c_0_133, c_0_134])).
% 0.06/0.98  cnf(c_0_149, negated_conjecture, (esk45_1(esk54_0)=esk52_0|esk54_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_136]), c_0_137])).
% 0.06/0.98  cnf(c_0_150, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_138])).
% 0.06/0.98  cnf(c_0_151, negated_conjecture, (app(esk7_2(esk50_0,nil),nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_48]), c_0_75])])).
% 0.06/0.98  cnf(c_0_152, negated_conjecture, (ssList(esk7_2(esk50_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_140]), c_0_48]), c_0_75])])).
% 0.06/0.98  cnf(c_0_153, negated_conjecture, (frontsegP(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_124]), c_0_125]), c_0_39])])).
% 0.06/0.98  cnf(c_0_154, negated_conjecture, (esk56_0=esk52_0|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_130]), c_0_49])])).
% 0.06/0.98  cnf(c_0_155, negated_conjecture, (esk53_0=nil|app(cons(esk52_0,nil),esk54_0)!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_47]), c_0_48]), c_0_49])])).
% 0.06/0.98  cnf(c_0_156, negated_conjecture, (cons(esk52_0,nil)=esk50_0|esk50_0=nil), inference(spm,[status(thm)],[c_0_113, c_0_145])).
% 0.06/0.98  cnf(c_0_157, negated_conjecture, (esk53_0=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_47]), c_0_38]), c_0_49])])).
% 0.06/0.98  cnf(c_0_158, plain, (~lt(X1,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_147])).
% 0.06/0.98  cnf(c_0_159, negated_conjecture, (esk54_0=nil|esk51_0=nil|lt(esk52_0,esk52_0)), inference(spm,[status(thm)],[c_0_148, c_0_149])).
% 0.06/0.98  cnf(c_0_160, negated_conjecture, (esk7_2(esk50_0,nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_152])])).
% 0.06/0.98  cnf(c_0_161, negated_conjecture, (app(nil,esk6_2(esk50_0,nil))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_153]), c_0_48]), c_0_75])])).
% 0.06/0.98  cnf(c_0_162, negated_conjecture, (ssList(esk6_2(esk50_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_153]), c_0_48]), c_0_75])])).
% 0.06/0.98  cnf(c_0_163, negated_conjecture, (cons(esk52_0,nil)=esk50_0|esk51_0=nil), inference(spm,[status(thm)],[c_0_62, c_0_154])).
% 0.06/0.98  cnf(c_0_164, negated_conjecture, (esk53_0=nil|app(esk50_0,esk54_0)!=esk50_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156]), c_0_157])).
% 0.06/0.98  cnf(c_0_165, negated_conjecture, (esk51_0=nil|esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_49])])).
% 0.06/0.98  cnf(c_0_166, negated_conjecture, (app(esk50_0,nil)=esk50_0), inference(rw,[status(thm)],[c_0_151, c_0_160])).
% 0.06/0.98  cnf(c_0_167, negated_conjecture, (esk6_2(esk50_0,nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_161]), c_0_162])])).
% 0.06/0.98  cnf(c_0_168, negated_conjecture, (app(app(esk53_0,esk50_0),esk54_0)=esk50_0|esk51_0=nil), inference(spm,[status(thm)],[c_0_36, c_0_163])).
% 0.06/0.98  cnf(c_0_169, negated_conjecture, (esk51_0=nil|esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_166])])).
% 0.06/0.98  cnf(c_0_170, negated_conjecture, (app(nil,esk50_0)=esk50_0), inference(rw,[status(thm)],[c_0_161, c_0_167])).
% 0.06/0.98  fof(c_0_171, plain, (![X5]:![X6]:![X7]:![X8]:((((((X7=X8|cons(X7,X5)!=cons(X8,X6))|~ssItem(X8))|~ssItem(X7))|~ssList(X6))|~ssList(X5))&(((((X6=X5|cons(X7,X5)!=cons(X8,X6))|~ssItem(X8))|~ssItem(X7))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])).
% 0.06/0.98  cnf(c_0_172, negated_conjecture, (app(esk50_0,X1)=cons(esk56_0,X1)|esk51_0=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_62]), c_0_63])).
% 0.06/0.98  cnf(c_0_173, negated_conjecture, (app(esk50_0,esk54_0)=esk50_0|esk51_0=nil), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169]), c_0_170])).
% 0.06/0.98  cnf(c_0_174, plain, (X2=X1|~ssList(X1)|~ssList(X2)|~ssItem(X3)|~ssItem(X4)|cons(X3,X1)!=cons(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_171])).
% 0.06/0.98  cnf(c_0_175, negated_conjecture, (cons(esk56_0,esk54_0)=esk50_0|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173]), c_0_38])])).
% 0.06/0.98  cnf(c_0_176, negated_conjecture, (esk50_0=nil|nil=X1|cons(X2,X1)!=esk50_0|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_113]), c_0_48])]), c_0_114])).
% 0.06/0.98  cnf(c_0_177, negated_conjecture, (cons(esk52_0,esk54_0)=esk50_0|esk51_0=nil), inference(spm,[status(thm)],[c_0_175, c_0_154])).
% 0.06/0.98  cnf(c_0_178, negated_conjecture, (esk54_0=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177]), c_0_38]), c_0_49])]), c_0_76]), c_0_137])).
% 0.06/0.98  cnf(c_0_179, negated_conjecture, (memberP(esk53_0,esk55_0)|memberP(esk54_0,esk55_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  cnf(c_0_180, negated_conjecture, (esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_164, c_0_178]), c_0_166])])).
% 0.06/0.98  cnf(c_0_181, negated_conjecture, (memberP(nil,esk55_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_179, c_0_178]), c_0_180])])).
% 0.06/0.98  cnf(c_0_182, negated_conjecture, (ssItem(esk55_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.06/0.98  cnf(c_0_183, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_181]), c_0_182])]), ['proof']).
% 0.06/0.98  # SZS output end CNFRefutation
% 0.06/0.98  # Training examples: 0 positive, 0 negative
% 0.06/0.98  
% 0.06/0.98  # -------------------------------------------------
% 0.06/0.98  # User time                : 0.732 s
% 0.06/0.98  # System time              : 0.021 s
% 0.06/0.98  # Total time               : 0.753 s
% 0.06/0.98  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
