<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC050+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC050+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC050+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC050%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC050+1.p">SWC050+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n150.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:26 EDT 2016

% Result     : Theorem 0.08s
% Output     : CNFRefutation 0.08s
% Verified   : 
% Statistics : Number of formulae       :   18 (  79 expanded)
%              Number of clauses        :   15 (  28 expanded)
%              Number of leaves         :    1 (  17 expanded)
%              Depth                    :    6
%              Number of atoms          :   97 (1234 expanded)
%              Number of equality atoms :    8 ( 109 expanded)
%              Maximal formula depth    :   21 (   5 average)
%              Maximal clause size      :   36 (   2 average)
%              Maximal term depth       :    1 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ( ( ~ neq(X2,nil)
                        | ? [X5] :
                            ( ssList(X5)
                            & neq(X5,nil)
                            & rearsegP(X2,X5)
                            & rearsegP(X1,X5) )
                        | ! [X6] :
                            ( ssList(X6)
                           => ( ~ neq(X6,nil)
                              | ~ rearsegP(X4,X6)
                              | ~ rearsegP(X3,X6) ) ) )
                      & ( ~ neq(X2,nil)
                        | neq(X4,nil) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ( ( ~ neq(X2,nil)
                          | ? [X5] :
                              ( ssList(X5)
                              & neq(X5,nil)
                              & rearsegP(X2,X5)
                              & rearsegP(X1,X5) )
                          | ! [X6] :
                              ( ssList(X6)
                             => ( ~ neq(X6,nil)
                                | ~ rearsegP(X4,X6)
                                | ~ rearsegP(X3,X6) ) ) )
                        & ( ~ neq(X2,nil)
                          | neq(X4,nil) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_2"></A>fof(c_0_2,negated_conjecture,(
    ! [X11] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ( neq(esk49_0,nil)
        | neq(esk49_0,nil) )
      & ( ~ neq(esk51_0,nil)
        | neq(esk49_0,nil) )
      & ( neq(esk49_0,nil)
        | ~ ssList(X11)
        | ~ neq(X11,nil)
        | ~ rearsegP(esk49_0,X11)
        | ~ rearsegP(esk48_0,X11) )
      & ( ~ neq(esk51_0,nil)
        | ~ ssList(X11)
        | ~ neq(X11,nil)
        | ~ rearsegP(esk49_0,X11)
        | ~ rearsegP(esk48_0,X11) )
      & ( neq(esk49_0,nil)
        | ssList(esk52_0) )
      & ( ~ neq(esk51_0,nil)
        | ssList(esk52_0) )
      & ( neq(esk49_0,nil)
        | neq(esk52_0,nil) )
      & ( ~ neq(esk51_0,nil)
        | neq(esk52_0,nil) )
      & ( neq(esk49_0,nil)
        | rearsegP(esk51_0,esk52_0) )
      & ( ~ neq(esk51_0,nil)
        | rearsegP(esk51_0,esk52_0) )
      & ( neq(esk49_0,nil)
        | rearsegP(esk50_0,esk52_0) )
      & ( ~ neq(esk51_0,nil)
        | rearsegP(esk50_0,esk52_0) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( neq(esk49_0,nil)
    | neq(esk49_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ~ rearsegP(esk48_0,X1)
    | ~ rearsegP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( neq(esk49_0,nil) ),
    inference(cn,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( rearsegP(esk50_0,esk52_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( rearsegP(esk51_0,esk52_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( ssList(esk52_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( neq(esk52_0,nil)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( ~ rearsegP(esk48_0,X1)
    | ~ rearsegP(esk49_0,X1)
    | ~ ssList(X1)
    | ~ neq(X1,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,c_0_5]),c_0_6])])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( rearsegP(esk48_0,esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_8]),c_0_5]),c_0_6])])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( rearsegP(esk49_0,esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,c_0_5]),c_0_5]),c_0_6])])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( ssList(esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_5]),c_0_6])])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( neq(esk52_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_5]),c_0_6])])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_13]),c_0_14]),c_0_15]),c_0_16])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC050+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.25  % Computer   : n150.star.cs.uiowa.edu
% 0.03/0.25  % Model      : x86_64 x86_64
% 0.03/0.25  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.25  % Memory     : 32218.75MB
% 0.03/0.25  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.25  % CPULimit   : 300
% 0.03/0.25  % DateTime   : Tue May 10 20:34:53 CDT 2016
% 0.03/0.25  % CPUTime    : 
% 0.03/0.25  # Version: 1.9.1pre011
% 0.03/0.25  # No SInE strategy applied
% 0.03/0.25  # Trying AutoSched0 for 151 seconds
% 0.08/0.30  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.08/0.30  # and selection function SelectCQPrecWNTNp.
% 0.08/0.30  #
% 0.08/0.30  # Preprocessing time       : 0.027 s
% 0.08/0.30  # Presaturation interreduction done
% 0.08/0.30  
% 0.08/0.30  # Proof found!
% 0.08/0.30  # SZS status Theorem
% 0.08/0.30  # SZS output start CNFRefutation
% 0.08/0.30  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((~(neq(X2,nil))|?[X5]:(((ssList(X5)&neq(X5,nil))&rearsegP(X2,X5))&rearsegP(X1,X5)))|![X6]:(ssList(X6)=>((~(neq(X6,nil))|~(rearsegP(X4,X6)))|~(rearsegP(X3,X6)))))&(~(neq(X2,nil))|neq(X4,nil))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.08/0.30  fof(c_0_1, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((~neq(X2,nil)|?[X5]:(((ssList(X5)&neq(X5,nil))&rearsegP(X2,X5))&rearsegP(X1,X5)))|![X6]:(ssList(X6)=>((~neq(X6,nil)|~rearsegP(X4,X6))|~rearsegP(X3,X6))))&(~neq(X2,nil)|neq(X4,nil)))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.08/0.30  fof(c_0_2, negated_conjecture, (![X11]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((esk49_0=esk51_0&esk48_0=esk50_0)&((((neq(esk49_0,nil)|neq(esk49_0,nil))&(~neq(esk51_0,nil)|neq(esk49_0,nil)))&((neq(esk49_0,nil)|(((~ssList(X11)|~neq(X11,nil))|~rearsegP(esk49_0,X11))|~rearsegP(esk48_0,X11)))&(~neq(esk51_0,nil)|(((~ssList(X11)|~neq(X11,nil))|~rearsegP(esk49_0,X11))|~rearsegP(esk48_0,X11)))))&(((neq(esk49_0,nil)|ssList(esk52_0))&(~neq(esk51_0,nil)|ssList(esk52_0)))&((((neq(esk49_0,nil)|neq(esk52_0,nil))&(~neq(esk51_0,nil)|neq(esk52_0,nil)))&((neq(esk49_0,nil)|rearsegP(esk51_0,esk52_0))&(~neq(esk51_0,nil)|rearsegP(esk51_0,esk52_0))))&((neq(esk49_0,nil)|rearsegP(esk50_0,esk52_0))&(~neq(esk51_0,nil)|rearsegP(esk50_0,esk52_0)))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
% 0.08/0.30  cnf(c_0_3, negated_conjecture, (neq(esk49_0,nil)|neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.08/0.30  cnf(c_0_4, negated_conjecture, (~rearsegP(esk48_0,X1)|~rearsegP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.08/0.30  cnf(c_0_5, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.08/0.30  cnf(c_0_6, negated_conjecture, (neq(esk49_0,nil)), inference(cn,[status(thm)],[c_0_3])).
% 0.08/0.30  cnf(c_0_7, negated_conjecture, (rearsegP(esk50_0,esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.08/0.30  cnf(c_0_8, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.08/0.30  cnf(c_0_9, negated_conjecture, (rearsegP(esk51_0,esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.08/0.30  cnf(c_0_10, negated_conjecture, (ssList(esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.08/0.30  cnf(c_0_11, negated_conjecture, (neq(esk52_0,nil)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.08/0.30  cnf(c_0_12, negated_conjecture, (~rearsegP(esk48_0,X1)|~rearsegP(esk49_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_5]), c_0_6])])).
% 0.08/0.30  cnf(c_0_13, negated_conjecture, (rearsegP(esk48_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_8]), c_0_5]), c_0_6])])).
% 0.08/0.30  cnf(c_0_14, negated_conjecture, (rearsegP(esk49_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_5]), c_0_5]), c_0_6])])).
% 0.08/0.30  cnf(c_0_15, negated_conjecture, (ssList(esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_5]), c_0_6])])).
% 0.08/0.30  cnf(c_0_16, negated_conjecture, (neq(esk52_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_5]), c_0_6])])).
% 0.08/0.30  cnf(c_0_17, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15]), c_0_16])]), ['proof']).
% 0.08/0.30  # SZS output end CNFRefutation
% 0.08/0.30  # Training examples: 0 positive, 0 negative
% 0.08/0.30  
% 0.08/0.30  # -------------------------------------------------
% 0.08/0.30  # User time                : 0.045 s
% 0.08/0.30  # System time              : 0.006 s
% 0.08/0.30  # Total time               : 0.051 s
% 0.08/0.30  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
