<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC062+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC062+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC062+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC062%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC062+1.p">SWC062+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n169.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:29 EDT 2016

% Result     : Theorem 225.39s
% Output     : CNFRefutation 225.39s
% Verified   : 
% Statistics : Number of formulae       :  164 (2612 expanded)
%              Number of clauses        :  112 ( 968 expanded)
%              Number of leaves         :   26 ( 640 expanded)
%              Depth                    :   16
%              Number of atoms          :  606 (29239 expanded)
%              Number of equality atoms :  218 (11739 expanded)
%              Maximal formula depth    :   20 (   5 average)
%              Maximal clause size      :   26 (   3 average)
%              Maximal term depth       :    4 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_4"></A>fof(c_0_4,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ? [X5] :
                        ( ssItem(X5)
                        & ? [X6] :
                            ( ssList(X6)
                            & app(X6,cons(X5,nil)) != X3
                            & app(cons(X5,nil),X6) = X4 ) )
                    | ( nil != X3
                      & nil = X4 )
                    | ( ( nil != X2
                        | nil = X1 )
                      & ( ~ neq(X2,nil)
                        | ? [X7] :
                            ( ssList(X7)
                            & neq(X7,nil)
                            & segmentP(X2,X7)
                            & segmentP(X1,X7) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,X1) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax55)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax27)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => hd(cons(X2,X1)) = X2 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax23)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil = X1
        | ? [X2] :
            ( ssList(X2)
            & ? [X3] :
                ( ssItem(X3)
                & cons(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax20)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( segmentP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax58)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ? [X2] :
            ( ssItem(X2)
            & hd(X1) = X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax75)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) != X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax18)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ? [X2] :
            ( ssList(X2)
            & tl(X1) = X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax76)).

<A NAME="c_0_17"></A>fof(c_0_17,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( segmentP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax7)).

<A NAME="c_0_18"></A>fof(c_0_18,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_19"></A>fof(c_0_19,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => cons(hd(X1),tl(X1)) = X1 ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax78)).

<A NAME="c_0_20"></A>fof(c_0_20,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssItem(hd(X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax22)).

<A NAME="c_0_21"></A>fof(c_0_21,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssList(tl(X1)) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax24)).

<A NAME="c_0_22"></A>fof(c_0_22,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_23"></A>fof(c_0_23,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( rearsegP(X1,X2)
               => rearsegP(app(X3,X1),X2) ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax50)).

<A NAME="c_0_24"></A>fof(c_0_24,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( rearsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax6)).

<A NAME="c_0_25"></A>fof(c_0_25,axiom,(
    ! [X1] :
      ( ssList(X1)
     => rearsegP(X1,X1) ) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax49)).

<A NAME="c_0_26"></A>fof(c_0_26,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])).

<A NAME="c_0_27"></A>fof(c_0_27,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).

<A NAME="c_0_28"></A>fof(c_0_28,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).

<A NAME="c_0_29"></A>cnf(c_0_29,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_32"></A>fof(c_0_32,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ? [X5] :
                          ( ssItem(X5)
                          & ? [X6] :
                              ( ssList(X6)
                              & app(X6,cons(X5,nil)) != X3
                              & app(cons(X5,nil),X6) = X4 ) )
                      | ( nil != X3
                        & nil = X4 )
                      | ( ( nil != X2
                          | nil = X1 )
                        & ( ~ neq(X2,nil)
                          | ? [X7] :
                              ( ssList(X7)
                              & neq(X7,nil)
                              & segmentP(X2,X7)
                              & segmentP(X1,X7) ) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_4])])).

<A NAME="c_0_33"></A>cnf(c_0_33,plain,
    ( nil = X2
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,plain,
    ( app(X1,app(nil,X2)) = app(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_30]),c_0_31])])).

<A NAME="c_0_35"></A>fof(c_0_35,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_36"></A>fof(c_0_36,negated_conjecture,(
    ! [X12,X13,X14] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ( ~ ssItem(X12)
        | ~ ssList(X13)
        | app(X13,cons(X12,nil)) = esk50_0
        | app(cons(X12,nil),X13) != esk51_0 )
      & ( nil = esk50_0
        | nil != esk51_0 )
      & ( neq(esk49_0,nil)
        | nil = esk49_0 )
      & ( ~ ssList(X14)
        | ~ neq(X14,nil)
        | ~ segmentP(esk49_0,X14)
        | ~ segmentP(esk48_0,X14)
        | nil = esk49_0 )
      & ( neq(esk49_0,nil)
        | nil != esk48_0 )
      & ( ~ ssList(X14)
        | ~ neq(X14,nil)
        | ~ segmentP(esk49_0,X14)
        | ~ segmentP(esk48_0,X14)
        | nil != esk48_0 ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])])])).

<A NAME="c_0_37"></A>cnf(c_0_37,plain,
    ( app(nil,X1) = nil
    | app(X2,X1) != nil
    | ~ ssList(app(nil,X1))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_39"></A>fof(c_0_39,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( nil != esk48_0
    | ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( neq(esk49_0,nil)
    | nil != esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( nil = esk50_0
    | nil != esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,plain,
    ( app(nil,X1) = nil
    | app(X2,X1) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_38]),c_0_31])])).

<A NAME="c_0_47"></A>cnf(c_0_47,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( esk48_0 != nil
    | ~ segmentP(esk48_0,esk49_0)
    | ~ segmentP(esk49_0,esk49_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_41]),c_0_42])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( esk48_0 = nil
    | esk49_0 != nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_44]),<A HREF="#c_0_45">c_0_45</A>])).

<A NAME="c_0_50"></A>fof(c_0_50,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

<A NAME="c_0_51"></A>fof(c_0_51,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

<A NAME="c_0_52"></A>fof(c_0_52,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssItem(X6)
      | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( app(nil,X1) = nil
    | X1 != nil
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_47]),c_0_31])])).

<A NAME="c_0_54"></A>fof(c_0_54,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | hd(cons(X4,X3)) = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).

<A NAME="c_0_55"></A>fof(c_0_55,plain,(
    ! [X4] :
      ( ( ssList(esk44_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( ssItem(esk45_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( cons(esk45_1(X4),esk44_1(X4)) = X4
        | nil = X4
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( esk49_0 != nil
    | ~ segmentP(nil,esk49_0)
    | ~ segmentP(esk49_0,esk49_0) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,plain,
    ( segmentP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_58"></A>fof(c_0_58,plain,(
    ! [X2] :
      ( ( ~ segmentP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | segmentP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).

<A NAME="c_0_59"></A>fof(c_0_59,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).

<A NAME="c_0_60"></A>fof(c_0_60,plain,(
    ! [X3] :
      ( ( ssItem(esk46_1(X3))
        | nil = X3
        | ~ ssList(X3) )
      & ( hd(X3) = esk46_1(X3)
        | nil = X3
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( app(X2,cons(X1,nil)) = esk50_0
    | app(cons(X1,nil),X2) != esk51_0
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_62"></A>fof(c_0_62,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) != X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).

<A NAME="c_0_63"></A>cnf(c_0_63,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>])).

<A NAME="c_0_64"></A>cnf(c_0_64,plain,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( app(nil,app(X1,X2)) = app(nil,X2)
    | X1 != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_53]),c_0_31])])).

<A NAME="c_0_66"></A>cnf(c_0_66,plain,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_67"></A>cnf(c_0_67,plain,
    ( nil = X1
    | cons(esk45_1(X1),esk44_1(X1)) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_68"></A>cnf(c_0_68,plain,
    ( nil = X1
    | ssItem(esk45_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( nil = X1
    | ssList(esk44_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( esk49_0 != nil
    | ~ segmentP(nil,esk49_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_57]),c_0_42])])).

<A NAME="c_0_71"></A>cnf(c_0_71,plain,
    ( segmentP(nil,X1)
    | ~ ssList(X1)
    | nil != X1 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>])).

<A NAME="c_0_73"></A>fof(c_0_73,plain,(
    ! [X3] :
      ( ( ssList(esk47_1(X3))
        | nil = X3
        | ~ ssList(X3) )
      & ( tl(X3) = esk47_1(X3)
        | nil = X3
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])).

<A NAME="c_0_74"></A>cnf(c_0_74,plain,
    ( nil = X1
    | hd(X1) = esk46_1(X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( app(X1,cons(X2,nil)) = esk48_0
    | app(cons(X2,nil),X1) != esk49_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_44]),<A HREF="#c_0_45">c_0_45</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,plain,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,plain,
    ( cons(X1,app(nil,X2)) = cons(X1,X2)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>,c_0_64]),c_0_31])])).

<A NAME="c_0_78"></A>fof(c_0_78,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk8_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ssList(esk9_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(app(X9,X6),X10) != X5
        | segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])])).

<A NAME="c_0_79"></A>cnf(c_0_79,plain,
    ( app(nil,app(X1,X2)) = app(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_47]),c_0_31])])).

<A NAME="c_0_80"></A>cnf(c_0_80,plain,
    ( nil = app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != X1
    | nil != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( app(X1,app(X2,nil)) = app(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_29]),c_0_31])]),<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,plain,
    ( ssList(app(nil,X1))
    | X2 != nil
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_65]),c_0_31])]),<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,plain,
    ( esk45_1(X1) = hd(X1)
    | nil = X1
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>,c_0_67]),c_0_68]),<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( esk49_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>,c_0_71]),c_0_42])])).

<A NAME="c_0_85"></A>cnf(c_0_85,plain,
    ( esk44_1(X1) = tl(X1)
    | nil = X1
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_67]),c_0_68]),<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,plain,
    ( nil = X1
    | tl(X1) = esk47_1(X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( esk46_1(esk49_0) = hd(esk49_0)
    | esk49_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( app(X1,cons(X2,nil)) = esk48_0
    | cons(X2,X1) != esk49_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_89"></A>fof(c_0_89,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).

<A NAME="c_0_90"></A>cnf(c_0_90,plain,
    ( cons(X1,X2) != app(nil,X2)
    | ~ ssList(app(nil,X2))
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>,<A HREF="#c_0_77">c_0_77</A>])).

<A NAME="c_0_91"></A>cnf(c_0_91,plain,
    ( segmentP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(app(X3,X2),X4) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>])).

<A NAME="c_0_92"></A>cnf(c_0_92,plain,
    ( app(nil,nil) = nil
    | nil != X1
    | nil != X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_79">c_0_79</A>,<A HREF="#c_0_80">c_0_80</A>])).

<A NAME="c_0_93"></A>cnf(c_0_93,plain,
    ( app(nil,X1) = app(X1,nil)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_79">c_0_79</A>,c_0_81]),c_0_31])])).

<A NAME="c_0_94"></A>cnf(c_0_94,plain,
    ( ssList(app(nil,X1))
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_95"></A>cnf(c_0_95,negated_conjecture,
    ( esk45_1(esk49_0) = hd(esk49_0) ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_42]),<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_96"></A>cnf(c_0_96,negated_conjecture,
    ( esk44_1(esk49_0) = tl(esk49_0) ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_42]),<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,plain,
    ( nil = X1
    | ssList(esk47_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>])).

<A NAME="c_0_98"></A>cnf(c_0_98,negated_conjecture,
    ( esk47_1(esk49_0) = tl(esk49_0) ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>,c_0_42]),<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( nil = X1
    | ssItem(esk46_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,negated_conjecture,
    ( esk46_1(esk49_0) = hd(esk49_0) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_101"></A>cnf(c_0_101,negated_conjecture,
    ( app(nil,esk48_0) = esk48_0
    | cons(X1,X2) != esk49_0
    | ~ ssList(cons(X1,nil))
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_79">c_0_79</A>,<A HREF="#c_0_88">c_0_88</A>])).

<A NAME="c_0_102"></A>cnf(c_0_102,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>])).

<A NAME="c_0_103"></A>cnf(c_0_103,plain,
    ( cons(X1,nil) != nil
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,c_0_30]),c_0_31])])).

<A NAME="c_0_104"></A>cnf(c_0_104,plain,
    ( segmentP(X1,X2)
    | app(X2,X3) != X1
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_47]),c_0_31])])).

<A NAME="c_0_105"></A>cnf(c_0_105,plain,
    ( app(nil,nil) = nil
    | nil != X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_106"></A>fof(c_0_106,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | cons(hd(X2),tl(X2)) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).

<A NAME="c_0_107"></A>fof(c_0_107,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssItem(hd(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])).

<A NAME="c_0_108"></A>fof(c_0_108,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssList(tl(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).

<A NAME="c_0_109"></A>cnf(c_0_109,plain,
    ( app(X1,app(X2,X3)) = nil
    | app(X1,X2) != nil
    | nil != X3
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,c_0_29]),<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_110"></A>cnf(c_0_110,plain,
    ( app(nil,app(X1,nil)) = app(X1,nil)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_79">c_0_79</A>,c_0_93]),c_0_31])])).

<A NAME="c_0_111"></A>cnf(c_0_111,plain,
    ( tl(cons(X1,X2)) = app(nil,X2)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_77]),<A HREF="#c_0_94">c_0_94</A>])).

<A NAME="c_0_112"></A>cnf(c_0_112,negated_conjecture,
    ( cons(hd(esk49_0),tl(esk49_0)) = esk49_0 ),
    inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,c_0_95]),c_0_42])]),c_0_84]),<A HREF="#c_0_96">c_0_96</A>])).

<A NAME="c_0_113"></A>cnf(c_0_113,negated_conjecture,
    ( ssList(tl(esk49_0)) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_98]),c_0_42])]),<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_114"></A>cnf(c_0_114,negated_conjecture,
    ( ssItem(hd(esk49_0)) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_100]),c_0_42])]),<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_115"></A>fof(c_0_115,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( app(nil,esk48_0) = esk48_0
    | cons(X1,X2) != esk49_0
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_101">c_0_101</A>,c_0_102]),c_0_31])])).

<A NAME="c_0_117"></A>cnf(c_0_117,negated_conjecture,
    ( cons(X1,X2) != esk49_0
    | esk48_0 != nil
    | ~ ssList(cons(X1,nil))
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_88]),<A HREF="#c_0_103">c_0_103</A>])).

<A NAME="c_0_118"></A>cnf(c_0_118,plain,
    ( segmentP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_104]),<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_119"></A>fof(c_0_119,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ rearsegP(X4,X5)
      | rearsegP(app(X6,X4),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).

<A NAME="c_0_120"></A>cnf(c_0_120,plain,
    ( ssList(cons(X1,app(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_64]),<A HREF="#c_0_102">c_0_102</A>])).

<A NAME="c_0_121"></A>cnf(c_0_121,plain,
    ( app(nil,nil) = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_105">c_0_105</A>,<A HREF="#c_0_31">c_0_31</A>])).

<A NAME="c_0_122"></A>cnf(c_0_122,negated_conjecture,
    ( app(nil,cons(X1,nil)) = esk48_0
    | cons(X1,nil) != esk49_0
    | ~ ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_30]),c_0_31])])).

<A NAME="c_0_123"></A>cnf(c_0_123,plain,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>])).

<A NAME="c_0_124"></A>cnf(c_0_124,plain,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>])).

<A NAME="c_0_125"></A>cnf(c_0_125,plain,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>])).

<A NAME="c_0_126"></A>cnf(c_0_126,plain,
    ( app(X1,nil) = nil
    | app(nil,X1) != nil
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>,c_0_110]),c_0_31])])).

<A NAME="c_0_127"></A>cnf(c_0_127,negated_conjecture,
    ( app(nil,tl(esk49_0)) = tl(esk49_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_111">c_0_111</A>,c_0_112]),c_0_113]),c_0_114])])).

<A NAME="c_0_128"></A>cnf(c_0_128,negated_conjecture,
    ( nil = esk49_0
    | ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_129"></A>cnf(c_0_129,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_115">c_0_115</A>])).

<A NAME="c_0_130"></A>cnf(c_0_130,negated_conjecture,
    ( app(nil,esk48_0) = esk48_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_116">c_0_116</A>,c_0_112]),c_0_113]),c_0_114])])).

<A NAME="c_0_131"></A>cnf(c_0_131,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_132"></A>cnf(c_0_132,negated_conjecture,
    ( cons(X1,X2) != esk49_0
    | esk48_0 != nil
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_117">c_0_117</A>,c_0_102]),c_0_31])])).

<A NAME="c_0_133"></A>cnf(c_0_133,negated_conjecture,
    ( segmentP(esk48_0,X1)
    | cons(X2,X1) != esk49_0
    | ~ ssList(cons(X2,nil))
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_118">c_0_118</A>,<A HREF="#c_0_88">c_0_88</A>])).

<A NAME="c_0_134"></A>fof(c_0_134,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk7_2(X4,X5))
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(esk7_2(X4,X5),X5) = X4
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X7,X5) != X4
        | rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])])])).

<A NAME="c_0_135"></A>cnf(c_0_135,plain,
    ( rearsegP(app(X1,X2),X3)
    | ~ rearsegP(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>])).

<A NAME="c_0_136"></A>cnf(c_0_136,plain,
    ( ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_120">c_0_120</A>,c_0_121]),c_0_31])])).

<A NAME="c_0_137"></A>cnf(c_0_137,negated_conjecture,
    ( cons(X1,nil) = esk48_0
    | cons(X1,nil) != esk49_0
    | ~ ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,<A HREF="#c_0_122">c_0_122</A>])).

<A NAME="c_0_138"></A>cnf(c_0_138,plain,
    ( cons(hd(X1),app(tl(X1),X2)) = app(X1,X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,c_0_123]),c_0_124]),<A HREF="#c_0_125">c_0_125</A>])).

<A NAME="c_0_139"></A>cnf(c_0_139,negated_conjecture,
    ( app(tl(esk49_0),nil) = nil
    | tl(esk49_0) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_126">c_0_126</A>,c_0_127]),c_0_113])])).

<A NAME="c_0_140"></A>cnf(c_0_140,negated_conjecture,
    ( X1 = nil
    | ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ ssList(X1) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_128">c_0_128</A>,c_0_129]),c_0_31])]),<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_141"></A>cnf(c_0_141,negated_conjecture,
    ( segmentP(X1,esk48_0)
    | app(esk48_0,X2) != X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_130]),c_0_31]),c_0_131])])).

<A NAME="c_0_142"></A>cnf(c_0_142,negated_conjecture,
    ( nil = X1
    | esk48_0 != nil
    | X1 != esk49_0
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_132">c_0_132</A>,c_0_67]),c_0_68]),<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_143"></A>cnf(c_0_143,negated_conjecture,
    ( segmentP(esk48_0,X1)
    | cons(X2,X1) != esk49_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_133">c_0_133</A>,c_0_102]),c_0_31])])).

<A NAME="c_0_144"></A>cnf(c_0_144,plain,
    ( app(esk7_2(X1,X2),X2) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_134">c_0_134</A>])).

<A NAME="c_0_145"></A>cnf(c_0_145,plain,
    ( ssList(esk7_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_134">c_0_134</A>])).

<A NAME="c_0_146"></A>cnf(c_0_146,plain,
    ( rearsegP(cons(X1,X2),X3)
    | ~ rearsegP(X2,X3)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_135">c_0_135</A>,c_0_63]),<A HREF="#c_0_136">c_0_136</A>])).

<A NAME="c_0_147"></A>fof(c_0_147,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | rearsegP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])).

<A NAME="c_0_148"></A>cnf(c_0_148,negated_conjecture,
    ( cons(X1,nil) = esk48_0
    | cons(X1,nil) != esk49_0
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_137">c_0_137</A>,c_0_102]),c_0_31])])).

<A NAME="c_0_149"></A>cnf(c_0_149,negated_conjecture,
    ( cons(hd(esk49_0),nil) = app(esk49_0,nil)
    | tl(esk49_0) != nil ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_138">c_0_138</A>,c_0_139]),c_0_31]),c_0_42])]),<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_150"></A>cnf(c_0_150,negated_conjecture,
    ( esk48_0 = nil
    | ~ segmentP(esk49_0,esk48_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_140">c_0_140</A>,c_0_57]),c_0_131])])).

<A NAME="c_0_151"></A>cnf(c_0_151,negated_conjecture,
    ( segmentP(X1,esk48_0)
    | esk48_0 != X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_141">c_0_141</A>,c_0_30]),c_0_31]),c_0_131])])).

<A NAME="c_0_152"></A>cnf(c_0_152,negated_conjecture,
    ( esk48_0 != nil ),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_142">c_0_142</A>,c_0_42]),<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_153"></A>cnf(c_0_153,negated_conjecture,
    ( segmentP(esk48_0,tl(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_143">c_0_143</A>,c_0_112]),c_0_113]),c_0_114])])).

<A NAME="c_0_154"></A>cnf(c_0_154,plain,
    ( segmentP(X1,X2)
    | app(X3,X4) != X1
    | ~ rearsegP(X3,X2)
    | ~ ssList(X4)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_144]),<A HREF="#c_0_145">c_0_145</A>])).

<A NAME="c_0_155"></A>cnf(c_0_155,negated_conjecture,
    ( rearsegP(esk49_0,X1)
    | ~ rearsegP(tl(esk49_0),X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_146">c_0_146</A>,c_0_112]),c_0_113]),c_0_114])])).

<A NAME="c_0_156"></A>cnf(c_0_156,plain,
    ( rearsegP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_147">c_0_147</A>])).

<A NAME="c_0_157"></A>cnf(c_0_157,negated_conjecture,
    ( app(esk49_0,nil) = esk48_0
    | app(esk49_0,nil) != esk49_0
    | tl(esk49_0) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_148">c_0_148</A>,c_0_149]),c_0_114])])).

<A NAME="c_0_158"></A>cnf(c_0_158,negated_conjecture,
    ( esk48_0 != esk49_0 ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_150">c_0_150</A>,c_0_151]),c_0_42])]),<A HREF="#c_0_152">c_0_152</A>])).

<A NAME="c_0_159"></A>cnf(c_0_159,negated_conjecture,
    ( tl(esk49_0) = nil
    | ~ segmentP(esk49_0,tl(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_140">c_0_140</A>,c_0_153]),c_0_113])])).

<A NAME="c_0_160"></A>cnf(c_0_160,plain,
    ( segmentP(X1,X2)
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_154">c_0_154</A>,c_0_30]),c_0_31])])])).

<A NAME="c_0_161"></A>cnf(c_0_161,negated_conjecture,
    ( rearsegP(esk49_0,tl(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_155">c_0_155</A>,c_0_156]),c_0_113])])).

<A NAME="c_0_162"></A>cnf(c_0_162,negated_conjecture,
    ( tl(esk49_0) != nil ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_157">c_0_157</A>,c_0_30]),c_0_42])]),<A HREF="#c_0_158">c_0_158</A>])).

<A NAME="c_0_163"></A>cnf(c_0_163,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_159">c_0_159</A>,c_0_160]),c_0_161]),c_0_113]),c_0_42])]),<A HREF="#c_0_162">c_0_162</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC062+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n169.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:34:39 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 151.14/151.30  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 151.14/151.30  # and selection function SelectCQPrecWNTNp.
% 151.14/151.30  #
% 151.14/151.30  # Preprocessing time       : 0.014 s
% 151.14/151.30  # Presaturation interreduction done
% 151.45/151.65  # No success with AutoSched0
% 151.45/151.65  # Trying AutoSched1 for 73 seconds
% 224.48/224.65  # AutoSched1-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S00
% 224.48/224.65  # and selection function NoSelection.
% 224.48/224.65  #
% 224.48/224.65  # Preprocessing time       : 0.015 s
% 224.48/224.65  # Presaturation interreduction done
% 224.69/224.85  # No success with AutoSched1
% 224.69/224.85  # Trying AutoSched2 for 23 seconds
% 225.39/225.59  # AutoSched2-Mode selected heuristic G_E___107_C37_F1_PI_AE_Q4_CS_SP_PS_S0Y
% 225.39/225.59  # and selection function SelectMaxLComplexAvoidPosPred.
% 225.39/225.59  #
% 225.39/225.59  # Preprocessing time       : 0.015 s
% 225.39/225.59  # Presaturation interreduction done
% 225.39/225.59  
% 225.39/225.59  # Proof found!
% 225.39/225.59  # SZS status Theorem
% 225.39/225.59  # SZS output start CNFRefutation
% 225.39/225.59  fof(c_0_0, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax82)).
% 225.39/225.59  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax84)).
% 225.39/225.59  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax83)).
% 225.39/225.59  fof(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
% 225.39/225.59  fof(c_0_4, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|?[X5]:(ssItem(X5)&?[X6]:((ssList(X6)&app(X6,cons(X5,nil))!=X3)&app(cons(X5,nil),X6)=X4)))|(nil!=X3&nil=X4))|((nil!=X2|nil=X1)&(~(neq(X2,nil))|?[X7]:(((ssList(X7)&neq(X7,nil))&segmentP(X2,X7))&segmentP(X1,X7)))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
% 225.39/225.59  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax26)).
% 225.39/225.59  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
% 225.39/225.59  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>segmentP(X1,X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax55)).
% 225.39/225.59  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax81)).
% 225.39/225.59  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax27)).
% 225.39/225.59  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax23)).
% 225.39/225.59  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax20)).
% 225.39/225.59  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>(segmentP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax58)).
% 225.39/225.59  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax25)).
% 225.39/225.59  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>?[X2]:(ssItem(X2)&hd(X1)=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax75)).
% 225.39/225.59  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax18)).
% 225.39/225.59  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>?[X2]:(ssList(X2)&tl(X1)=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax76)).
% 225.39/225.59  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax7)).
% 225.39/225.59  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
% 225.39/225.59  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax78)).
% 225.39/225.59  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax22)).
% 225.39/225.59  fof(c_0_21, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax24)).
% 225.39/225.59  fof(c_0_22, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax15)).
% 225.39/225.59  fof(c_0_23, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(rearsegP(X1,X2)=>rearsegP(app(X3,X1),X2)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax50)).
% 225.39/225.59  fof(c_0_24, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax6)).
% 225.39/225.59  fof(c_0_25, axiom, (![X1]:(ssList(X1)=>rearsegP(X1,X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax49)).
% 225.39/225.59  fof(c_0_26, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])).
% 225.39/225.59  fof(c_0_27, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
% 225.39/225.59  fof(c_0_28, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
% 225.39/225.59  cnf(c_0_29, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
% 225.39/225.59  cnf(c_0_30, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
% 225.39/225.59  cnf(c_0_31, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 225.39/225.59  fof(c_0_32, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|?[X5]:(ssItem(X5)&?[X6]:((ssList(X6)&app(X6,cons(X5,nil))!=X3)&app(cons(X5,nil),X6)=X4)))|(nil!=X3&nil=X4))|((nil!=X2|nil=X1)&(~neq(X2,nil)|?[X7]:(((ssList(X7)&neq(X7,nil))&segmentP(X2,X7))&segmentP(X1,X7))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_4])])).
% 225.39/225.59  cnf(c_0_33, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
% 225.39/225.59  cnf(c_0_34, plain, (app(X1,app(nil,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
% 225.39/225.59  fof(c_0_35, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 225.39/225.59  fof(c_0_36, negated_conjecture, (![X12]:![X13]:![X14]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(~ssItem(X12)|((~ssList(X13)|app(X13,cons(X12,nil))=esk50_0)|app(cons(X12,nil),X13)!=esk51_0)))&(nil=esk50_0|nil!=esk51_0))&(((neq(esk49_0,nil)|nil=esk49_0)&((((~ssList(X14)|~neq(X14,nil))|~segmentP(esk49_0,X14))|~segmentP(esk48_0,X14))|nil=esk49_0))&((neq(esk49_0,nil)|nil!=esk48_0)&((((~ssList(X14)|~neq(X14,nil))|~segmentP(esk49_0,X14))|~segmentP(esk48_0,X14))|nil!=esk48_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])])])).
% 225.39/225.59  cnf(c_0_37, plain, (app(nil,X1)=nil|app(X2,X1)!=nil|~ssList(app(nil,X1))|~ssList(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
% 225.39/225.59  cnf(c_0_38, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
% 225.39/225.59  fof(c_0_39, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
% 225.39/225.59  cnf(c_0_40, negated_conjecture, (nil!=esk48_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
% 225.39/225.59  cnf(c_0_41, negated_conjecture, (neq(esk49_0,nil)|nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_36])).
% 225.39/225.59  cnf(c_0_42, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
% 225.39/225.59  cnf(c_0_43, negated_conjecture, (nil=esk50_0|nil!=esk51_0), inference(split_conjunct,[status(thm)],[c_0_36])).
% 225.39/225.59  cnf(c_0_44, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_36])).
% 225.39/225.59  cnf(c_0_45, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_36])).
% 225.39/225.59  cnf(c_0_46, plain, (app(nil,X1)=nil|app(X2,X1)!=nil|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_31])])).
% 225.39/225.59  cnf(c_0_47, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 225.39/225.59  cnf(c_0_48, negated_conjecture, (esk48_0!=nil|~segmentP(esk48_0,esk49_0)|~segmentP(esk49_0,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])])).
% 225.39/225.59  cnf(c_0_49, negated_conjecture, (esk48_0=nil|esk49_0!=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
% 225.39/225.59  fof(c_0_50, plain, (![X2]:(~ssList(X2)|segmentP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
% 225.39/225.59  fof(c_0_51, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
% 225.39/225.59  fof(c_0_52, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 225.39/225.59  cnf(c_0_53, plain, (app(nil,X1)=nil|X1!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_31])])).
% 225.39/225.59  fof(c_0_54, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
% 225.39/225.59  fof(c_0_55, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).
% 225.39/225.59  cnf(c_0_56, negated_conjecture, (esk49_0!=nil|~segmentP(nil,esk49_0)|~segmentP(esk49_0,esk49_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
% 225.39/225.59  cnf(c_0_57, plain, (segmentP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
% 225.39/225.59  fof(c_0_58, plain, (![X2]:(((~segmentP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|segmentP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
% 225.39/225.59  fof(c_0_59, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
% 225.39/225.59  fof(c_0_60, plain, (![X3]:(((ssItem(esk46_1(X3))|nil=X3)|~ssList(X3))&((hd(X3)=esk46_1(X3)|nil=X3)|~ssList(X3)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])).
% 225.39/225.59  cnf(c_0_61, negated_conjecture, (app(X2,cons(X1,nil))=esk50_0|app(cons(X1,nil),X2)!=esk51_0|~ssList(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
% 225.39/225.59  fof(c_0_62, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
% 225.39/225.59  cnf(c_0_63, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
% 225.39/225.59  cnf(c_0_64, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
% 225.39/225.59  cnf(c_0_65, plain, (app(nil,app(X1,X2))=app(nil,X2)|X1!=nil|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_53]), c_0_31])])).
% 225.39/225.59  cnf(c_0_66, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
% 225.39/225.59  cnf(c_0_67, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
% 225.39/225.59  cnf(c_0_68, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
% 225.39/225.59  cnf(c_0_69, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
% 225.39/225.59  cnf(c_0_70, negated_conjecture, (esk49_0!=nil|~segmentP(nil,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_42])])).
% 225.39/225.59  cnf(c_0_71, plain, (segmentP(nil,X1)|~ssList(X1)|nil!=X1), inference(split_conjunct,[status(thm)],[c_0_58])).
% 225.39/225.59  cnf(c_0_72, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
% 225.39/225.59  fof(c_0_73, plain, (![X3]:(((ssList(esk47_1(X3))|nil=X3)|~ssList(X3))&((tl(X3)=esk47_1(X3)|nil=X3)|~ssList(X3)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])).
% 225.39/225.59  cnf(c_0_74, plain, (nil=X1|hd(X1)=esk46_1(X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
% 225.39/225.59  cnf(c_0_75, negated_conjecture, (app(X1,cons(X2,nil))=esk48_0|app(cons(X2,nil),X1)!=esk49_0|~ssList(X1)|~ssItem(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_44]), c_0_45])).
% 225.39/225.59  cnf(c_0_76, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
% 225.39/225.59  cnf(c_0_77, plain, (cons(X1,app(nil,X2))=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_31])])).
% 225.39/225.59  fof(c_0_78, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])])).
% 225.39/225.59  cnf(c_0_79, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_47]), c_0_31])])).
% 225.39/225.59  cnf(c_0_80, plain, (nil=app(X1,X2)|~ssList(X1)|~ssList(X2)|nil!=X1|nil!=X2), inference(split_conjunct,[status(thm)],[c_0_28])).
% 225.39/225.59  cnf(c_0_81, plain, (app(X1,app(X2,nil))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_29]), c_0_31])]), c_0_38])).
% 225.39/225.59  cnf(c_0_82, plain, (ssList(app(nil,X1))|X2!=nil|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_65]), c_0_31])]), c_0_38])).
% 225.39/225.59  cnf(c_0_83, plain, (esk45_1(X1)=hd(X1)|nil=X1|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68]), c_0_69])).
% 225.39/225.59  cnf(c_0_84, negated_conjecture, (esk49_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_42])])).
% 225.39/225.59  cnf(c_0_85, plain, (esk44_1(X1)=tl(X1)|nil=X1|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_67]), c_0_68]), c_0_69])).
% 225.39/225.59  cnf(c_0_86, plain, (nil=X1|tl(X1)=esk47_1(X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
% 225.39/225.59  cnf(c_0_87, negated_conjecture, (esk46_1(esk49_0)=hd(esk49_0)|esk49_0=nil), inference(spm,[status(thm)],[c_0_74, c_0_42])).
% 225.39/225.59  cnf(c_0_88, negated_conjecture, (app(X1,cons(X2,nil))=esk48_0|cons(X2,X1)!=esk49_0|~ssList(X1)|~ssItem(X2)), inference(spm,[status(thm)],[c_0_75, c_0_63])).
% 225.39/225.59  fof(c_0_89, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
% 225.39/225.59  cnf(c_0_90, plain, (cons(X1,X2)!=app(nil,X2)|~ssList(app(nil,X2))|~ssList(X2)|~ssItem(X1)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
% 225.39/225.59  cnf(c_0_91, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_78])).
% 225.39/225.59  cnf(c_0_92, plain, (app(nil,nil)=nil|nil!=X1|nil!=X2|~ssList(X1)|~ssList(X2)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
% 225.39/225.59  cnf(c_0_93, plain, (app(nil,X1)=app(X1,nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_81]), c_0_31])])).
% 225.39/225.59  cnf(c_0_94, plain, (ssList(app(nil,X1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_82, c_0_31])).
% 225.39/225.59  cnf(c_0_95, negated_conjecture, (esk45_1(esk49_0)=hd(esk49_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_42]), c_0_84])).
% 225.39/225.59  cnf(c_0_96, negated_conjecture, (esk44_1(esk49_0)=tl(esk49_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_42]), c_0_84])).
% 225.39/225.59  cnf(c_0_97, plain, (nil=X1|ssList(esk47_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
% 225.39/225.59  cnf(c_0_98, negated_conjecture, (esk47_1(esk49_0)=tl(esk49_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_42]), c_0_84])).
% 225.39/225.59  cnf(c_0_99, plain, (nil=X1|ssItem(esk46_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
% 225.39/225.59  cnf(c_0_100, negated_conjecture, (esk46_1(esk49_0)=hd(esk49_0)), inference(sr,[status(thm)],[c_0_87, c_0_84])).
% 225.39/225.59  cnf(c_0_101, negated_conjecture, (app(nil,esk48_0)=esk48_0|cons(X1,X2)!=esk49_0|~ssList(cons(X1,nil))|~ssList(X2)|~ssItem(X1)), inference(spm,[status(thm)],[c_0_79, c_0_88])).
% 225.39/225.59  cnf(c_0_102, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
% 225.39/225.59  cnf(c_0_103, plain, (cons(X1,nil)!=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_30]), c_0_31])])).
% 225.39/225.59  cnf(c_0_104, plain, (segmentP(X1,X2)|app(X2,X3)!=X1|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_47]), c_0_31])])).
% 225.39/225.59  cnf(c_0_105, plain, (app(nil,nil)=nil|nil!=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_92, c_0_31])).
% 225.39/225.59  fof(c_0_106, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).
% 225.39/225.59  fof(c_0_107, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])).
% 225.39/225.59  fof(c_0_108, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).
% 225.39/225.59  cnf(c_0_109, plain, (app(X1,app(X2,X3))=nil|app(X1,X2)!=nil|nil!=X3|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_29]), c_0_38])).
% 225.39/225.59  cnf(c_0_110, plain, (app(nil,app(X1,nil))=app(X1,nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_93]), c_0_31])])).
% 225.39/225.59  cnf(c_0_111, plain, (tl(cons(X1,X2))=app(nil,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_77]), c_0_94])).
% 225.39/225.59  cnf(c_0_112, negated_conjecture, (cons(hd(esk49_0),tl(esk49_0))=esk49_0), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_95]), c_0_42])]), c_0_84]), c_0_96])).
% 225.39/225.59  cnf(c_0_113, negated_conjecture, (ssList(tl(esk49_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_42])]), c_0_84])).
% 225.39/225.59  cnf(c_0_114, negated_conjecture, (ssItem(hd(esk49_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_42])]), c_0_84])).
% 225.39/225.59  fof(c_0_115, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).
% 225.39/225.59  cnf(c_0_116, negated_conjecture, (app(nil,esk48_0)=esk48_0|cons(X1,X2)!=esk49_0|~ssList(X2)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_31])])).
% 225.39/225.59  cnf(c_0_117, negated_conjecture, (cons(X1,X2)!=esk49_0|esk48_0!=nil|~ssList(cons(X1,nil))|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_88]), c_0_103])).
% 225.39/225.59  cnf(c_0_118, plain, (segmentP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_104]), c_0_38])).
% 225.39/225.59  fof(c_0_119, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~rearsegP(X4,X5)|rearsegP(app(X6,X4),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).
% 225.39/225.59  cnf(c_0_120, plain, (ssList(cons(X1,app(X2,X3)))|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_64]), c_0_102])).
% 225.39/225.59  cnf(c_0_121, plain, (app(nil,nil)=nil), inference(spm,[status(thm)],[c_0_105, c_0_31])).
% 225.39/225.59  cnf(c_0_122, negated_conjecture, (app(nil,cons(X1,nil))=esk48_0|cons(X1,nil)!=esk49_0|~ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_30]), c_0_31])])).
% 225.39/225.59  cnf(c_0_123, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
% 225.39/225.59  cnf(c_0_124, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_107])).
% 225.39/225.59  cnf(c_0_125, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
% 225.39/225.59  cnf(c_0_126, plain, (app(X1,nil)=nil|app(nil,X1)!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_31])])).
% 225.39/225.59  cnf(c_0_127, negated_conjecture, (app(nil,tl(esk49_0))=tl(esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_113]), c_0_114])])).
% 225.39/225.59  cnf(c_0_128, negated_conjecture, (nil=esk49_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
% 225.39/225.59  cnf(c_0_129, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_115])).
% 225.39/225.59  cnf(c_0_130, negated_conjecture, (app(nil,esk48_0)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_112]), c_0_113]), c_0_114])])).
% 225.39/225.59  cnf(c_0_131, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
% 225.39/225.59  cnf(c_0_132, negated_conjecture, (cons(X1,X2)!=esk49_0|esk48_0!=nil|~ssList(X2)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_102]), c_0_31])])).
% 225.39/225.59  cnf(c_0_133, negated_conjecture, (segmentP(esk48_0,X1)|cons(X2,X1)!=esk49_0|~ssList(cons(X2,nil))|~ssList(X1)|~ssItem(X2)), inference(spm,[status(thm)],[c_0_118, c_0_88])).
% 225.39/225.59  fof(c_0_134, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])])])).
% 225.39/225.59  cnf(c_0_135, plain, (rearsegP(app(X1,X2),X3)|~rearsegP(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_119])).
% 225.39/225.59  cnf(c_0_136, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_31])])).
% 225.39/225.59  cnf(c_0_137, negated_conjecture, (cons(X1,nil)=esk48_0|cons(X1,nil)!=esk49_0|~ssList(cons(X1,nil))|~ssItem(X1)), inference(spm,[status(thm)],[c_0_47, c_0_122])).
% 225.39/225.59  cnf(c_0_138, plain, (cons(hd(X1),app(tl(X1),X2))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_123]), c_0_124]), c_0_125])).
% 225.39/225.59  cnf(c_0_139, negated_conjecture, (app(tl(esk49_0),nil)=nil|tl(esk49_0)!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_113])])).
% 225.39/225.59  cnf(c_0_140, negated_conjecture, (X1=nil|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~ssList(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_31])]), c_0_84])).
% 225.39/225.59  cnf(c_0_141, negated_conjecture, (segmentP(X1,esk48_0)|app(esk48_0,X2)!=X1|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_130]), c_0_31]), c_0_131])])).
% 225.39/225.59  cnf(c_0_142, negated_conjecture, (nil=X1|esk48_0!=nil|X1!=esk49_0|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_67]), c_0_68]), c_0_69])).
% 225.39/225.59  cnf(c_0_143, negated_conjecture, (segmentP(esk48_0,X1)|cons(X2,X1)!=esk49_0|~ssList(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_102]), c_0_31])])).
% 225.39/225.59  cnf(c_0_144, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_134])).
% 225.39/225.59  cnf(c_0_145, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_134])).
% 225.39/225.59  cnf(c_0_146, plain, (rearsegP(cons(X1,X2),X3)|~rearsegP(X2,X3)|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_63]), c_0_136])).
% 225.39/225.59  fof(c_0_147, plain, (![X2]:(~ssList(X2)|rearsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])).
% 225.39/225.59  cnf(c_0_148, negated_conjecture, (cons(X1,nil)=esk48_0|cons(X1,nil)!=esk49_0|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_102]), c_0_31])])).
% 225.39/225.59  cnf(c_0_149, negated_conjecture, (cons(hd(esk49_0),nil)=app(esk49_0,nil)|tl(esk49_0)!=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_31]), c_0_42])]), c_0_84])).
% 225.39/225.59  cnf(c_0_150, negated_conjecture, (esk48_0=nil|~segmentP(esk49_0,esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_57]), c_0_131])])).
% 225.39/225.59  cnf(c_0_151, negated_conjecture, (segmentP(X1,esk48_0)|esk48_0!=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_30]), c_0_31]), c_0_131])])).
% 225.39/225.59  cnf(c_0_152, negated_conjecture, (esk48_0!=nil), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_42]), c_0_84])).
% 225.39/225.59  cnf(c_0_153, negated_conjecture, (segmentP(esk48_0,tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_112]), c_0_113]), c_0_114])])).
% 225.39/225.59  cnf(c_0_154, plain, (segmentP(X1,X2)|app(X3,X4)!=X1|~rearsegP(X3,X2)|~ssList(X4)|~ssList(X2)|~ssList(X1)|~ssList(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_144]), c_0_145])).
% 225.39/225.59  cnf(c_0_155, negated_conjecture, (rearsegP(esk49_0,X1)|~rearsegP(tl(esk49_0),X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_112]), c_0_113]), c_0_114])])).
% 225.39/225.59  cnf(c_0_156, plain, (rearsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_147])).
% 225.39/225.59  cnf(c_0_157, negated_conjecture, (app(esk49_0,nil)=esk48_0|app(esk49_0,nil)!=esk49_0|tl(esk49_0)!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_114])])).
% 225.39/225.59  cnf(c_0_158, negated_conjecture, (esk48_0!=esk49_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_42])]), c_0_152])).
% 225.39/225.59  cnf(c_0_159, negated_conjecture, (tl(esk49_0)=nil|~segmentP(esk49_0,tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_153]), c_0_113])])).
% 225.39/225.59  cnf(c_0_160, plain, (segmentP(X1,X2)|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_30]), c_0_31])])])).
% 225.39/225.59  cnf(c_0_161, negated_conjecture, (rearsegP(esk49_0,tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156]), c_0_113])])).
% 225.39/225.59  cnf(c_0_162, negated_conjecture, (tl(esk49_0)!=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_30]), c_0_42])]), c_0_158])).
% 225.39/225.59  cnf(c_0_163, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_161]), c_0_113]), c_0_42])]), c_0_162]), ['proof']).
% 225.39/225.59  # SZS output end CNFRefutation
% 225.39/225.59  # Training examples: 0 positive, 0 negative
% 225.39/225.59  
% 225.39/225.59  # -------------------------------------------------
% 225.39/225.59  # User time                : 220.146 s
% 225.39/225.59  # System time              : 5.300 s
% 225.39/225.59  # Total time               : 225.447 s
% 225.39/225.59  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
