<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC420+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC420+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC420+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC420%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC420+1.p">SWC420+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n098.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:53 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :   60 ( 195 expanded)
%              Number of clauses        :   40 (  80 expanded)
%              Number of leaves         :   10 (  49 expanded)
%              Depth                    :   21
%              Number of atoms          :  240 (1668 expanded)
%              Number of equality atoms :  101 ( 721 expanded)
%              Maximal formula depth    :   21 (   5 average)
%              Maximal clause size      :   18 (   3 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ~ ssList(X4)
                  | X2 != X4
                  | X1 != X3
                  | ~ neq(X2,nil)
                  | ? [X5] :
                      ( ssItem(X5)
                      & ? [X6] :
                          ( ssList(X6)
                          & ? [X7] :
                              ( ssList(X7)
                              & app(app(X6,cons(X5,nil)),X7) = X2
                              & app(app(X7,cons(X5,nil)),X6) = X1 ) ) )
                  | ? [X8] :
                      ( ssItem(X8)
                      & ? [X9] :
                          ( ssList(X9)
                          & app(X9,cons(X8,nil)) != X3
                          & app(cons(X8,nil),X9) = X4 ) )
                  | ( nil != X3
                    & nil = X4 ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil = X1
        | ? [X2] :
            ( ssList(X2)
            & ? [X3] :
                ( ssItem(X3)
                & cons(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax20)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_10"></A>fof(c_0_10,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ~ ssList(X4)
                    | X2 != X4
                    | X1 != X3
                    | ~ neq(X2,nil)
                    | ? [X5] :
                        ( ssItem(X5)
                        & ? [X6] :
                            ( ssList(X6)
                            & ? [X7] :
                                ( ssList(X7)
                                & app(app(X6,cons(X5,nil)),X7) = X2
                                & app(app(X7,cons(X5,nil)),X6) = X1 ) ) )
                    | ? [X8] :
                        ( ssItem(X8)
                        & ? [X9] :
                            ( ssList(X9)
                            & app(X9,cons(X8,nil)) != X3
                            & app(cons(X8,nil),X9) = X4 ) )
                    | ( nil != X3
                      & nil = X4 ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_11"></A>fof(c_0_11,negated_conjecture,(
    ! [X14,X15,X16,X17,X18] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & neq(esk49_0,nil)
      & ( ~ ssItem(X14)
        | ~ ssList(X15)
        | ~ ssList(X16)
        | app(app(X15,cons(X14,nil)),X16) != esk49_0
        | app(app(X16,cons(X14,nil)),X15) != esk48_0 )
      & ( ~ ssItem(X17)
        | ~ ssList(X18)
        | app(X18,cons(X17,nil)) = esk50_0
        | app(cons(X17,nil),X18) != esk51_0 )
      & ( nil = esk50_0
        | nil != esk51_0 ) ) ),
    inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( app(app(X1,cons(X2,nil)),X3) != esk48_0
    | app(app(X3,cons(X2,nil)),X1) != esk49_0
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssItem(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_15"></A>fof(c_0_15,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( app(app(X1,cons(X2,nil)),X3) != esk51_0
    | app(app(X3,cons(X2,nil)),X1) != esk50_0
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssItem(X2) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_13]),<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_19"></A>fof(c_0_19,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( app(app(X1,cons(X2,nil)),nil) != esk50_0
    | app(cons(X2,nil),X1) != esk51_0
    | ~ ssList(cons(X2,nil))
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,c_0_17]),c_0_18])])).

<A NAME="c_0_21"></A>cnf(c_0_21,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_22"></A>fof(c_0_22,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( app(app(X1,cons(X2,nil)),nil) != esk50_0
    | app(cons(X2,nil),X1) != esk51_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_21]),c_0_18])])).

<A NAME="c_0_24"></A>cnf(c_0_24,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_25"></A>fof(c_0_25,plain,(
    ! [X4] :
      ( ( ssList(esk44_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( ssItem(esk45_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( cons(esk45_1(X4),esk44_1(X4)) = X4
        | nil = X4
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( app(app(X1,cons(X2,nil)),nil) != esk50_0
    | cons(X2,X1) != esk51_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,plain,
    ( nil = X1
    | cons(esk45_1(X1),esk44_1(X1)) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( esk51_0 = nil
    | app(app(esk44_1(esk51_0),cons(esk45_1(esk51_0),nil)),nil) != esk50_0
    | ~ ssList(esk44_1(esk51_0))
    | ~ ssItem(esk45_1(esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_28])]),c_0_29])])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( nil = X1
    | ssList(esk44_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( esk51_0 = nil
    | app(app(esk44_1(esk51_0),cons(esk45_1(esk51_0),nil)),nil) != esk50_0
    | ~ ssItem(esk45_1(esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31]),c_0_29])])).

<A NAME="c_0_33"></A>cnf(c_0_33,plain,
    ( nil = X1
    | ssItem(esk45_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>])).

<A NAME="c_0_34"></A>fof(c_0_34,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( esk51_0 = nil
    | app(app(esk44_1(esk51_0),cons(esk45_1(esk51_0),nil)),nil) != esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_33]),c_0_29])])).

<A NAME="c_0_36"></A>cnf(c_0_36,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( esk51_0 = nil
    | app(esk44_1(esk51_0),app(cons(esk45_1(esk51_0),nil),nil)) != esk50_0
    | ~ ssList(cons(esk45_1(esk51_0),nil))
    | ~ ssList(esk44_1(esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_36]),c_0_18])])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( esk51_0 = nil
    | app(esk44_1(esk51_0),app(cons(esk45_1(esk51_0),nil),nil)) != esk50_0
    | ~ ssList(cons(esk45_1(esk51_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_31]),c_0_29])])).

<A NAME="c_0_39"></A>fof(c_0_39,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( esk51_0 = nil
    | app(esk44_1(esk51_0),app(cons(esk45_1(esk51_0),nil),nil)) != esk50_0
    | ~ ssItem(esk45_1(esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>,c_0_21]),c_0_18])])).

<A NAME="c_0_41"></A>fof(c_0_41,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssItem(X6)
      | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

<A NAME="c_0_42"></A>cnf(c_0_42,plain,
    ( nil = app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != X1
    | nil != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( esk51_0 = nil
    | app(esk44_1(esk51_0),app(cons(esk45_1(esk51_0),nil),nil)) != esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_33]),c_0_29])])).

<A NAME="c_0_44"></A>cnf(c_0_44,plain,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_42])]),c_0_18])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( app(X2,cons(X1,nil)) = esk50_0
    | app(cons(X1,nil),X2) != esk51_0
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( esk51_0 = nil
    | app(esk44_1(esk51_0),cons(esk45_1(esk51_0),nil)) != esk50_0
    | ~ ssItem(esk45_1(esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_44]),c_0_45]),c_0_18])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( app(X1,cons(X2,nil)) = esk50_0
    | cons(X2,X1) != esk51_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( esk51_0 = nil
    | app(esk44_1(esk51_0),cons(esk45_1(esk51_0),nil)) != esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_33]),c_0_29])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( esk51_0 = nil
    | ~ ssList(esk44_1(esk51_0))
    | ~ ssItem(esk45_1(esk51_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_28])]),c_0_29])]),<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_51"></A>fof(c_0_51,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( neq(esk49_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( esk51_0 = nil
    | ~ ssItem(esk45_1(esk51_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_31]),c_0_29])])).

<A NAME="c_0_54"></A>cnf(c_0_54,plain,
    ( ~ ssList(X1)
    | ~ ssList(X2)
    | X1 != X2
    | ~ neq(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( neq(esk51_0,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( esk51_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_33]),c_0_29])])).

<A NAME="c_0_57"></A>cnf(c_0_57,plain,
    ( ~ ssList(X1)
    | ~ neq(X1,X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( neq(nil,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,c_0_58]),c_0_18])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC420+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n098.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:44:07 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.31  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.31  # and selection function SelectCQPrecWNTNp.
% 0.06/0.31  #
% 0.06/0.31  # Preprocessing time       : 0.014 s
% 0.06/0.31  # Presaturation interreduction done
% 0.06/0.31  
% 0.06/0.31  # Proof found!
% 0.06/0.31  # SZS status Theorem
% 0.06/0.31  # SZS output start CNFRefutation
% 0.06/0.31  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((((~(ssList(X4))|X2!=X4)|X1!=X3)|~(neq(X2,nil)))|?[X5]:(ssItem(X5)&?[X6]:(ssList(X6)&?[X7]:((ssList(X7)&app(app(X6,cons(X5,nil)),X7)=X2)&app(app(X7,cons(X5,nil)),X6)=X1))))|?[X8]:(ssItem(X8)&?[X9]:((ssList(X9)&app(X9,cons(X8,nil))!=X3)&app(cons(X8,nil),X9)=X4)))|(nil!=X3&nil=X4)))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.06/0.31  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
% 0.06/0.31  fof(c_0_2, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.31  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.06/0.31  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.06/0.31  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax20)).
% 0.06/0.31  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.06/0.31  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.06/0.31  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
% 0.06/0.31  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.06/0.31  fof(c_0_10, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((((~ssList(X4)|X2!=X4)|X1!=X3)|~neq(X2,nil))|?[X5]:(ssItem(X5)&?[X6]:(ssList(X6)&?[X7]:((ssList(X7)&app(app(X6,cons(X5,nil)),X7)=X2)&app(app(X7,cons(X5,nil)),X6)=X1))))|?[X8]:(ssItem(X8)&?[X9]:((ssList(X9)&app(X9,cons(X8,nil))!=X3)&app(cons(X8,nil),X9)=X4)))|(nil!=X3&nil=X4))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.31  fof(c_0_11, negated_conjecture, (![X14]:![X15]:![X16]:![X17]:![X18]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&((((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&neq(esk49_0,nil))&(~ssItem(X14)|(~ssList(X15)|((~ssList(X16)|app(app(X15,cons(X14,nil)),X16)!=esk49_0)|app(app(X16,cons(X14,nil)),X15)!=esk48_0))))&(~ssItem(X17)|((~ssList(X18)|app(X18,cons(X17,nil))=esk50_0)|app(cons(X17,nil),X18)!=esk51_0)))&(nil=esk50_0|nil!=esk51_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).
% 0.06/0.31  cnf(c_0_12, negated_conjecture, (app(app(X1,cons(X2,nil)),X3)!=esk48_0|app(app(X3,cons(X2,nil)),X1)!=esk49_0|~ssList(X1)|~ssList(X3)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.06/0.31  cnf(c_0_13, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.06/0.31  cnf(c_0_14, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.06/0.31  fof(c_0_15, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
% 0.06/0.31  cnf(c_0_16, negated_conjecture, (app(app(X1,cons(X2,nil)),X3)!=esk51_0|app(app(X3,cons(X2,nil)),X1)!=esk50_0|~ssList(X1)|~ssList(X3)|~ssItem(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13]), c_0_14])).
% 0.06/0.31  cnf(c_0_17, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
% 0.06/0.31  cnf(c_0_18, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
% 0.06/0.31  fof(c_0_19, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
% 0.06/0.31  cnf(c_0_20, negated_conjecture, (app(app(X1,cons(X2,nil)),nil)!=esk50_0|app(cons(X2,nil),X1)!=esk51_0|~ssList(cons(X2,nil))|~ssList(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])])).
% 0.06/0.31  cnf(c_0_21, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.31  fof(c_0_22, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 0.06/0.31  cnf(c_0_23, negated_conjecture, (app(app(X1,cons(X2,nil)),nil)!=esk50_0|app(cons(X2,nil),X1)!=esk51_0|~ssList(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_18])])).
% 0.06/0.31  cnf(c_0_24, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
% 0.06/0.31  fof(c_0_25, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
% 0.06/0.31  cnf(c_0_26, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.06/0.31  cnf(c_0_27, negated_conjecture, (app(app(X1,cons(X2,nil)),nil)!=esk50_0|cons(X2,X1)!=esk51_0|~ssList(X1)|~ssItem(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
% 0.06/0.31  cnf(c_0_28, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
% 0.06/0.31  cnf(c_0_29, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_26, c_0_13])).
% 0.06/0.31  cnf(c_0_30, negated_conjecture, (esk51_0=nil|app(app(esk44_1(esk51_0),cons(esk45_1(esk51_0),nil)),nil)!=esk50_0|~ssList(esk44_1(esk51_0))|~ssItem(esk45_1(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28])]), c_0_29])])).
% 0.06/0.31  cnf(c_0_31, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
% 0.06/0.31  cnf(c_0_32, negated_conjecture, (esk51_0=nil|app(app(esk44_1(esk51_0),cons(esk45_1(esk51_0),nil)),nil)!=esk50_0|~ssItem(esk45_1(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_29])])).
% 0.06/0.31  cnf(c_0_33, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
% 0.06/0.31  fof(c_0_34, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
% 0.06/0.31  cnf(c_0_35, negated_conjecture, (esk51_0=nil|app(app(esk44_1(esk51_0),cons(esk45_1(esk51_0),nil)),nil)!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_29])])).
% 0.06/0.31  cnf(c_0_36, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
% 0.06/0.31  cnf(c_0_37, negated_conjecture, (esk51_0=nil|app(esk44_1(esk51_0),app(cons(esk45_1(esk51_0),nil),nil))!=esk50_0|~ssList(cons(esk45_1(esk51_0),nil))|~ssList(esk44_1(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_18])])).
% 0.06/0.31  cnf(c_0_38, negated_conjecture, (esk51_0=nil|app(esk44_1(esk51_0),app(cons(esk45_1(esk51_0),nil),nil))!=esk50_0|~ssList(cons(esk45_1(esk51_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_31]), c_0_29])])).
% 0.06/0.31  fof(c_0_39, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
% 0.06/0.31  cnf(c_0_40, negated_conjecture, (esk51_0=nil|app(esk44_1(esk51_0),app(cons(esk45_1(esk51_0),nil),nil))!=esk50_0|~ssItem(esk45_1(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_21]), c_0_18])])).
% 0.06/0.31  fof(c_0_41, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
% 0.06/0.31  cnf(c_0_42, plain, (nil=app(X1,X2)|~ssList(X1)|~ssList(X2)|nil!=X1|nil!=X2), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.31  cnf(c_0_43, negated_conjecture, (esk51_0=nil|app(esk44_1(esk51_0),app(cons(esk45_1(esk51_0),nil),nil))!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_33]), c_0_29])])).
% 0.06/0.31  cnf(c_0_44, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 0.06/0.31  cnf(c_0_45, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_42])]), c_0_18])])).
% 0.06/0.31  cnf(c_0_46, negated_conjecture, (app(X2,cons(X1,nil))=esk50_0|app(cons(X1,nil),X2)!=esk51_0|~ssList(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.06/0.31  cnf(c_0_47, negated_conjecture, (esk51_0=nil|app(esk44_1(esk51_0),cons(esk45_1(esk51_0),nil))!=esk50_0|~ssItem(esk45_1(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_18])])).
% 0.06/0.31  cnf(c_0_48, negated_conjecture, (app(X1,cons(X2,nil))=esk50_0|cons(X2,X1)!=esk51_0|~ssList(X1)|~ssItem(X2)), inference(spm,[status(thm)],[c_0_46, c_0_24])).
% 0.06/0.31  cnf(c_0_49, negated_conjecture, (esk51_0=nil|app(esk44_1(esk51_0),cons(esk45_1(esk51_0),nil))!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_33]), c_0_29])])).
% 0.06/0.31  cnf(c_0_50, negated_conjecture, (esk51_0=nil|~ssList(esk44_1(esk51_0))|~ssItem(esk45_1(esk51_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_28])]), c_0_29])]), c_0_49])).
% 0.06/0.31  fof(c_0_51, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
% 0.06/0.31  cnf(c_0_52, negated_conjecture, (neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.06/0.31  cnf(c_0_53, negated_conjecture, (esk51_0=nil|~ssItem(esk45_1(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_31]), c_0_29])])).
% 0.06/0.31  cnf(c_0_54, plain, (~ssList(X1)|~ssList(X2)|X1!=X2|~neq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
% 0.06/0.31  cnf(c_0_55, negated_conjecture, (neq(esk51_0,nil)), inference(rw,[status(thm)],[c_0_52, c_0_13])).
% 0.06/0.31  cnf(c_0_56, negated_conjecture, (esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_33]), c_0_29])])).
% 0.06/0.31  cnf(c_0_57, plain, (~ssList(X1)|~neq(X1,X1)), inference(er,[status(thm)],[c_0_54])).
% 0.06/0.31  cnf(c_0_58, negated_conjecture, (neq(nil,nil)), inference(rw,[status(thm)],[c_0_55, c_0_56])).
% 0.06/0.31  cnf(c_0_59, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_18])]), ['proof']).
% 0.06/0.31  # SZS output end CNFRefutation
% 0.06/0.31  # Training examples: 0 positive, 0 negative
% 0.06/0.31  
% 0.06/0.31  # -------------------------------------------------
% 0.06/0.31  # User time                : 0.078 s
% 0.06/0.31  # System time              : 0.007 s
% 0.06/0.31  # Total time               : 0.084 s
% 0.06/0.31  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
