   --------------------------------------------------------------------------

                   TSTP Solution File: SWC051+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC051+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n180.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:27 EDT 2016

 % Result     : Theorem 0.02s
 % Output     : CNFRefutation 0.02s
 % Verified   :
 % Statistics : Number of formulae       :   19 (  54 expanded)
 %              Number of clauses        :   14 (  19 expanded)
 %              Number of leaves         :    2 (  12 expanded)
 %              Depth                    :    6
 %              Number of atoms          :   76 ( 571 expanded)
 %              Number of equality atoms :   14 ( 137 expanded)
 %              Maximal formula depth    :   18 (   4 average)
 %              Maximal clause size      :   17 (   2 average)
 %              Maximal term depth       :    1 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ~ neq(X2,nil)
                     | ? [X5] :
                         ( ssList(X5)
                         & neq(X5,nil)
                         & rearsegP(X2,X5)
                         & rearsegP(X1,X5) )
                     | ( nil != X3
                       & nil = X4 )
                     | ( neq(X4,nil)
                       & ( ~ neq(X3,nil)
                         | ~ rearsegP(X4,X3) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => rearsegP(X1,X1) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax49)).

 fof(c_0_2,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ~ neq(X2,nil)
                       | ? [X5] :
                           ( ssList(X5)
                           & neq(X5,nil)
                           & rearsegP(X2,X5)
                           & rearsegP(X1,X5) )
                       | ( nil != X3
                         & nil = X4 )
                       | ( neq(X4,nil)
                         & ( ~ neq(X3,nil)
                           | ~ rearsegP(X4,X3) ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_3,negated_conjecture,(
     ! [X10] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & neq(esk49_0,nil)
       & ( ~ ssList(X10)
         | ~ neq(X10,nil)
         | ~ rearsegP(esk49_0,X10)
         | ~ rearsegP(esk48_0,X10) )
       & ( nil = esk50_0
         | nil != esk51_0 )
       & ( neq(esk50_0,nil)
         | ~ neq(esk51_0,nil) )
       & ( rearsegP(esk51_0,esk50_0)
         | ~ neq(esk51_0,nil) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).

 cnf(c_0_4,negated_conjecture,
     ( neq(esk49_0,nil) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_3])).

 cnf(c_0_5,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_3])).

 cnf(c_0_6,negated_conjecture,
     ( ~ rearsegP(esk48_0,X1)
     | ~ rearsegP(esk49_0,X1)
     | ~ neq(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[6]c_0_3])).

 cnf(c_0_7,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[7]c_0_3])).

 fof(c_0_8,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | rearsegP(X2,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).

 cnf(c_0_9,negated_conjecture,
     ( rearsegP(esk51_0,esk50_0)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[8]c_0_3])).

 cnf(c_0_10,negated_conjecture,
     ( neq(esk51_0,nil) ),
     inference(rw,[status(thm)],[[9]c_0_4,[10]c_0_5])).

 cnf(c_0_11,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_3])).

 cnf(c_0_12,negated_conjecture,
     ( neq(esk50_0,nil)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_3])).

 cnf(c_0_13,negated_conjecture,
     ( ~ rearsegP(esk50_0,X1)
     | ~ rearsegP(esk51_0,X1)
     | ~ ssList(X1)
     | ~ neq(X1,nil) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[13]c_0_6,c_0_7]),[14]c_0_5])).

 cnf(c_0_14,plain,
     ( rearsegP(X1,X1)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_8])).

 cnf(c_0_15,negated_conjecture,
     ( rearsegP(esk51_0,esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[16]c_0_9,c_0_10])])).

 cnf(c_0_16,negated_conjecture,
     ( ssList(esk50_0) ),
     inference(rw,[status(thm)],[[17]c_0_11,[18]c_0_7])).

 cnf(c_0_17,negated_conjecture,
     ( neq(esk50_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[19]c_0_12,c_0_10])])).

 cnf(c_0_18,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[20]c_0_13,c_0_14]),c_0_15]),c_0_16]),c_0_17])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC051+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.22  % Computer   : n180.star.cs.uiowa.edu
 % 0.02/0.22  % Model      : x86_64 x86_64
 % 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.22  % Memory     : 32218.75MB
 % 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.22  % CPULimit   : 300
 % 0.02/0.22  % DateTime   : Tue May 10 20:34:53 CDT 2016
 % 0.02/0.22  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.02/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.02/0.25  # and selection function SelectCQPrecWNTNp.
 % 0.02/0.25  #
 % 0.02/0.25  # Preprocessing time       : 0.014 s
 % 0.02/0.25  # Presaturation interreduction done
 % 0.02/0.25
 % 0.02/0.25  # Proof found!
 % 0.02/0.25  # SZS status Theorem
 % 0.02/0.25  # SZS output start CNFRefutation
 % 0.02/0.25  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~(neq(X2,nil)))|?[X5]:(((ssList(X5)&neq(X5,nil))&rearsegP(X2,X5))&rearsegP(X1,X5)))|(nil!=X3&nil=X4))|(neq(X4,nil)&(~(neq(X3,nil))|~(rearsegP(X4,X3)))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.02/0.25  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>rearsegP(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax49)).
 % 0.02/0.25  fof(c_0_2, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~neq(X2,nil))|?[X5]:(((ssList(X5)&neq(X5,nil))&rearsegP(X2,X5))&rearsegP(X1,X5)))|(nil!=X3&nil=X4))|(neq(X4,nil)&(~neq(X3,nil)|~rearsegP(X4,X3)))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.02/0.25  fof(c_0_3, negated_conjecture, (![X10]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((((esk49_0=esk51_0&esk48_0=esk50_0)&neq(esk49_0,nil))&(((~ssList(X10)|~neq(X10,nil))|~rearsegP(esk49_0,X10))|~rearsegP(esk48_0,X10)))&(nil=esk50_0|nil!=esk51_0))&((neq(esk50_0,nil)|~neq(esk51_0,nil))&(rearsegP(esk51_0,esk50_0)|~neq(esk51_0,nil))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
 % 0.02/0.25  cnf(c_0_4, negated_conjecture, (neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
 % 0.02/0.25  cnf(c_0_5, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_3])).
 % 0.02/0.25  cnf(c_0_6, negated_conjecture, (~rearsegP(esk48_0,X1)|~rearsegP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
 % 0.02/0.25  cnf(c_0_7, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_3])).
 % 0.02/0.25  fof(c_0_8, plain, (![X2]:(~ssList(X2)|rearsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
 % 0.02/0.25  cnf(c_0_9, negated_conjecture, (rearsegP(esk51_0,esk50_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
 % 0.02/0.25  cnf(c_0_10, negated_conjecture, (neq(esk51_0,nil)), inference(rw,[status(thm)],[c_0_4, c_0_5])).
 % 0.02/0.25  cnf(c_0_11, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_3])).
 % 0.02/0.25  cnf(c_0_12, negated_conjecture, (neq(esk50_0,nil)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
 % 0.02/0.25  cnf(c_0_13, negated_conjecture, (~rearsegP(esk50_0,X1)|~rearsegP(esk51_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_7]), c_0_5])).
 % 0.02/0.25  cnf(c_0_14, plain, (rearsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
 % 0.02/0.25  cnf(c_0_15, negated_conjecture, (rearsegP(esk51_0,esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_10])])).
 % 0.02/0.25  cnf(c_0_16, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_11, c_0_7])).
 % 0.02/0.25  cnf(c_0_17, negated_conjecture, (neq(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_10])])).
 % 0.02/0.25  cnf(c_0_18, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17])]), ['proof']).
 % 0.02/0.25  # SZS output end CNFRefutation
 % 0.02/0.25  # Training examples: 0 positive, 0 negative
 % 0.02/0.25
 % 0.02/0.25  # -------------------------------------------------
 % 0.02/0.25  # User time                : 0.025 s
 % 0.02/0.25  # System time              : 0.002 s
 % 0.02/0.25  # Total time               : 0.028 s
 % 0.02/0.25  # Maximum resident set size: 1876 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
