   --------------------------------------------------------------------------

                   TSTP Solution File: SWC054+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC054+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n128.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:27 EDT 2016

 % Result     : Theorem 0.02s
 % Output     : CNFRefutation 0.02s
 % Verified   :
 % Statistics : Number of formulae       :   40 ( 621 expanded)
 %              Number of clauses        :   32 ( 217 expanded)
 %              Number of leaves         :    4 ( 136 expanded)
 %              Depth                    :    9
 %              Number of atoms          :  154 (9716 expanded)
 %              Number of equality atoms :   46 (2852 expanded)
 %              Maximal formula depth    :   21 (   4 average)
 %              Maximal clause size      :   30 (   2 average)
 %              Maximal term depth       :    1 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ~ ssList(X4)
                   | X2 != X4
                   | X1 != X3
                   | ( nil != X3
                     & nil = X4 )
                   | ( ! [X5] :
                         ( ~ ssList(X5)
                         | ~ neq(X5,nil)
                         | ~ segmentP(X4,X5)
                         | ~ segmentP(X3,X5) )
                     & neq(X4,nil) )
                   | ( ( nil != X2
                       | nil = X1 )
                     & ( ~ neq(X2,nil)
                       | ? [X6] :
                           ( ssList(X6)
                           & neq(X6,nil)
                           & segmentP(X2,X6)
                           & segmentP(X1,X6) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => segmentP(X1,X1) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax55)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( segmentP(nil,X1)
       <=> nil = X1 ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax58)).

 fof(c_0_3,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_4,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ~ ssList(X4)
                     | X2 != X4
                     | X1 != X3
                     | ( nil != X3
                       & nil = X4 )
                     | ( ! [X5] :
                           ( ~ ssList(X5)
                           | ~ neq(X5,nil)
                           | ~ segmentP(X4,X5)
                           | ~ segmentP(X3,X5) )
                       & neq(X4,nil) )
                     | ( ( nil != X2
                         | nil = X1 )
                       & ( ~ neq(X2,nil)
                         | ? [X6] :
                             ( ssList(X6)
                             & neq(X6,nil)
                             & segmentP(X2,X6)
                             & segmentP(X1,X6) ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_5,negated_conjecture,(
     ! [X12] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ( nil = esk50_0
         | nil != esk51_0 )
       & ( ssList(esk52_0)
         | ~ neq(esk51_0,nil) )
       & ( neq(esk52_0,nil)
         | ~ neq(esk51_0,nil) )
       & ( segmentP(esk51_0,esk52_0)
         | ~ neq(esk51_0,nil) )
       & ( segmentP(esk50_0,esk52_0)
         | ~ neq(esk51_0,nil) )
       & ( neq(esk49_0,nil)
         | nil = esk49_0 )
       & ( ~ ssList(X12)
         | ~ neq(X12,nil)
         | ~ segmentP(esk49_0,X12)
         | ~ segmentP(esk48_0,X12)
         | nil = esk49_0 )
       & ( neq(esk49_0,nil)
         | nil != esk48_0 )
       & ( ~ ssList(X12)
         | ~ neq(X12,nil)
         | ~ segmentP(esk49_0,X12)
         | ~ segmentP(esk48_0,X12)
         | nil != esk48_0 ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).

 cnf(c_0_6,negated_conjecture,
     ( segmentP(esk50_0,esk52_0)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_5])).

 cnf(c_0_7,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_5])).

 cnf(c_0_8,negated_conjecture,
     ( nil = esk49_0
     | neq(esk49_0,nil) ),
     inference(split_conjunct,[status(thm)],[[6]c_0_5])).

 cnf(c_0_9,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[7]c_0_5])).

 cnf(c_0_10,negated_conjecture,
     ( nil != esk48_0
     | ~ segmentP(esk48_0,X1)
     | ~ segmentP(esk49_0,X1)
     | ~ neq(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[8]c_0_5])).

 fof(c_0_11,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | segmentP(X2,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).

 cnf(c_0_12,negated_conjecture,
     ( nil = esk49_0
     | ~ segmentP(esk48_0,X1)
     | ~ segmentP(esk49_0,X1)
     | ~ neq(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[9]c_0_5])).

 cnf(c_0_13,negated_conjecture,
     ( segmentP(esk48_0,esk52_0)
     | ~ neq(esk51_0,nil) ),
     inference(rw,[status(thm)],[[10]c_0_6,[11]c_0_7])).

 cnf(c_0_14,negated_conjecture,
     ( esk51_0 = nil
     | neq(esk51_0,nil) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[12]c_0_8,c_0_9]),[13]c_0_9])).

 cnf(c_0_15,negated_conjecture,
     ( neq(esk52_0,nil)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_5])).

 cnf(c_0_16,negated_conjecture,
     ( ssList(esk52_0)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_5])).

 cnf(c_0_17,negated_conjecture,
     ( segmentP(esk51_0,esk52_0)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_5])).

 cnf(c_0_18,negated_conjecture,
     ( esk48_0 != nil
     | ~ segmentP(esk48_0,X1)
     | ~ segmentP(esk51_0,X1)
     | ~ ssList(X1)
     | ~ neq(X1,nil) ),
     inference(rw,[status(thm)],[[17]c_0_10,[18]c_0_9])).

 cnf(c_0_19,plain,
     ( segmentP(X1,X1)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_11])).

 cnf(c_0_20,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[20]c_0_5])).

 cnf(c_0_21,negated_conjecture,
     ( esk51_0 = nil
     | ~ segmentP(esk48_0,X1)
     | ~ segmentP(esk51_0,X1)
     | ~ ssList(X1)
     | ~ neq(X1,nil) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[21]c_0_12,c_0_9]),[22]c_0_9])).

 cnf(c_0_22,negated_conjecture,
     ( esk51_0 = nil
     | segmentP(esk48_0,esk52_0) ),
     inference(spm,[status(thm)],[[23]c_0_13,[24]c_0_14])).

 cnf(c_0_23,negated_conjecture,
     ( esk51_0 = nil
     | neq(esk52_0,nil) ),
     inference(spm,[status(thm)],[[25]c_0_15,[26]c_0_14])).

 cnf(c_0_24,negated_conjecture,
     ( esk51_0 = nil
     | ssList(esk52_0) ),
     inference(spm,[status(thm)],[[27]c_0_16,[28]c_0_14])).

 cnf(c_0_25,negated_conjecture,
     ( esk51_0 = nil
     | segmentP(esk51_0,esk52_0) ),
     inference(spm,[status(thm)],[[29]c_0_17,[30]c_0_14])).

 cnf(c_0_26,negated_conjecture,
     ( nil = esk50_0
     | nil != esk51_0 ),
     inference(split_conjunct,[status(thm)],[[31]c_0_5])).

 fof(c_0_27,plain,(
     ! [X2] :
       ( ( ~ segmentP(nil,X2)
         | nil = X2
         | ~ ssList(X2) )
       & ( nil != X2
         | segmentP(nil,X2)
         | ~ ssList(X2) ) ) ),
     inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

 cnf(c_0_28,negated_conjecture,
     ( esk48_0 != nil
     | ~ segmentP(esk51_0,esk48_0)
     | ~ neq(esk48_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[32]c_0_18,c_0_19]),c_0_20])])).

 cnf(c_0_29,negated_conjecture,
     ( esk51_0 = nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[33]c_0_21,c_0_22]),c_0_23]),c_0_24]),[34]c_0_25])).

 cnf(c_0_30,negated_conjecture,
     ( esk48_0 = nil
     | esk51_0 != nil ),
     inference(rw,[status(thm)],[[35]c_0_26,[36]c_0_7])).

 cnf(c_0_31,plain,
     ( segmentP(nil,X1)
     | ~ ssList(X1)
     | nil != X1 ),
     inference(split_conjunct,[status(thm)],[[37]c_0_27])).

 cnf(c_0_32,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[38]c_0_3])).

 cnf(c_0_33,negated_conjecture,
     ( neq(esk49_0,nil)
     | nil != esk48_0 ),
     inference(split_conjunct,[status(thm)],[[39]c_0_5])).

 cnf(c_0_34,negated_conjecture,
     ( esk48_0 != nil
     | ~ segmentP(nil,esk48_0)
     | ~ neq(esk48_0,nil) ),
     inference(rw,[status(thm)],[[40]c_0_28,[41]c_0_29])).

 cnf(c_0_35,negated_conjecture,
     ( esk48_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[42]c_0_30,c_0_29])])).

 cnf(c_0_36,plain,
     ( segmentP(nil,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_31]),c_0_32])])).

 cnf(c_0_37,negated_conjecture,
     ( neq(esk51_0,nil)
     | esk48_0 != nil ),
     inference(rw,[status(thm)],[[43]c_0_33,[44]c_0_9])).

 cnf(c_0_38,negated_conjecture,
     ( ~ neq(nil,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[45]c_0_34,c_0_35]),c_0_35]),c_0_36]),c_0_35])])).

 cnf(c_0_39,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[46]c_0_37,c_0_29]),c_0_35])]),[47]c_0_38]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC054+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.22  % Computer   : n128.star.cs.uiowa.edu
 % 0.02/0.22  % Model      : x86_64 x86_64
 % 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.22  % Memory     : 32218.75MB
 % 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.22  % CPULimit   : 300
 % 0.02/0.22  % DateTime   : Tue May 10 20:34:53 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.02/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.02/0.25  # and selection function SelectCQPrecWNTNp.
 % 0.02/0.25  #
 % 0.02/0.25  # Preprocessing time       : 0.014 s
 % 0.02/0.25  # Presaturation interreduction done
 % 0.02/0.25
 % 0.02/0.25  # Proof found!
 % 0.02/0.25  # SZS status Theorem
 % 0.02/0.25  # SZS output start CNFRefutation
 % 0.02/0.25  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~(ssList(X4))|X2!=X4)|X1!=X3)|(nil!=X3&nil=X4))|(![X5]:(((~(ssList(X5))|~(neq(X5,nil)))|~(segmentP(X4,X5)))|~(segmentP(X3,X5)))&neq(X4,nil)))|((nil!=X2|nil=X1)&(~(neq(X2,nil))|?[X6]:(((ssList(X6)&neq(X6,nil))&segmentP(X2,X6))&segmentP(X1,X6))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
 % 0.02/0.25  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>segmentP(X1,X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax55)).
 % 0.02/0.25  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>(segmentP(nil,X1)<=>nil=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax58)).
 % 0.02/0.25  fof(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.02/0.25  fof(c_0_4, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~ssList(X4)|X2!=X4)|X1!=X3)|(nil!=X3&nil=X4))|(![X5]:(((~ssList(X5)|~neq(X5,nil))|~segmentP(X4,X5))|~segmentP(X3,X5))&neq(X4,nil)))|((nil!=X2|nil=X1)&(~neq(X2,nil)|?[X6]:(((ssList(X6)&neq(X6,nil))&segmentP(X2,X6))&segmentP(X1,X6)))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.02/0.25  fof(c_0_5, negated_conjecture, (![X12]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(nil=esk50_0|nil!=esk51_0))&((((ssList(esk52_0)|~neq(esk51_0,nil))&(neq(esk52_0,nil)|~neq(esk51_0,nil)))&(segmentP(esk51_0,esk52_0)|~neq(esk51_0,nil)))&(segmentP(esk50_0,esk52_0)|~neq(esk51_0,nil))))&(((neq(esk49_0,nil)|nil=esk49_0)&((((~ssList(X12)|~neq(X12,nil))|~segmentP(esk49_0,X12))|~segmentP(esk48_0,X12))|nil=esk49_0))&((neq(esk49_0,nil)|nil!=esk48_0)&((((~ssList(X12)|~neq(X12,nil))|~segmentP(esk49_0,X12))|~segmentP(esk48_0,X12))|nil!=esk48_0)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).
 % 0.02/0.25  cnf(c_0_6, negated_conjecture, (segmentP(esk50_0,esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  cnf(c_0_7, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  cnf(c_0_8, negated_conjecture, (nil=esk49_0|neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  cnf(c_0_9, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  cnf(c_0_10, negated_conjecture, (nil!=esk48_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  fof(c_0_11, plain, (![X2]:(~ssList(X2)|segmentP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
 % 0.02/0.25  cnf(c_0_12, negated_conjecture, (nil=esk49_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  cnf(c_0_13, negated_conjecture, (segmentP(esk48_0,esk52_0)|~neq(esk51_0,nil)), inference(rw,[status(thm)],[c_0_6, c_0_7])).
 % 0.02/0.25  cnf(c_0_14, negated_conjecture, (esk51_0=nil|neq(esk51_0,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_9]), c_0_9])).
 % 0.02/0.25  cnf(c_0_15, negated_conjecture, (neq(esk52_0,nil)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  cnf(c_0_16, negated_conjecture, (ssList(esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  cnf(c_0_17, negated_conjecture, (segmentP(esk51_0,esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  cnf(c_0_18, negated_conjecture, (esk48_0!=nil|~segmentP(esk48_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[c_0_10, c_0_9])).
 % 0.02/0.25  cnf(c_0_19, plain, (segmentP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.02/0.25  cnf(c_0_20, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  cnf(c_0_21, negated_conjecture, (esk51_0=nil|~segmentP(esk48_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_9]), c_0_9])).
 % 0.02/0.25  cnf(c_0_22, negated_conjecture, (esk51_0=nil|segmentP(esk48_0,esk52_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
 % 0.02/0.25  cnf(c_0_23, negated_conjecture, (esk51_0=nil|neq(esk52_0,nil)), inference(spm,[status(thm)],[c_0_15, c_0_14])).
 % 0.02/0.25  cnf(c_0_24, negated_conjecture, (esk51_0=nil|ssList(esk52_0)), inference(spm,[status(thm)],[c_0_16, c_0_14])).
 % 0.02/0.25  cnf(c_0_25, negated_conjecture, (esk51_0=nil|segmentP(esk51_0,esk52_0)), inference(spm,[status(thm)],[c_0_17, c_0_14])).
 % 0.02/0.25  cnf(c_0_26, negated_conjecture, (nil=esk50_0|nil!=esk51_0), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  fof(c_0_27, plain, (![X2]:(((~segmentP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|segmentP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
 % 0.02/0.25  cnf(c_0_28, negated_conjecture, (esk48_0!=nil|~segmentP(esk51_0,esk48_0)|~neq(esk48_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])])).
 % 0.02/0.25  cnf(c_0_29, negated_conjecture, (esk51_0=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24]), c_0_25])).
 % 0.02/0.25  cnf(c_0_30, negated_conjecture, (esk48_0=nil|esk51_0!=nil), inference(rw,[status(thm)],[c_0_26, c_0_7])).
 % 0.02/0.25  cnf(c_0_31, plain, (segmentP(nil,X1)|~ssList(X1)|nil!=X1), inference(split_conjunct,[status(thm)],[c_0_27])).
 % 0.02/0.25  cnf(c_0_32, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
 % 0.02/0.25  cnf(c_0_33, negated_conjecture, (neq(esk49_0,nil)|nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  cnf(c_0_34, negated_conjecture, (esk48_0!=nil|~segmentP(nil,esk48_0)|~neq(esk48_0,nil)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
 % 0.02/0.25  cnf(c_0_35, negated_conjecture, (esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_29])])).
 % 0.02/0.25  cnf(c_0_36, plain, (segmentP(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_31]), c_0_32])])).
 % 0.02/0.25  cnf(c_0_37, negated_conjecture, (neq(esk51_0,nil)|esk48_0!=nil), inference(rw,[status(thm)],[c_0_33, c_0_9])).
 % 0.02/0.25  cnf(c_0_38, negated_conjecture, (~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35]), c_0_35]), c_0_36]), c_0_35])])).
 % 0.02/0.25  cnf(c_0_39, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_29]), c_0_35])]), c_0_38]), ['proof']).
 % 0.02/0.25  # SZS output end CNFRefutation
 % 0.02/0.25  # Training examples: 0 positive, 0 negative
 % 0.02/0.25
 % 0.02/0.25  # -------------------------------------------------
 % 0.02/0.25  # User time                : 0.025 s
 % 0.02/0.25  # System time              : 0.003 s
 % 0.02/0.25  # Total time               : 0.029 s
 % 0.02/0.25  # Maximum resident set size: 1872 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
