   --------------------------------------------------------------------------

                   TSTP Solution File: SWC013+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC013+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n150.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:17 EDT 2016

 % Result     : Theorem 0.05s
 % Output     : CNFRefutation 0.05s
 % Verified   :
 % Statistics : Number of formulae       :   30 (  77 expanded)
 %              Number of clauses        :   20 (  30 expanded)
 %              Number of leaves         :    5 (  19 expanded)
 %              Depth                    :   11
 %              Number of atoms          :  156 (1062 expanded)
 %              Number of equality atoms :   49 ( 385 expanded)
 %              Maximal formula depth    :   23 (   5 average)
 %              Maximal clause size      :   38 (   3 average)
 %              Maximal term depth       :    3 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ( ( ~ neq(X2,nil)
                         | ? [X5] :
                             ( ssList(X5)
                             & X3 != X5
                             & ? [X6] :
                                 ( ssItem(X6)
                                 & cons(X6,nil) = X5
                                 & hd(X4) = X6
                                 & neq(nil,X4) ) )
                         | ? [X7] :
                             ( ssList(X7)
                             & X1 = X7
                             & ? [X8] :
                                 ( ssItem(X8)
                                 & cons(X8,nil) = X7
                                 & hd(X2) = X8
                                 & neq(nil,X2) ) ) )
                       & ( ~ neq(X2,nil)
                         | neq(X4,nil) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_2,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssItem(hd(X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax22)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_5,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ( ( ~ neq(X2,nil)
                           | ? [X5] :
                               ( ssList(X5)
                               & X3 != X5
                               & ? [X6] :
                                   ( ssItem(X6)
                                   & cons(X6,nil) = X5
                                   & hd(X4) = X6
                                   & neq(nil,X4) ) )
                           | ? [X7] :
                               ( ssList(X7)
                               & X1 = X7
                               & ? [X8] :
                                   ( ssItem(X8)
                                   & cons(X8,nil) = X7
                                   & hd(X2) = X8
                                   & neq(nil,X2) ) ) )
                         & ( ~ neq(X2,nil)
                           | neq(X4,nil) ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_6,negated_conjecture,(
     ! [X13,X14,X15,X16] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ( neq(esk49_0,nil)
         | neq(esk49_0,nil) )
       & ( ~ neq(esk51_0,nil)
         | neq(esk49_0,nil) )
       & ( neq(esk49_0,nil)
         | ~ ssList(X13)
         | esk50_0 = X13
         | ~ ssItem(X14)
         | cons(X14,nil) != X13
         | hd(esk51_0) != X14
         | ~ neq(nil,esk51_0) )
       & ( ~ neq(esk51_0,nil)
         | ~ ssList(X13)
         | esk50_0 = X13
         | ~ ssItem(X14)
         | cons(X14,nil) != X13
         | hd(esk51_0) != X14
         | ~ neq(nil,esk51_0) )
       & ( neq(esk49_0,nil)
         | ~ ssList(X15)
         | esk48_0 != X15
         | ~ ssItem(X16)
         | cons(X16,nil) != X15
         | hd(esk49_0) != X16
         | ~ neq(nil,esk49_0) )
       & ( ~ neq(esk51_0,nil)
         | ~ ssList(X15)
         | esk48_0 != X15
         | ~ ssItem(X16)
         | cons(X16,nil) != X15
         | hd(esk49_0) != X16
         | ~ neq(nil,esk49_0) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])).

 cnf(c_0_7,negated_conjecture,
     ( neq(esk49_0,nil)
     | neq(esk49_0,nil) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_6])).

 cnf(c_0_8,negated_conjecture,
     ( ~ neq(nil,esk49_0)
     | hd(esk49_0) != X1
     | cons(X1,nil) != X2
     | ~ ssItem(X1)
     | esk48_0 != X2
     | ~ ssList(X2)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[5]c_0_6])).

 cnf(c_0_9,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_6])).

 cnf(c_0_10,negated_conjecture,
     ( neq(esk49_0,nil) ),
     inference(cn,[status(thm)],[[7]c_0_7])).

 cnf(c_0_11,negated_conjecture,
     ( esk50_0 = X2
     | ~ neq(nil,esk51_0)
     | hd(esk51_0) != X1
     | cons(X1,nil) != X2
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[8]c_0_6])).

 cnf(c_0_12,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[9]c_0_6])).

 cnf(c_0_13,negated_conjecture,
     ( cons(hd(esk49_0),nil) != esk48_0
     | ~ ssList(cons(hd(esk49_0),nil))
     | ~ neq(nil,esk49_0)
     | ~ ssItem(hd(esk49_0)) ),
     inference(er,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[10]c_0_8,c_0_9]),c_0_10])])])])).

 fof(c_0_14,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

 cnf(c_0_15,negated_conjecture,
     ( ~ ssList(cons(hd(esk49_0),nil))
     | ~ neq(nil,esk49_0)
     | ~ ssItem(hd(esk49_0)) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[11]c_0_11,c_0_12]),c_0_9]),c_0_9]),c_0_9]),c_0_10])])])]),[12]c_0_13])).

 cnf(c_0_16,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[13]c_0_14])).

 cnf(c_0_17,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_2])).

 fof(c_0_18,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssItem(hd(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

 cnf(c_0_19,negated_conjecture,
     ( ~ neq(nil,esk49_0)
     | ~ ssItem(hd(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_15,c_0_16]),c_0_17])])).

 cnf(c_0_20,plain,
     ( ssItem(hd(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_18])).

 cnf(c_0_21,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[17]c_0_6])).

 fof(c_0_22,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).

 cnf(c_0_23,negated_conjecture,
     ( esk49_0 = nil
     | ~ neq(nil,esk49_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_19,c_0_20]),c_0_21])])).

 cnf(c_0_24,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_22])).

 cnf(c_0_25,plain,
     ( ~ ssList(X1)
     | ~ ssList(X2)
     | X1 != X2
     | ~ neq(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[20]c_0_22])).

 cnf(c_0_26,negated_conjecture,
     ( esk49_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[21]c_0_23,c_0_24]),c_0_21]),c_0_17])])).

 cnf(c_0_27,plain,
     ( ~ ssList(X1)
     | ~ neq(X1,X1) ),
     inference(er,[status(thm)],[[22]c_0_25])).

 cnf(c_0_28,negated_conjecture,
     ( neq(nil,nil) ),
     inference(rw,[status(thm)],[[23]c_0_10,[24]c_0_26])).

 cnf(c_0_29,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_27,c_0_28]),c_0_17])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.05  % Problem    : SWC013+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.06  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.05/0.26  % Computer   : n150.star.cs.uiowa.edu
 % 0.05/0.26  % Model      : x86_64 x86_64
 % 0.05/0.26  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.05/0.26  % Memory     : 32218.75MB
 % 0.05/0.26  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.05/0.26  % CPULimit   : 300
 % 0.05/0.26  % DateTime   : Tue May 10 20:34:38 CDT 2016
 % 0.05/0.26  % CPUTime    :
 % 0.05/0.26  # Version: 1.9.1pre011
 % 0.05/0.26  # No SInE strategy applied
 % 0.05/0.26  # Trying AutoSched0 for 151 seconds
 % 0.05/0.29  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.05/0.29  # and selection function SelectCQPrecWNTNp.
 % 0.05/0.29  #
 % 0.05/0.29  # Preprocessing time       : 0.014 s
 % 0.05/0.29  # Presaturation interreduction done
 % 0.05/0.29
 % 0.05/0.29  # Proof found!
 % 0.05/0.29  # SZS status Theorem
 % 0.05/0.29  # SZS output start CNFRefutation
 % 0.05/0.29  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((~(neq(X2,nil))|?[X5]:((ssList(X5)&X3!=X5)&?[X6]:(((ssItem(X6)&cons(X6,nil)=X5)&hd(X4)=X6)&neq(nil,X4))))|?[X7]:((ssList(X7)&X1=X7)&?[X8]:(((ssItem(X8)&cons(X8,nil)=X7)&hd(X2)=X8)&neq(nil,X2))))&(~(neq(X2,nil))|neq(X4,nil))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.05/0.29  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.05/0.29  fof(c_0_2, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.05/0.29  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax22)).
 % 0.05/0.29  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 0.05/0.29  fof(c_0_5, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((~neq(X2,nil)|?[X5]:((ssList(X5)&X3!=X5)&?[X6]:(((ssItem(X6)&cons(X6,nil)=X5)&hd(X4)=X6)&neq(nil,X4))))|?[X7]:((ssList(X7)&X1=X7)&?[X8]:(((ssItem(X8)&cons(X8,nil)=X7)&hd(X2)=X8)&neq(nil,X2))))&(~neq(X2,nil)|neq(X4,nil)))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.05/0.29  fof(c_0_6, negated_conjecture, (![X13]:![X14]:![X15]:![X16]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((esk49_0=esk51_0&esk48_0=esk50_0)&((((neq(esk49_0,nil)|neq(esk49_0,nil))&(~neq(esk51_0,nil)|neq(esk49_0,nil)))&((neq(esk49_0,nil)|((~ssList(X13)|esk50_0=X13)|(((~ssItem(X14)|cons(X14,nil)!=X13)|hd(esk51_0)!=X14)|~neq(nil,esk51_0))))&(~neq(esk51_0,nil)|((~ssList(X13)|esk50_0=X13)|(((~ssItem(X14)|cons(X14,nil)!=X13)|hd(esk51_0)!=X14)|~neq(nil,esk51_0))))))&((neq(esk49_0,nil)|((~ssList(X15)|esk48_0!=X15)|(((~ssItem(X16)|cons(X16,nil)!=X15)|hd(esk49_0)!=X16)|~neq(nil,esk49_0))))&(~neq(esk51_0,nil)|((~ssList(X15)|esk48_0!=X15)|(((~ssItem(X16)|cons(X16,nil)!=X15)|hd(esk49_0)!=X16)|~neq(nil,esk49_0)))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])])).
 % 0.05/0.29  cnf(c_0_7, negated_conjecture, (neq(esk49_0,nil)|neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.05/0.29  cnf(c_0_8, negated_conjecture, (~neq(nil,esk49_0)|hd(esk49_0)!=X1|cons(X1,nil)!=X2|~ssItem(X1)|esk48_0!=X2|~ssList(X2)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.05/0.29  cnf(c_0_9, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.05/0.29  cnf(c_0_10, negated_conjecture, (neq(esk49_0,nil)), inference(cn,[status(thm)],[c_0_7])).
 % 0.05/0.29  cnf(c_0_11, negated_conjecture, (esk50_0=X2|~neq(nil,esk51_0)|hd(esk51_0)!=X1|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.05/0.29  cnf(c_0_12, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.05/0.29  cnf(c_0_13, negated_conjecture, (cons(hd(esk49_0),nil)!=esk48_0|~ssList(cons(hd(esk49_0),nil))|~neq(nil,esk49_0)|~ssItem(hd(esk49_0))), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_9]), c_0_10])])])])).
 % 0.05/0.29  fof(c_0_14, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
 % 0.05/0.29  cnf(c_0_15, negated_conjecture, (~ssList(cons(hd(esk49_0),nil))|~neq(nil,esk49_0)|~ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_12]), c_0_9]), c_0_9]), c_0_9]), c_0_10])])])]), c_0_13])).
 % 0.05/0.29  cnf(c_0_16, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
 % 0.05/0.29  cnf(c_0_17, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.05/0.29  fof(c_0_18, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
 % 0.05/0.29  cnf(c_0_19, negated_conjecture, (~neq(nil,esk49_0)|~ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])])).
 % 0.05/0.29  cnf(c_0_20, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
 % 0.05/0.29  cnf(c_0_21, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.05/0.29  fof(c_0_22, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
 % 0.05/0.29  cnf(c_0_23, negated_conjecture, (esk49_0=nil|~neq(nil,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
 % 0.05/0.29  cnf(c_0_24, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.05/0.29  cnf(c_0_25, plain, (~ssList(X1)|~ssList(X2)|X1!=X2|~neq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.05/0.29  cnf(c_0_26, negated_conjecture, (esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_21]), c_0_17])])).
 % 0.05/0.29  cnf(c_0_27, plain, (~ssList(X1)|~neq(X1,X1)), inference(er,[status(thm)],[c_0_25])).
 % 0.05/0.29  cnf(c_0_28, negated_conjecture, (neq(nil,nil)), inference(rw,[status(thm)],[c_0_10, c_0_26])).
 % 0.05/0.29  cnf(c_0_29, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_17])]), ['proof']).
 % 0.05/0.29  # SZS output end CNFRefutation
 % 0.05/0.29  # Training examples: 0 positive, 0 negative
 % 0.05/0.29
 % 0.05/0.29  # -------------------------------------------------
 % 0.05/0.29  # User time                : 0.026 s
 % 0.05/0.29  # System time              : 0.004 s
 % 0.05/0.29  # Total time               : 0.030 s
 % 0.05/0.29  # Maximum resident set size: 1876 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
