   --------------------------------------------------------------------------

                   TSTP Solution File: SWC046+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC046+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n161.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:25 EDT 2016

 % Result     : Theorem 0.03s
 % Output     : CNFRefutation 0.03s
 % Verified   :
 % Statistics : Number of formulae       :   33 (  73 expanded)
 %              Number of clauses        :   19 (  27 expanded)
 %              Number of leaves         :    6 (  17 expanded)
 %              Depth                    :    7
 %              Number of atoms          :  140 ( 726 expanded)
 %              Number of equality atoms :   35 ( 163 expanded)
 %              Maximal formula depth    :   22 (   5 average)
 %              Maximal clause size      :   23 (   2 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( nil != X2
                     | X2 != X4
                     | X1 != X3
                     | nil = X1
                     | ? [X5] :
                         ( ssItem(X5)
                         & ( ( ~ memberP(X3,X5)
                             & ! [X6] :
                                 ( ssList(X6)
                                => ~ segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) )
                             & memberP(X4,X5) )
                           | ( memberP(X3,X5)
                             & ( ~ memberP(X4,X5)
                               | ? [X6] :
                                   ( ssList(X6)
                                   & segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) ) ) ) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ~ memberP(nil,X1) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax38)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( memberP(cons(X2,X3),X1)
               <=> ( X1 = X2
                   | memberP(X3,X1) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax37)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => cons(hd(X1),tl(X1)) = X1 ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax78)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssItem(hd(X1)) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax22)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssList(tl(X1)) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax24)).

 fof(c_0_6,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( nil != X2
                       | X2 != X4
                       | X1 != X3
                       | nil = X1
                       | ? [X5] :
                           ( ssItem(X5)
                           & ( ( ~ memberP(X3,X5)
                               & ! [X6] :
                                   ( ssList(X6)
                                  => ~ segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) )
                               & memberP(X4,X5) )
                             | ( memberP(X3,X5)
                               & ( ~ memberP(X4,X5)
                                 | ? [X6] :
                                     ( ssList(X6)
                                     & segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) ) ) ) ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_7,negated_conjecture,(
     ! [X11,X13] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & nil = esk49_0
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & nil != esk48_0
       & ( ssList(esk52_1(X11))
         | memberP(esk50_0,X11)
         | ~ memberP(esk51_0,X11)
         | ~ ssItem(X11) )
       & ( segmentP(esk51_0,app(app(cons(X11,nil),esk52_1(X11)),cons(X11,nil)))
         | memberP(esk50_0,X11)
         | ~ memberP(esk51_0,X11)
         | ~ ssItem(X11) )
       & ( memberP(esk51_0,X11)
         | ~ memberP(esk50_0,X11)
         | ~ ssItem(X11) )
       & ( ~ ssList(X13)
         | ~ segmentP(esk51_0,app(app(cons(X11,nil),X13),cons(X11,nil)))
         | ~ memberP(esk50_0,X11)
         | ~ ssItem(X11) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).

 fof(c_0_8,plain,(
     ! [X1] :
       ( ssItem(X1)
      => ~ memberP(nil,X1) ) ),
     inference(fof_simplification,[status(thm)],[[4]c_0_1])).

 fof(c_0_9,plain,(
     ! [X4,X5,X6] :
       ( ( ~ memberP(cons(X5,X6),X4)
         | X4 = X5
         | memberP(X6,X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) )
       & ( X4 != X5
         | memberP(cons(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) )
       & ( ~ memberP(X6,X4)
         | memberP(cons(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).

 cnf(c_0_10,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_7])).

 cnf(c_0_11,negated_conjecture,
     ( nil = esk49_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_7])).

 fof(c_0_12,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | ~ memberP(nil,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).

 cnf(c_0_13,plain,
     ( memberP(cons(X2,X3),X1)
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | X1 != X2 ),
     inference(split_conjunct,[status(thm)],[[7]c_0_9])).

 fof(c_0_14,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | cons(hd(X2),tl(X2)) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

 fof(c_0_15,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssItem(hd(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).

 fof(c_0_16,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssList(tl(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).

 cnf(c_0_17,negated_conjecture,
     ( memberP(esk51_0,X1)
     | ~ ssItem(X1)
     | ~ memberP(esk50_0,X1) ),
     inference(split_conjunct,[status(thm)],[[8]c_0_7])).

 cnf(c_0_18,negated_conjecture,
     ( esk51_0 = nil ),
     inference(rw,[status(thm)],[[9]c_0_10,[10]c_0_11])).

 cnf(c_0_19,plain,
     ( ~ memberP(nil,X1)
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_12])).

 cnf(c_0_20,plain,
     ( memberP(cons(X1,X2),X1)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(er,[status(thm)],[[12]c_0_13])).

 cnf(c_0_21,plain,
     ( cons(hd(X1),tl(X1)) = X1
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[13]c_0_14])).

 cnf(c_0_22,plain,
     ( ssItem(hd(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_15])).

 cnf(c_0_23,plain,
     ( ssList(tl(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_16])).

 cnf(c_0_24,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_7])).

 cnf(c_0_25,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[17]c_0_7])).

 cnf(c_0_26,negated_conjecture,
     ( nil != esk48_0 ),
     inference(split_conjunct,[status(thm)],[[18]c_0_7])).

 cnf(c_0_27,negated_conjecture,
     ( ~ memberP(esk50_0,X1)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(rw,[status(thm)],[[19]c_0_17,c_0_18]),[20]c_0_19])).

 cnf(c_0_28,plain,
     ( nil = X1
     | memberP(X1,hd(X1))
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[21]c_0_20,c_0_21]),c_0_22]),[22]c_0_23])).

 cnf(c_0_29,negated_conjecture,
     ( ssList(esk50_0) ),
     inference(rw,[status(thm)],[[23]c_0_24,[24]c_0_25])).

 cnf(c_0_30,negated_conjecture,
     ( esk50_0 != nil ),
     inference(rw,[status(thm)],[[25]c_0_26,[26]c_0_25])).

 cnf(c_0_31,negated_conjecture,
     ( ~ ssItem(hd(esk50_0)) ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_27,c_0_28]),c_0_29])]),[28]c_0_30])).

 cnf(c_0_32,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[29]c_0_31,c_0_22]),c_0_29])]),[30]c_0_30]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC046+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.22  % Computer   : n161.star.cs.uiowa.edu
 % 0.03/0.22  % Model      : x86_64 x86_64
 % 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.22  % Memory     : 32218.75MB
 % 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.22  % CPULimit   : 300
 % 0.03/0.22  % DateTime   : Tue May 10 20:34:53 CDT 2016
 % 0.03/0.23  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.03/0.28  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.03/0.28  # and selection function SelectCQPrecWNTNp.
 % 0.03/0.28  #
 % 0.03/0.28  # Preprocessing time       : 0.014 s
 % 0.03/0.28  # Presaturation interreduction done
 % 0.03/0.28
 % 0.03/0.28  # Proof found!
 % 0.03/0.28  # SZS status Theorem
 % 0.03/0.28  # SZS output start CNFRefutation
 % 0.03/0.28  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((nil!=X2|X2!=X4)|X1!=X3)|nil=X1)|?[X5]:(ssItem(X5)&(((~(memberP(X3,X5))&![X6]:(ssList(X6)=>~(segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))))))&memberP(X4,X5))|(memberP(X3,X5)&(~(memberP(X4,X5))|?[X6]:(ssList(X6)&segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil)))))))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
 % 0.03/0.28  fof(c_0_1, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax38)).
 % 0.03/0.28  fof(c_0_2, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)<=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax37)).
 % 0.03/0.28  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax78)).
 % 0.03/0.28  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax22)).
 % 0.03/0.28  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax24)).
 % 0.03/0.28  fof(c_0_6, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((nil!=X2|X2!=X4)|X1!=X3)|nil=X1)|?[X5]:(ssItem(X5)&(((~memberP(X3,X5)&![X6]:(ssList(X6)=>~segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil)))))&memberP(X4,X5))|(memberP(X3,X5)&(~memberP(X4,X5)|?[X6]:(ssList(X6)&segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))))))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.03/0.28  fof(c_0_7, negated_conjecture, (![X11]:![X13]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((nil=esk49_0&esk49_0=esk51_0)&esk48_0=esk50_0)&nil!=esk48_0)&(((((ssList(esk52_1(X11))|memberP(esk50_0,X11))|~memberP(esk51_0,X11))|~ssItem(X11))&(((segmentP(esk51_0,app(app(cons(X11,nil),esk52_1(X11)),cons(X11,nil)))|memberP(esk50_0,X11))|~memberP(esk51_0,X11))|~ssItem(X11)))&(((memberP(esk51_0,X11)|~memberP(esk50_0,X11))|~ssItem(X11))&(((~ssList(X13)|~segmentP(esk51_0,app(app(cons(X11,nil),X13),cons(X11,nil))))|~memberP(esk50_0,X11))|~ssItem(X11)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
 % 0.03/0.28  fof(c_0_8, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_1])).
 % 0.03/0.28  fof(c_0_9, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
 % 0.03/0.28  cnf(c_0_10, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.03/0.28  cnf(c_0_11, negated_conjecture, (nil=esk49_0), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.03/0.28  fof(c_0_12, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
 % 0.03/0.28  cnf(c_0_13, plain, (memberP(cons(X2,X3),X1)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.03/0.28  fof(c_0_14, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
 % 0.03/0.28  fof(c_0_15, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
 % 0.03/0.28  fof(c_0_16, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
 % 0.03/0.28  cnf(c_0_17, negated_conjecture, (memberP(esk51_0,X1)|~ssItem(X1)|~memberP(esk50_0,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.03/0.28  cnf(c_0_18, negated_conjecture, (esk51_0=nil), inference(rw,[status(thm)],[c_0_10, c_0_11])).
 % 0.03/0.28  cnf(c_0_19, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
 % 0.03/0.28  cnf(c_0_20, plain, (memberP(cons(X1,X2),X1)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_13])).
 % 0.03/0.28  cnf(c_0_21, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
 % 0.03/0.28  cnf(c_0_22, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
 % 0.03/0.28  cnf(c_0_23, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
 % 0.03/0.28  cnf(c_0_24, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.03/0.28  cnf(c_0_25, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.03/0.28  cnf(c_0_26, negated_conjecture, (nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.03/0.28  cnf(c_0_27, negated_conjecture, (~memberP(esk50_0,X1)|~ssItem(X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
 % 0.03/0.28  cnf(c_0_28, plain, (nil=X1|memberP(X1,hd(X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])).
 % 0.03/0.28  cnf(c_0_29, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_24, c_0_25])).
 % 0.03/0.28  cnf(c_0_30, negated_conjecture, (esk50_0!=nil), inference(rw,[status(thm)],[c_0_26, c_0_25])).
 % 0.03/0.28  cnf(c_0_31, negated_conjecture, (~ssItem(hd(esk50_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])]), c_0_30])).
 % 0.03/0.28  cnf(c_0_32, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_22]), c_0_29])]), c_0_30]), ['proof']).
 % 0.03/0.28  # SZS output end CNFRefutation
 % 0.03/0.28  # Training examples: 0 positive, 0 negative
 % 0.03/0.28
 % 0.03/0.28  # -------------------------------------------------
 % 0.03/0.28  # User time                : 0.051 s
 % 0.03/0.28  # System time              : 0.006 s
 % 0.03/0.28  # Total time               : 0.057 s
 % 0.03/0.28  # Maximum resident set size: 1872 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
