   --------------------------------------------------------------------------

                   TSTP Solution File: SWC002-1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC002-1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n154.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:15 EDT 2016

 % Result     : Unsatisfiable 0.06s
 % Output     : CNFRefutation 0.06s
 % Verified   :
 % Statistics : Number of clauses        :   50 ( 377 expanded)
 %              Number of leaves         :   21 ( 219 expanded)
 %              Depth                    :   15
 %              Number of atoms          :  142 ( 669 expanded)
 %              Number of equality atoms :   39 ( 138 expanded)
 %              Maximal clause size      :    7 (   3 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 cnf(c_0_0,negated_conjecture,
     ( neq(sk2,nil)
     | neq(sk2,nil) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

 cnf(c_0_1,negated_conjecture,
     ( app(app(sk7,cons(sk6,nil)),sk8) = sk4
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_26)).

 cnf(c_0_2,negated_conjecture,
     ( sk2 = sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

 cnf(c_0_3,negated_conjecture,
     ( ssList(sk7)
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_24)).

 cnf(c_0_4,negated_conjecture,
     ( ssList(sk8)
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_25)).

 cnf(c_0_5,axiom,
     ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

 cnf(c_0_6,negated_conjecture,
     ( ssItem(sk6)
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_23)).

 cnf(c_0_7,axiom,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

 cnf(c_0_8,axiom,
     ( ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

 cnf(c_0_9,axiom,
     ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2)
     | ~ duplicatefreeP(X5)
     | ~ ssList(X5) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

 cnf(c_0_10,axiom,
     ( app(cons(X1,nil),X2) = cons(X1,X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

 cnf(c_0_11,axiom,
     ( duplicatefreeP(X1)
     | ssItem(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

 cnf(c_0_12,axiom,
     ( ssList(app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

 cnf(c_0_13,negated_conjecture,
     ( ssList(sk2) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

 cnf(c_0_14,negated_conjecture,
     ( geq(sk5(X3,X2,X1),X1)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | app(app(X2,cons(X1,nil)),X3) != sk2
     | app(X2,X3) != sk1
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_21)).

 cnf(c_0_15,axiom,
     ( ssList(skaf42(X1,X2)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause53)).

 cnf(c_0_16,negated_conjecture,
     ( memberP(sk2,sk5(X3,X2,X1))
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | app(app(X2,cons(X1,nil)),X3) != sk2
     | app(X2,X3) != sk1
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_20)).

 cnf(c_0_17,negated_conjecture,
     ( app(sk7,sk8) = sk3
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_27)).

 cnf(c_0_18,negated_conjecture,
     ( sk1 = sk3 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

 cnf(c_0_19,negated_conjecture,
     ( sk6 = X1
     | ~ ssItem(X1)
     | ~ memberP(sk4,X1)
     | ~ geq(X1,sk6)
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_28)).

 cnf(c_0_20,negated_conjecture,
     ( ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | app(app(X2,cons(X1,nil)),X3) != sk2
     | app(X2,X3) != sk1
     | X1 != sk5(X3,X2,X1)
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_22)).

 cnf(c_0_21,negated_conjecture,
     ( neq(sk2,nil) ),
     inference(cn,[status(thm)],[[4]c_0_0])).

 cnf(c_0_22,negated_conjecture,
     ( app(app(sk7,cons(sk6,nil)),sk8) = sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[5]c_0_1,c_0_2]),c_0_2]),c_0_21])])).

 cnf(c_0_23,negated_conjecture,
     ( ssList(sk7) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[6]c_0_3,c_0_2]),c_0_21])])).

 cnf(c_0_24,negated_conjecture,
     ( ssList(sk8) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[7]c_0_4,c_0_2]),c_0_21])])).

 cnf(c_0_25,negated_conjecture,
     ( app(sk7,app(cons(sk6,nil),sk8)) = sk2
     | ~ ssList(cons(sk6,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[8]c_0_22,c_0_5]),c_0_23]),c_0_24])])).

 cnf(c_0_26,negated_conjecture,
     ( ssItem(sk6) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[9]c_0_6,c_0_2]),c_0_21])])).

 cnf(c_0_27,negated_conjecture,
     ( app(sk7,app(cons(sk6,nil),sk8)) = sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[10]c_0_25,c_0_7]),c_0_26]),c_0_8])])).

 cnf(c_0_28,plain,
     ( ~ ssItem(X1)
     | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
     inference(er,[status(thm)],[[11]c_0_9])).

 cnf(c_0_29,negated_conjecture,
     ( app(sk7,cons(sk6,sk8)) = sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[12]c_0_27,c_0_10]),c_0_26]),c_0_24])])).

 cnf(c_0_30,negated_conjecture,
     ( ~ ssList(app(sk2,cons(sk6,X1)))
     | ~ ssList(X1)
     | ~ duplicatefreeP(app(sk2,cons(sk6,X1))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[13]c_0_28,c_0_29]),c_0_26]),c_0_24]),c_0_23])])).

 cnf(c_0_31,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(app(sk2,cons(sk6,X2)))
     | ~ ssList(X2) ),
     inference(spm,[status(thm)],[[14]c_0_30,[15]c_0_11])).

 cnf(c_0_32,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(cons(sk6,X2))
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[16]c_0_31,c_0_12]),c_0_13])])).

 cnf(c_0_33,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_32,c_0_7]),c_0_26])])).

 cnf(c_0_34,negated_conjecture,
     ( geq(sk5(X1,X2,X3),X3)
     | app(app(X2,cons(X3,nil)),X1) != sk2
     | app(X2,X1) != sk1
     | ~ ssItem(X3)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[18]c_0_14,c_0_2]),c_0_21])])).

 cnf(c_0_35,negated_conjecture,
     ( ssItem(X1) ),
     inference(spm,[status(thm)],[[19]c_0_33,[20]c_0_15])).

 cnf(c_0_36,negated_conjecture,
     ( memberP(sk2,sk5(X1,X2,X3))
     | app(app(X2,cons(X3,nil)),X1) != sk2
     | app(X2,X1) != sk1
     | ~ ssItem(X3)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[21]c_0_16,c_0_2]),c_0_21])])).

 cnf(c_0_37,negated_conjecture,
     ( geq(sk5(X1,X2,X3),X3)
     | app(app(X2,cons(X3,nil)),X1) != sk2
     | app(X2,X1) != sk1
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[22]c_0_34,c_0_35])])).

 cnf(c_0_38,negated_conjecture,
     ( app(sk7,sk8) = sk1 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[23]c_0_17,c_0_18]),c_0_2]),c_0_21])])).

 cnf(c_0_39,negated_conjecture,
     ( memberP(sk2,sk5(X1,X2,X3))
     | app(app(X2,cons(X3,nil)),X1) != sk2
     | app(X2,X1) != sk1
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[24]c_0_36,c_0_35])])).

 cnf(c_0_40,negated_conjecture,
     ( sk6 = X1
     | ~ memberP(sk2,X1)
     | ~ geq(X1,sk6)
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[25]c_0_19,c_0_2]),c_0_2]),c_0_21])])).

 cnf(c_0_41,negated_conjecture,
     ( geq(sk5(sk8,sk7,X1),X1)
     | app(app(sk7,cons(X1,nil)),sk8) != sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_37,c_0_38]),c_0_24]),c_0_23])])).

 cnf(c_0_42,negated_conjecture,
     ( memberP(sk2,sk5(sk8,sk7,X1))
     | app(app(sk7,cons(X1,nil)),sk8) != sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_39,c_0_38]),c_0_24]),c_0_23])])).

 cnf(c_0_43,negated_conjecture,
     ( app(app(X1,cons(X2,nil)),X3) != sk2
     | app(X1,X3) != sk1
     | sk5(X3,X1,X2) != X2
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[28]c_0_20,c_0_2]),c_0_21])])).

 cnf(c_0_44,negated_conjecture,
     ( sk6 = X1
     | ~ memberP(sk2,X1)
     | ~ geq(X1,sk6) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[29]c_0_40,c_0_35])])).

 cnf(c_0_45,negated_conjecture,
     ( geq(sk5(sk8,sk7,sk6),sk6) ),
     inference(spm,[status(thm)],[[30]c_0_41,[31]c_0_22])).

 cnf(c_0_46,negated_conjecture,
     ( memberP(sk2,sk5(sk8,sk7,sk6)) ),
     inference(spm,[status(thm)],[[32]c_0_42,[33]c_0_22])).

 cnf(c_0_47,negated_conjecture,
     ( app(app(X1,cons(X2,nil)),X3) != sk2
     | app(X1,X3) != sk1
     | sk5(X3,X1,X2) != X2
     | ~ ssList(X3)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[34]c_0_43,c_0_35])])).

 cnf(c_0_48,negated_conjecture,
     ( sk5(sk8,sk7,sk6) = sk6 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[35]c_0_44,c_0_45]),c_0_46])])).

 cnf(c_0_49,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[36]c_0_47,c_0_48]),c_0_22]),c_0_38]),c_0_24]),c_0_23])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC002-1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.23  % Computer   : n154.star.cs.uiowa.edu
 % 0.02/0.23  % Model      : x86_64 x86_64
 % 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.23  % Memory     : 32218.75MB
 % 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.23  % CPULimit   : 300
 % 0.02/0.23  % DateTime   : Tue May 10 20:34:23 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.06/0.80  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4d
 % 0.06/0.80  # and selection function SelectCQIPrecWNTNp.
 % 0.06/0.80  #
 % 0.06/0.80  # Preprocessing time       : 0.010 s
 % 0.06/0.80  # Presaturation interreduction done
 % 0.06/0.80
 % 0.06/0.80  # Proof found!
 % 0.06/0.80  # SZS status Unsatisfiable
 % 0.06/0.80  # SZS output start CNFRefutation
 % 0.06/0.80  cnf(c_0_0, negated_conjecture, (neq(sk2,nil)|neq(sk2,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
 % 0.06/0.80  cnf(c_0_1, negated_conjecture, (app(app(sk7,cons(sk6,nil)),sk8)=sk4|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_26)).
 % 0.06/0.80  cnf(c_0_2, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
 % 0.06/0.80  cnf(c_0_3, negated_conjecture, (ssList(sk7)|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_24)).
 % 0.06/0.80  cnf(c_0_4, negated_conjecture, (ssList(sk8)|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_25)).
 % 0.06/0.80  cnf(c_0_5, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
 % 0.06/0.80  cnf(c_0_6, negated_conjecture, (ssItem(sk6)|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_23)).
 % 0.06/0.80  cnf(c_0_7, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
 % 0.06/0.80  cnf(c_0_8, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
 % 0.06/0.80  cnf(c_0_9, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
 % 0.06/0.80  cnf(c_0_10, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
 % 0.06/0.80  cnf(c_0_11, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
 % 0.06/0.80  cnf(c_0_12, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
 % 0.06/0.80  cnf(c_0_13, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
 % 0.06/0.80  cnf(c_0_14, negated_conjecture, (geq(sk5(X3,X2,X1),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|app(app(X2,cons(X1,nil)),X3)!=sk2|app(X2,X3)!=sk1|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_21)).
 % 0.06/0.80  cnf(c_0_15, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause53)).
 % 0.06/0.80  cnf(c_0_16, negated_conjecture, (memberP(sk2,sk5(X3,X2,X1))|~ssItem(X1)|~ssList(X2)|~ssList(X3)|app(app(X2,cons(X1,nil)),X3)!=sk2|app(X2,X3)!=sk1|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_20)).
 % 0.06/0.80  cnf(c_0_17, negated_conjecture, (app(sk7,sk8)=sk3|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_27)).
 % 0.06/0.80  cnf(c_0_18, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
 % 0.06/0.80  cnf(c_0_19, negated_conjecture, (sk6=X1|~ssItem(X1)|~memberP(sk4,X1)|~geq(X1,sk6)|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_28)).
 % 0.06/0.80  cnf(c_0_20, negated_conjecture, (~ssItem(X1)|~ssList(X2)|~ssList(X3)|app(app(X2,cons(X1,nil)),X3)!=sk2|app(X2,X3)!=sk1|X1!=sk5(X3,X2,X1)|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_22)).
 % 0.06/0.80  cnf(c_0_21, negated_conjecture, (neq(sk2,nil)), inference(cn,[status(thm)],[c_0_0])).
 % 0.06/0.80  cnf(c_0_22, negated_conjecture, (app(app(sk7,cons(sk6,nil)),sk8)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_1, c_0_2]), c_0_2]), c_0_21])])).
 % 0.06/0.80  cnf(c_0_23, negated_conjecture, (ssList(sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_2]), c_0_21])])).
 % 0.06/0.80  cnf(c_0_24, negated_conjecture, (ssList(sk8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_2]), c_0_21])])).
 % 0.06/0.80  cnf(c_0_25, negated_conjecture, (app(sk7,app(cons(sk6,nil),sk8))=sk2|~ssList(cons(sk6,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_5]), c_0_23]), c_0_24])])).
 % 0.06/0.80  cnf(c_0_26, negated_conjecture, (ssItem(sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_2]), c_0_21])])).
 % 0.06/0.80  cnf(c_0_27, negated_conjecture, (app(sk7,app(cons(sk6,nil),sk8))=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_7]), c_0_26]), c_0_8])])).
 % 0.06/0.80  cnf(c_0_28, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_9])).
 % 0.06/0.80  cnf(c_0_29, negated_conjecture, (app(sk7,cons(sk6,sk8))=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_10]), c_0_26]), c_0_24])])).
 % 0.06/0.80  cnf(c_0_30, negated_conjecture, (~ssList(app(sk2,cons(sk6,X1)))|~ssList(X1)|~duplicatefreeP(app(sk2,cons(sk6,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_26]), c_0_24]), c_0_23])])).
 % 0.06/0.80  cnf(c_0_31, negated_conjecture, (ssItem(X1)|~ssList(app(sk2,cons(sk6,X2)))|~ssList(X2)), inference(spm,[status(thm)],[c_0_30, c_0_11])).
 % 0.06/0.80  cnf(c_0_32, negated_conjecture, (ssItem(X1)|~ssList(cons(sk6,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_12]), c_0_13])])).
 % 0.06/0.80  cnf(c_0_33, negated_conjecture, (ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_7]), c_0_26])])).
 % 0.06/0.80  cnf(c_0_34, negated_conjecture, (geq(sk5(X1,X2,X3),X3)|app(app(X2,cons(X3,nil)),X1)!=sk2|app(X2,X1)!=sk1|~ssItem(X3)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_2]), c_0_21])])).
 % 0.06/0.80  cnf(c_0_35, negated_conjecture, (ssItem(X1)), inference(spm,[status(thm)],[c_0_33, c_0_15])).
 % 0.06/0.80  cnf(c_0_36, negated_conjecture, (memberP(sk2,sk5(X1,X2,X3))|app(app(X2,cons(X3,nil)),X1)!=sk2|app(X2,X1)!=sk1|~ssItem(X3)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_2]), c_0_21])])).
 % 0.06/0.80  cnf(c_0_37, negated_conjecture, (geq(sk5(X1,X2,X3),X3)|app(app(X2,cons(X3,nil)),X1)!=sk2|app(X2,X1)!=sk1|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35])])).
 % 0.06/0.80  cnf(c_0_38, negated_conjecture, (app(sk7,sk8)=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_2]), c_0_21])])).
 % 0.06/0.80  cnf(c_0_39, negated_conjecture, (memberP(sk2,sk5(X1,X2,X3))|app(app(X2,cons(X3,nil)),X1)!=sk2|app(X2,X1)!=sk1|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_35])])).
 % 0.06/0.80  cnf(c_0_40, negated_conjecture, (sk6=X1|~memberP(sk2,X1)|~geq(X1,sk6)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_2]), c_0_2]), c_0_21])])).
 % 0.06/0.80  cnf(c_0_41, negated_conjecture, (geq(sk5(sk8,sk7,X1),X1)|app(app(sk7,cons(X1,nil)),sk8)!=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_24]), c_0_23])])).
 % 0.06/0.80  cnf(c_0_42, negated_conjecture, (memberP(sk2,sk5(sk8,sk7,X1))|app(app(sk7,cons(X1,nil)),sk8)!=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_38]), c_0_24]), c_0_23])])).
 % 0.06/0.80  cnf(c_0_43, negated_conjecture, (app(app(X1,cons(X2,nil)),X3)!=sk2|app(X1,X3)!=sk1|sk5(X3,X1,X2)!=X2|~ssItem(X2)|~ssList(X3)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_2]), c_0_21])])).
 % 0.06/0.80  cnf(c_0_44, negated_conjecture, (sk6=X1|~memberP(sk2,X1)|~geq(X1,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_35])])).
 % 0.06/0.80  cnf(c_0_45, negated_conjecture, (geq(sk5(sk8,sk7,sk6),sk6)), inference(spm,[status(thm)],[c_0_41, c_0_22])).
 % 0.06/0.80  cnf(c_0_46, negated_conjecture, (memberP(sk2,sk5(sk8,sk7,sk6))), inference(spm,[status(thm)],[c_0_42, c_0_22])).
 % 0.06/0.80  cnf(c_0_47, negated_conjecture, (app(app(X1,cons(X2,nil)),X3)!=sk2|app(X1,X3)!=sk1|sk5(X3,X1,X2)!=X2|~ssList(X3)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_35])])).
 % 0.06/0.80  cnf(c_0_48, negated_conjecture, (sk5(sk8,sk7,sk6)=sk6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])])).
 % 0.06/0.80  cnf(c_0_49, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_22]), c_0_38]), c_0_24]), c_0_23])]), ['proof']).
 % 0.06/0.80  # SZS output end CNFRefutation
 % 0.06/0.80  # Training examples: 0 positive, 0 negative
 % 0.06/0.81
 % 0.06/0.81  # -------------------------------------------------
 % 0.06/0.81  # User time                : 0.565 s
 % 0.06/0.81  # System time              : 0.010 s
 % 0.06/0.81  # Total time               : 0.576 s
 % 0.06/0.81  # Maximum resident set size: 1872 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
