   --------------------------------------------------------------------------

                   TSTP Solution File: SWC083+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC083+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n153.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:35 EDT 2016

 % Result     : Theorem 0.06s
 % Output     : CNFRefutation 0.06s
 % Verified   :
 % Statistics : Number of formulae       :  115 (10809 expanded)
 %              Number of clauses        :   83 (3968 expanded)
 %              Number of leaves         :   16 (2627 expanded)
 %              Depth                    :   25
 %              Number of atoms          :  370 (97365 expanded)
 %              Number of equality atoms :   97 (23052 expanded)
 %              Maximal formula depth    :   22 (   4 average)
 %              Maximal clause size      :   18 (   2 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ~ ssList(X4)
                   | X2 != X4
                   | X1 != X3
                   | ~ neq(X2,nil)
                   | ? [X5] :
                       ( ssList(X5)
                       & neq(X5,nil)
                       & segmentP(X2,X5)
                       & segmentP(X1,X5) )
                   | ! [X6] :
                       ( ssList(X6)
                      => ! [X7] :
                           ( ssList(X7)
                          => ! [X8] :
                               ( ~ ssList(X8)
                               | app(app(X6,X7),X8) != X3
                               | app(X6,X8) != X4 ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( rearsegP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & app(X3,X2) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax6)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( app(X3,X2) = app(X1,X2)
                => X3 = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax79)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax82)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( segmentP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & ? [X4] :
                     ( ssList(X4)
                     & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax7)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(X1,nil) = X1 ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax84)).

 fof(c_0_7,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_8,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_9,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( frontsegP(X1,X2)
                => frontsegP(app(X1,X3),X2) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax43)).

 fof(c_0_10,axiom,(
     ! [X1] :
       ( ssList(X1)
      => frontsegP(X1,nil) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax45)).

 fof(c_0_11,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( frontsegP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & app(X2,X3) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax5)).

 fof(c_0_12,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_13,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( app(X2,X3) = app(X2,X1)
                => X3 = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax80)).

 fof(c_0_14,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(nil,X1) = X1 ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax28)).

 fof(c_0_15,axiom,(
     ! [X1] :
       ( ssList(X1)
      => frontsegP(X1,X1) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax42)).

 fof(c_0_16,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ~ ssList(X4)
                     | X2 != X4
                     | X1 != X3
                     | ~ neq(X2,nil)
                     | ? [X5] :
                         ( ssList(X5)
                         & neq(X5,nil)
                         & segmentP(X2,X5)
                         & segmentP(X1,X5) )
                     | ! [X6] :
                         ( ssList(X6)
                        => ! [X7] :
                             ( ssList(X7)
                            => ! [X8] :
                                 ( ~ ssList(X8)
                                 | app(app(X6,X7),X8) != X3
                                 | app(X6,X8) != X4 ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_17,plain,(
     ! [X4,X5,X7] :
       ( ( ssList(esk7_2(X4,X5))
         | ~ rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( app(esk7_2(X4,X5),X5) = X4
         | ~ rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( ~ ssList(X7)
         | app(X7,X5) != X4
         | rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

 fof(c_0_18,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

 fof(c_0_19,negated_conjecture,(
     ! [X13] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & neq(esk49_0,nil)
       & ( ~ ssList(X13)
         | ~ neq(X13,nil)
         | ~ segmentP(esk49_0,X13)
         | ~ segmentP(esk48_0,X13) )
       & ssList(esk52_0)
       & ssList(esk53_0)
       & ssList(esk54_0)
       & app(app(esk52_0,esk53_0),esk54_0) = esk50_0
       & app(esk52_0,esk54_0) = esk51_0 ) ),
     inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])).

 cnf(c_0_20,plain,
     ( rearsegP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(X3,X2) != X1
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_17])).

 cnf(c_0_21,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[5]c_0_18])).

 cnf(c_0_22,negated_conjecture,
     ( app(app(esk52_0,esk53_0),esk54_0) = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_19])).

 cnf(c_0_23,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[7]c_0_19])).

 fof(c_0_24,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(X6,X5) != app(X4,X5)
       | X6 = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

 cnf(c_0_25,plain,
     ( rearsegP(app(X1,X2),X2)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_20]),[8]c_0_21])).

 cnf(c_0_26,negated_conjecture,
     ( app(app(esk52_0,esk53_0),esk54_0) = esk48_0 ),
     inference(rw,[status(thm)],[[9]c_0_22,[10]c_0_23])).

 cnf(c_0_27,negated_conjecture,
     ( ssList(esk54_0) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_19])).

 cnf(c_0_28,plain,
     ( X1 = X2
     | app(X1,X3) != app(X2,X3)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_24])).

 cnf(c_0_29,negated_conjecture,
     ( rearsegP(esk48_0,esk54_0)
     | ~ ssList(app(esk52_0,esk53_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[13]c_0_25,c_0_26]),c_0_27])])).

 cnf(c_0_30,negated_conjecture,
     ( ssList(esk53_0) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_19])).

 cnf(c_0_31,negated_conjecture,
     ( ssList(esk52_0) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_19])).

 cnf(c_0_32,negated_conjecture,
     ( X1 = app(esk52_0,esk53_0)
     | app(X1,esk54_0) != esk48_0
     | ~ ssList(app(esk52_0,esk53_0))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[16]c_0_28,c_0_26]),c_0_27])])).

 cnf(c_0_33,plain,
     ( app(esk7_2(X1,X2),X2) = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ rearsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[17]c_0_17])).

 cnf(c_0_34,negated_conjecture,
     ( rearsegP(esk48_0,esk54_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_29,c_0_21]),c_0_30]),c_0_31])])).

 cnf(c_0_35,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_19])).

 cnf(c_0_36,plain,
     ( ssList(esk7_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ rearsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[20]c_0_17])).

 cnf(c_0_37,negated_conjecture,
     ( X1 = app(esk52_0,esk53_0)
     | app(X1,esk54_0) != esk48_0
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[21]c_0_32,c_0_21]),c_0_30]),c_0_31])])).

 cnf(c_0_38,negated_conjecture,
     ( app(esk7_2(esk48_0,esk54_0),esk54_0) = esk48_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[22]c_0_33,c_0_34]),c_0_27]),c_0_35])])).

 cnf(c_0_39,negated_conjecture,
     ( ssList(esk7_2(esk48_0,esk54_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[23]c_0_36,c_0_34]),c_0_27]),c_0_35])])).

 fof(c_0_40,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

 cnf(c_0_41,negated_conjecture,
     ( esk7_2(esk48_0,esk54_0) = app(esk52_0,esk53_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[24]c_0_37,c_0_38]),c_0_39])])).

 cnf(c_0_42,plain,
     ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[25]c_0_40])).

 cnf(c_0_43,negated_conjecture,
     ( ssList(app(esk52_0,esk53_0)) ),
     inference(rw,[status(thm)],[[26]c_0_39,[27]c_0_41])).

 fof(c_0_44,plain,(
     ! [X5,X6,X9,X10] :
       ( ( ssList(esk8_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ssList(esk9_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ~ ssList(X9)
         | ~ ssList(X10)
         | app(app(X9,X6),X10) != X5
         | segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).

 cnf(c_0_45,negated_conjecture,
     ( app(app(esk52_0,esk53_0),app(esk54_0,X1)) = app(esk48_0,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[28]c_0_42,c_0_26]),c_0_27])]),c_0_43])])).

 cnf(c_0_46,plain,
     ( segmentP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(app(X3,X2),X4) != X1
     | ~ ssList(X4)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[29]c_0_44])).

 cnf(c_0_47,negated_conjecture,
     ( ssList(app(esk48_0,X1))
     | ~ ssList(app(esk54_0,X1))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[30]c_0_21,c_0_45]),c_0_43])])).

 cnf(c_0_48,negated_conjecture,
     ( app(esk52_0,esk54_0) = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[31]c_0_19])).

 cnf(c_0_49,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[32]c_0_19])).

 cnf(c_0_50,plain,
     ( segmentP(app(app(X1,X2),X3),X2)
     | ~ ssList(app(app(X1,X2),X3))
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(er,[status(thm)],[[33]c_0_46])).

 cnf(c_0_51,negated_conjecture,
     ( ssList(app(esk48_0,X1))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[34]c_0_47,c_0_21]),c_0_27])])).

 fof(c_0_52,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(X2,nil) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).

 cnf(c_0_53,negated_conjecture,
     ( app(esk52_0,esk54_0) = esk49_0 ),
     inference(rw,[status(thm)],[[35]c_0_48,[36]c_0_49])).

 cnf(c_0_54,negated_conjecture,
     ( segmentP(app(esk48_0,X1),esk54_0)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[37]c_0_50,c_0_26]),c_0_27])]),c_0_43])]),[38]c_0_51])).

 cnf(c_0_55,plain,
     ( app(X1,nil) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[39]c_0_52])).

 cnf(c_0_56,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[40]c_0_7])).

 cnf(c_0_57,negated_conjecture,
     ( segmentP(app(esk49_0,X1),esk54_0)
     | ~ ssList(app(esk49_0,X1))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[41]c_0_50,c_0_53]),c_0_31]),c_0_27])])).

 cnf(c_0_58,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[42]c_0_19])).

 cnf(c_0_59,negated_conjecture,
     ( ~ segmentP(esk48_0,X1)
     | ~ segmentP(esk49_0,X1)
     | ~ neq(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[43]c_0_19])).

 cnf(c_0_60,negated_conjecture,
     ( segmentP(esk48_0,esk54_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[44]c_0_54,c_0_55]),c_0_56]),c_0_35])])).

 cnf(c_0_61,negated_conjecture,
     ( segmentP(esk49_0,esk54_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[45]c_0_57,c_0_55]),c_0_58]),c_0_56])])).

 fof(c_0_62,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).

 cnf(c_0_63,negated_conjecture,
     ( ~ neq(esk54_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[46]c_0_59,c_0_60]),c_0_61]),c_0_27])])).

 cnf(c_0_64,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[47]c_0_62])).

 cnf(c_0_65,negated_conjecture,
     ( esk54_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[48]c_0_63,c_0_64]),c_0_56]),c_0_27])])).

 fof(c_0_66,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | ~ frontsegP(X4,X5)
       | frontsegP(app(X4,X6),X5) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

 fof(c_0_67,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | frontsegP(X2,nil) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).

 cnf(c_0_68,negated_conjecture,
     ( app(esk52_0,nil) = esk49_0 ),
     inference(rw,[status(thm)],[[49]c_0_53,[50]c_0_65])).

 fof(c_0_69,plain,(
     ! [X4,X5,X7] :
       ( ( ssList(esk6_2(X4,X5))
         | ~ frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( app(X5,esk6_2(X4,X5)) = X4
         | ~ frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( ~ ssList(X7)
         | app(X5,X7) != X4
         | frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).

 fof(c_0_70,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).

 cnf(c_0_71,plain,
     ( frontsegP(app(X1,X2),X3)
     | ~ frontsegP(X1,X3)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[51]c_0_66])).

 cnf(c_0_72,plain,
     ( frontsegP(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[52]c_0_67])).

 cnf(c_0_73,negated_conjecture,
     ( esk52_0 = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[53]c_0_55,c_0_68]),c_0_31])])).

 cnf(c_0_74,negated_conjecture,
     ( app(esk48_0,nil) = esk48_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[54]c_0_45,c_0_55]),c_0_26]),c_0_56]),c_0_27])])).

 cnf(c_0_75,plain,
     ( frontsegP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(X2,X3) != X1
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[55]c_0_69])).

 cnf(c_0_76,plain,
     ( nil = app(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != X1
     | nil != X2 ),
     inference(split_conjunct,[status(thm)],[[56]c_0_70])).

 cnf(c_0_77,plain,
     ( frontsegP(app(X1,X2),nil)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[57]c_0_71,c_0_72]),c_0_56])])).

 cnf(c_0_78,negated_conjecture,
     ( app(esk49_0,nil) = esk49_0 ),
     inference(rw,[status(thm)],[[58]c_0_68,[59]c_0_73])).

 cnf(c_0_79,negated_conjecture,
     ( rearsegP(esk48_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[60]c_0_25,c_0_74]),c_0_35]),c_0_56])])).

 fof(c_0_80,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(X5,X6) != app(X5,X4)
       | X6 = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).

 cnf(c_0_81,plain,
     ( frontsegP(app(X1,X2),X1)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_75]),[61]c_0_21])).

 cnf(c_0_82,plain,
     ( app(nil,nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_76])]),c_0_56])])).

 cnf(c_0_83,plain,
     ( app(X2,esk6_2(X1,X2)) = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ frontsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[62]c_0_69])).

 cnf(c_0_84,negated_conjecture,
     ( frontsegP(esk49_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[63]c_0_77,c_0_78]),c_0_56]),c_0_58])])).

 cnf(c_0_85,plain,
     ( ssList(esk6_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ frontsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[64]c_0_69])).

 cnf(c_0_86,negated_conjecture,
     ( app(esk49_0,app(nil,X1)) = app(esk49_0,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[65]c_0_42,c_0_78]),c_0_56]),c_0_58])])).

 fof(c_0_87,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(nil,X2) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).

 cnf(c_0_88,negated_conjecture,
     ( app(esk7_2(esk48_0,nil),nil) = esk48_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[66]c_0_33,c_0_79]),c_0_56]),c_0_35])])).

 cnf(c_0_89,negated_conjecture,
     ( ssList(esk7_2(esk48_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[67]c_0_36,c_0_79]),c_0_56]),c_0_35])])).

 cnf(c_0_90,plain,
     ( X1 = X2
     | app(X3,X1) != app(X3,X2)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[68]c_0_80])).

 cnf(c_0_91,negated_conjecture,
     ( frontsegP(esk49_0,esk52_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[69]c_0_81,c_0_53]),c_0_27]),c_0_31])])).

 cnf(c_0_92,plain,
     ( app(nil,app(nil,X1)) = app(nil,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[70]c_0_42,c_0_82]),c_0_56])])).

 cnf(c_0_93,negated_conjecture,
     ( app(nil,esk6_2(esk49_0,nil)) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[71]c_0_83,c_0_84]),c_0_56]),c_0_58])])).

 cnf(c_0_94,negated_conjecture,
     ( ssList(esk6_2(esk49_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[72]c_0_85,c_0_84]),c_0_56]),c_0_58])])).

 cnf(c_0_95,negated_conjecture,
     ( ssList(app(esk49_0,X1))
     | ~ ssList(app(nil,X1))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[73]c_0_21,c_0_86]),c_0_58])])).

 cnf(c_0_96,plain,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[74]c_0_87])).

 cnf(c_0_97,negated_conjecture,
     ( esk7_2(esk48_0,nil) = esk48_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[75]c_0_55,c_0_88]),c_0_89])])).

 fof(c_0_98,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | frontsegP(X2,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).

 cnf(c_0_99,negated_conjecture,
     ( X1 = esk54_0
     | app(esk52_0,X1) != esk49_0
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[76]c_0_90,c_0_53]),c_0_31]),c_0_27])])).

 cnf(c_0_100,negated_conjecture,
     ( app(esk52_0,esk6_2(esk49_0,esk52_0)) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[77]c_0_83,c_0_91]),c_0_31]),c_0_58])])).

 cnf(c_0_101,negated_conjecture,
     ( ssList(esk6_2(esk49_0,esk52_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[78]c_0_85,c_0_91]),c_0_31]),c_0_58])])).

 cnf(c_0_102,negated_conjecture,
     ( app(nil,esk49_0) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[79]c_0_92,c_0_93]),c_0_94])])).

 cnf(c_0_103,negated_conjecture,
     ( ssList(app(esk49_0,X1))
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[80]c_0_95,[81]c_0_96])).

 cnf(c_0_104,negated_conjecture,
     ( app(esk52_0,esk53_0) = esk48_0 ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[82]c_0_41,c_0_65]),[83]c_0_97])).

 cnf(c_0_105,plain,
     ( frontsegP(X1,X1)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[84]c_0_98])).

 cnf(c_0_106,negated_conjecture,
     ( esk6_2(esk49_0,esk52_0) = esk54_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[85]c_0_99,c_0_100]),c_0_101])])).

 cnf(c_0_107,negated_conjecture,
     ( segmentP(app(esk49_0,X1),esk49_0)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[86]c_0_50,c_0_102]),c_0_56]),c_0_58])]),[87]c_0_103])).

 cnf(c_0_108,negated_conjecture,
     ( app(esk49_0,esk53_0) = esk48_0 ),
     inference(rw,[status(thm)],[[88]c_0_104,[89]c_0_73])).

 cnf(c_0_109,plain,
     ( app(X1,esk6_2(X1,X1)) = X1
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[90]c_0_83,[91]c_0_105])).

 cnf(c_0_110,negated_conjecture,
     ( esk6_2(esk49_0,esk49_0) = nil ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[92]c_0_106,c_0_65]),[93]c_0_73])).

 cnf(c_0_111,negated_conjecture,
     ( segmentP(esk48_0,esk49_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[94]c_0_107,c_0_108]),c_0_30])])).

 cnf(c_0_112,negated_conjecture,
     ( segmentP(esk49_0,esk49_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[95]c_0_107,c_0_109]),c_0_110]),c_0_56]),c_0_58])])).

 cnf(c_0_113,negated_conjecture,
     ( neq(esk49_0,nil) ),
     inference(split_conjunct,[status(thm)],[[96]c_0_19])).

 cnf(c_0_114,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[97]c_0_59,c_0_111]),c_0_112]),c_0_58]),c_0_113])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC083+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.23  % Computer   : n153.star.cs.uiowa.edu
 % 0.02/0.23  % Model      : x86_64 x86_64
 % 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.23  % Memory     : 32218.75MB
 % 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.23  % CPULimit   : 300
 % 0.02/0.23  % DateTime   : Tue May 10 20:35:08 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.06/0.38  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.06/0.38  # and selection function SelectCQPrecWNTNp.
 % 0.06/0.38  #
 % 0.06/0.38  # Preprocessing time       : 0.014 s
 % 0.06/0.38  # Presaturation interreduction done
 % 0.06/0.38
 % 0.06/0.38  # Proof found!
 % 0.06/0.38  # SZS status Theorem
 % 0.06/0.38  # SZS output start CNFRefutation
 % 0.06/0.38  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~(ssList(X4))|X2!=X4)|X1!=X3)|~(neq(X2,nil)))|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))|![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>![X8]:((~(ssList(X8))|app(app(X6,X7),X8)!=X3)|app(X6,X8)!=X4)))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
 % 0.06/0.38  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)<=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax6)).
 % 0.06/0.38  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 0.06/0.38  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax79)).
 % 0.06/0.38  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax82)).
 % 0.06/0.38  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)<=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax7)).
 % 0.06/0.38  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax84)).
 % 0.06/0.38  fof(c_0_7, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.06/0.38  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 0.06/0.38  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(frontsegP(X1,X2)=>frontsegP(app(X1,X3),X2)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax43)).
 % 0.06/0.38  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax45)).
 % 0.06/0.38  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)<=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax5)).
 % 0.06/0.38  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 0.06/0.38  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X2,X3)=app(X2,X1)=>X3=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax80)).
 % 0.06/0.38  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 0.06/0.38  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax42)).
 % 0.06/0.38  fof(c_0_16, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~ssList(X4)|X2!=X4)|X1!=X3)|~neq(X2,nil))|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5)))|![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>![X8]:((~ssList(X8)|app(app(X6,X7),X8)!=X3)|app(X6,X8)!=X4))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.06/0.38  fof(c_0_17, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
 % 0.06/0.38  fof(c_0_18, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
 % 0.06/0.38  fof(c_0_19, negated_conjecture, (![X13]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&neq(esk49_0,nil))&(((~ssList(X13)|~neq(X13,nil))|~segmentP(esk49_0,X13))|~segmentP(esk48_0,X13)))&(ssList(esk52_0)&(ssList(esk53_0)&((ssList(esk54_0)&app(app(esk52_0,esk53_0),esk54_0)=esk50_0)&app(esk52_0,esk54_0)=esk51_0)))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])).
 % 0.06/0.38  cnf(c_0_20, plain, (rearsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X3,X2)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.06/0.38  cnf(c_0_21, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
 % 0.06/0.38  cnf(c_0_22, negated_conjecture, (app(app(esk52_0,esk53_0),esk54_0)=esk50_0), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.38  cnf(c_0_23, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.38  fof(c_0_24, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
 % 0.06/0.38  cnf(c_0_25, plain, (rearsegP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_20]), c_0_21])).
 % 0.06/0.38  cnf(c_0_26, negated_conjecture, (app(app(esk52_0,esk53_0),esk54_0)=esk48_0), inference(rw,[status(thm)],[c_0_22, c_0_23])).
 % 0.06/0.38  cnf(c_0_27, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.38  cnf(c_0_28, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
 % 0.06/0.38  cnf(c_0_29, negated_conjecture, (rearsegP(esk48_0,esk54_0)|~ssList(app(esk52_0,esk53_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
 % 0.06/0.38  cnf(c_0_30, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.38  cnf(c_0_31, negated_conjecture, (ssList(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.38  cnf(c_0_32, negated_conjecture, (X1=app(esk52_0,esk53_0)|app(X1,esk54_0)!=esk48_0|~ssList(app(esk52_0,esk53_0))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_26]), c_0_27])])).
 % 0.06/0.38  cnf(c_0_33, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.06/0.38  cnf(c_0_34, negated_conjecture, (rearsegP(esk48_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_21]), c_0_30]), c_0_31])])).
 % 0.06/0.38  cnf(c_0_35, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.38  cnf(c_0_36, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.06/0.38  cnf(c_0_37, negated_conjecture, (X1=app(esk52_0,esk53_0)|app(X1,esk54_0)!=esk48_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_21]), c_0_30]), c_0_31])])).
 % 0.06/0.38  cnf(c_0_38, negated_conjecture, (app(esk7_2(esk48_0,esk54_0),esk54_0)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_27]), c_0_35])])).
 % 0.06/0.38  cnf(c_0_39, negated_conjecture, (ssList(esk7_2(esk48_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_34]), c_0_27]), c_0_35])])).
 % 0.06/0.38  fof(c_0_40, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 0.06/0.38  cnf(c_0_41, negated_conjecture, (esk7_2(esk48_0,esk54_0)=app(esk52_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])])).
 % 0.06/0.38  cnf(c_0_42, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
 % 0.06/0.38  cnf(c_0_43, negated_conjecture, (ssList(app(esk52_0,esk53_0))), inference(rw,[status(thm)],[c_0_39, c_0_41])).
 % 0.06/0.38  fof(c_0_44, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
 % 0.06/0.38  cnf(c_0_45, negated_conjecture, (app(app(esk52_0,esk53_0),app(esk54_0,X1))=app(esk48_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_26]), c_0_27])]), c_0_43])])).
 % 0.06/0.38  cnf(c_0_46, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
 % 0.06/0.38  cnf(c_0_47, negated_conjecture, (ssList(app(esk48_0,X1))|~ssList(app(esk54_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_45]), c_0_43])])).
 % 0.06/0.38  cnf(c_0_48, negated_conjecture, (app(esk52_0,esk54_0)=esk51_0), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.38  cnf(c_0_49, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.38  cnf(c_0_50, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_46])).
 % 0.06/0.38  cnf(c_0_51, negated_conjecture, (ssList(app(esk48_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_21]), c_0_27])])).
 % 0.06/0.38  fof(c_0_52, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
 % 0.06/0.38  cnf(c_0_53, negated_conjecture, (app(esk52_0,esk54_0)=esk49_0), inference(rw,[status(thm)],[c_0_48, c_0_49])).
 % 0.06/0.38  cnf(c_0_54, negated_conjecture, (segmentP(app(esk48_0,X1),esk54_0)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_26]), c_0_27])]), c_0_43])]), c_0_51])).
 % 0.06/0.38  cnf(c_0_55, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
 % 0.06/0.38  cnf(c_0_56, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.06/0.38  cnf(c_0_57, negated_conjecture, (segmentP(app(esk49_0,X1),esk54_0)|~ssList(app(esk49_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_53]), c_0_31]), c_0_27])])).
 % 0.06/0.38  cnf(c_0_58, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.38  cnf(c_0_59, negated_conjecture, (~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.38  cnf(c_0_60, negated_conjecture, (segmentP(esk48_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), c_0_35])])).
 % 0.06/0.38  cnf(c_0_61, negated_conjecture, (segmentP(esk49_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_55]), c_0_58]), c_0_56])])).
 % 0.06/0.38  fof(c_0_62, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
 % 0.06/0.38  cnf(c_0_63, negated_conjecture, (~neq(esk54_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61]), c_0_27])])).
 % 0.06/0.38  cnf(c_0_64, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
 % 0.06/0.38  cnf(c_0_65, negated_conjecture, (esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_56]), c_0_27])])).
 % 0.06/0.38  fof(c_0_66, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~frontsegP(X4,X5)|frontsegP(app(X4,X6),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
 % 0.06/0.38  fof(c_0_67, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
 % 0.06/0.38  cnf(c_0_68, negated_conjecture, (app(esk52_0,nil)=esk49_0), inference(rw,[status(thm)],[c_0_53, c_0_65])).
 % 0.06/0.38  fof(c_0_69, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
 % 0.06/0.38  fof(c_0_70, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])).
 % 0.06/0.38  cnf(c_0_71, plain, (frontsegP(app(X1,X2),X3)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
 % 0.06/0.38  cnf(c_0_72, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
 % 0.06/0.38  cnf(c_0_73, negated_conjecture, (esk52_0=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_68]), c_0_31])])).
 % 0.06/0.38  cnf(c_0_74, negated_conjecture, (app(esk48_0,nil)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_55]), c_0_26]), c_0_56]), c_0_27])])).
 % 0.06/0.38  cnf(c_0_75, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_69])).
 % 0.06/0.38  cnf(c_0_76, plain, (nil=app(X1,X2)|~ssList(X1)|~ssList(X2)|nil!=X1|nil!=X2), inference(split_conjunct,[status(thm)],[c_0_70])).
 % 0.06/0.38  cnf(c_0_77, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_56])])).
 % 0.06/0.38  cnf(c_0_78, negated_conjecture, (app(esk49_0,nil)=esk49_0), inference(rw,[status(thm)],[c_0_68, c_0_73])).
 % 0.06/0.38  cnf(c_0_79, negated_conjecture, (rearsegP(esk48_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_74]), c_0_35]), c_0_56])])).
 % 0.06/0.38  fof(c_0_80, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X5,X6)!=app(X5,X4)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
 % 0.06/0.38  cnf(c_0_81, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_75]), c_0_21])).
 % 0.06/0.38  cnf(c_0_82, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_76])]), c_0_56])])).
 % 0.06/0.38  cnf(c_0_83, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
 % 0.06/0.38  cnf(c_0_84, negated_conjecture, (frontsegP(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_56]), c_0_58])])).
 % 0.06/0.38  cnf(c_0_85, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
 % 0.06/0.38  cnf(c_0_86, negated_conjecture, (app(esk49_0,app(nil,X1))=app(esk49_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_78]), c_0_56]), c_0_58])])).
 % 0.06/0.38  fof(c_0_87, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).
 % 0.06/0.38  cnf(c_0_88, negated_conjecture, (app(esk7_2(esk48_0,nil),nil)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_79]), c_0_56]), c_0_35])])).
 % 0.06/0.38  cnf(c_0_89, negated_conjecture, (ssList(esk7_2(esk48_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_79]), c_0_56]), c_0_35])])).
 % 0.06/0.38  cnf(c_0_90, plain, (X1=X2|app(X3,X1)!=app(X3,X2)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_80])).
 % 0.06/0.38  cnf(c_0_91, negated_conjecture, (frontsegP(esk49_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_53]), c_0_27]), c_0_31])])).
 % 0.06/0.38  cnf(c_0_92, plain, (app(nil,app(nil,X1))=app(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_82]), c_0_56])])).
 % 0.06/0.38  cnf(c_0_93, negated_conjecture, (app(nil,esk6_2(esk49_0,nil))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_56]), c_0_58])])).
 % 0.06/0.38  cnf(c_0_94, negated_conjecture, (ssList(esk6_2(esk49_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_84]), c_0_56]), c_0_58])])).
 % 0.06/0.38  cnf(c_0_95, negated_conjecture, (ssList(app(esk49_0,X1))|~ssList(app(nil,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_86]), c_0_58])])).
 % 0.06/0.38  cnf(c_0_96, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
 % 0.06/0.38  cnf(c_0_97, negated_conjecture, (esk7_2(esk48_0,nil)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_88]), c_0_89])])).
 % 0.06/0.38  fof(c_0_98, plain, (![X2]:(~ssList(X2)|frontsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
 % 0.06/0.38  cnf(c_0_99, negated_conjecture, (X1=esk54_0|app(esk52_0,X1)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_53]), c_0_31]), c_0_27])])).
 % 0.06/0.38  cnf(c_0_100, negated_conjecture, (app(esk52_0,esk6_2(esk49_0,esk52_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_91]), c_0_31]), c_0_58])])).
 % 0.06/0.38  cnf(c_0_101, negated_conjecture, (ssList(esk6_2(esk49_0,esk52_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_91]), c_0_31]), c_0_58])])).
 % 0.06/0.38  cnf(c_0_102, negated_conjecture, (app(nil,esk49_0)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94])])).
 % 0.06/0.38  cnf(c_0_103, negated_conjecture, (ssList(app(esk49_0,X1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
 % 0.06/0.38  cnf(c_0_104, negated_conjecture, (app(esk52_0,esk53_0)=esk48_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_65]), c_0_97])).
 % 0.06/0.38  cnf(c_0_105, plain, (frontsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_98])).
 % 0.06/0.38  cnf(c_0_106, negated_conjecture, (esk6_2(esk49_0,esk52_0)=esk54_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101])])).
 % 0.06/0.38  cnf(c_0_107, negated_conjecture, (segmentP(app(esk49_0,X1),esk49_0)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_102]), c_0_56]), c_0_58])]), c_0_103])).
 % 0.06/0.38  cnf(c_0_108, negated_conjecture, (app(esk49_0,esk53_0)=esk48_0), inference(rw,[status(thm)],[c_0_104, c_0_73])).
 % 0.06/0.38  cnf(c_0_109, plain, (app(X1,esk6_2(X1,X1))=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_83, c_0_105])).
 % 0.06/0.38  cnf(c_0_110, negated_conjecture, (esk6_2(esk49_0,esk49_0)=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_65]), c_0_73])).
 % 0.06/0.38  cnf(c_0_111, negated_conjecture, (segmentP(esk48_0,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_30])])).
 % 0.06/0.38  cnf(c_0_112, negated_conjecture, (segmentP(esk49_0,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_109]), c_0_110]), c_0_56]), c_0_58])])).
 % 0.06/0.38  cnf(c_0_113, negated_conjecture, (neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.38  cnf(c_0_114, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_111]), c_0_112]), c_0_58]), c_0_113])]), ['proof']).
 % 0.06/0.38  # SZS output end CNFRefutation
 % 0.06/0.38  # Training examples: 0 positive, 0 negative
 % 0.06/0.38
 % 0.06/0.38  # -------------------------------------------------
 % 0.06/0.38  # User time                : 0.145 s
 % 0.06/0.38  # System time              : 0.008 s
 % 0.06/0.38  # Total time               : 0.153 s
 % 0.06/0.38  # Maximum resident set size: 1872 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
