   --------------------------------------------------------------------------

                   TSTP Solution File: SWC045+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC045+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n115.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:25 EDT 2016

 % Result     : Theorem 162.80s
 % Output     : CNFRefutation 162.80s
 % Verified   :
 % Statistics : Number of formulae       :  171 ( 736 expanded)
 %              Number of clauses        :  114 ( 344 expanded)
 %              Number of leaves         :   28 ( 225 expanded)
 %              Depth                    :   30
 %              Number of atoms          :  690 (2858 expanded)
 %              Number of equality atoms :  154 ( 629 expanded)
 %              Maximal formula depth    :   21 (   5 average)
 %              Maximal clause size      :   18 (   3 average)
 %              Maximal term depth       :    5 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => tl(cons(X2,X1)) = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( singletonP(X1)
       <=> ? [X2] :
             ( ssItem(X2)
             & cons(X2,nil) = X1 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax4)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(X1,nil) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil != X1
            => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax86)).

 fof(c_0_5,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_6,axiom,(
     ~ singletonP(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax39)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_8,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(nil,X1) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

 fof(c_0_9,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_10,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_11,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ? [X2] :
             ( ssItem(X2)
             & hd(X1) = X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax75)).

 fof(c_0_12,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( rearsegP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & app(X3,X2) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax6)).

 fof(c_0_13,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( frontsegP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & app(X2,X3) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax5)).

 fof(c_0_14,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( rearsegP(nil,X1)
       <=> nil = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax52)).

 fof(c_0_15,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( segmentP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & ? [X4] :
                     ( ssList(X4)
                     & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax7)).

 fof(c_0_16,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( ( frontsegP(X1,X2)
                   & frontsegP(X2,X3) )
                => frontsegP(X1,X3) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax40)).

 fof(c_0_17,axiom,(
     ! [X1] :
       ( ssList(X1)
      => frontsegP(X1,nil) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax45)).

 fof(c_0_18,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( segmentP(nil,X1)
       <=> nil = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax58)).

 fof(c_0_19,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( ( segmentP(X1,X2)
                   & segmentP(X2,X3) )
                => segmentP(X1,X3) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax53)).

 fof(c_0_20,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( frontsegP(nil,X1)
       <=> nil = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax46)).

 fof(c_0_21,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ~ memberP(nil,X1) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax38)).

 fof(c_0_22,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( memberP(app(X2,X3),X1)
               <=> ( memberP(X2,X1)
                   | memberP(X3,X1) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax36)).

 fof(c_0_23,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( memberP(cons(X2,X3),X1)
               <=> ( X1 = X2
                   | memberP(X3,X1) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax37)).

 fof(c_0_24,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => hd(cons(X2,X1)) = X2 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax23)).

 fof(c_0_25,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( nil != X2
                     | X2 != X4
                     | X1 != X3
                     | nil = X1
                     | ? [X5] :
                         ( ssItem(X5)
                         & ( ( ~ memberP(X3,X5)
                             & ? [X6] :
                                 ( ssList(X6)
                                 & segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) ) )
                           | ( ! [X6] :
                                 ( ssList(X6)
                                => ~ segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) )
                             & memberP(X3,X5) ) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_26,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssItem(X3)
              => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

 fof(c_0_27,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil = X1
         | ? [X2] :
             ( ssList(X2)
             & ? [X3] :
                 ( ssItem(X3)
                 & cons(X3,X2) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax20)).

 fof(c_0_28,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | tl(cons(X4,X3)) = X3 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])).

 fof(c_0_29,plain,(
     ! [X3,X5] :
       ( ( ssItem(esk5_1(X3))
         | ~ singletonP(X3)
         | ~ ssList(X3) )
       & ( cons(esk5_1(X3),nil) = X3
         | ~ singletonP(X3)
         | ~ ssList(X3) )
       & ( ~ ssItem(X5)
         | cons(X5,nil) != X3
         | singletonP(X3)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

 fof(c_0_30,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

 fof(c_0_31,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(X2,nil) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

 fof(c_0_32,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | nil = X3
       | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

 cnf(c_0_33,plain,
     ( tl(cons(X1,X2)) = X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_28])).

 cnf(c_0_34,plain,
     ( cons(esk5_1(X1),nil) = X1
     | ~ ssList(X1)
     | ~ singletonP(X1) ),
     inference(split_conjunct,[status(thm)],[[5]c_0_29])).

 cnf(c_0_35,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[6]c_0_5])).

 cnf(c_0_36,plain,
     ( ssItem(esk5_1(X1))
     | ~ ssList(X1)
     | ~ singletonP(X1) ),
     inference(split_conjunct,[status(thm)],[[7]c_0_29])).

 cnf(c_0_37,plain,
     ( singletonP(X1)
     | ~ ssList(X1)
     | cons(X2,nil) != X1
     | ~ ssItem(X2) ),
     inference(split_conjunct,[status(thm)],[[8]c_0_29])).

 fof(c_0_38,plain,(
     ~ singletonP(nil) ),
     inference(fof_simplification,[status(thm)],[[9]c_0_6])).

 fof(c_0_39,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).

 cnf(c_0_40,plain,
     ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[10]c_0_30])).

 cnf(c_0_41,plain,
     ( app(X1,nil) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_31])).

 cnf(c_0_42,plain,
     ( tl(app(X1,X2)) = app(tl(X1),X2)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_32])).

 cnf(c_0_43,plain,
     ( tl(X1) = nil
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[13]c_0_33,c_0_34]),c_0_35])]),[14]c_0_36])).

 cnf(c_0_44,plain,
     ( singletonP(cons(X1,nil))
     | ~ ssList(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(er,[status(thm)],[[15]c_0_37])).

 cnf(c_0_45,plain,
     ( ~ singletonP(nil) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_38])).

 fof(c_0_46,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(nil,X2) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).

 fof(c_0_47,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

 cnf(c_0_48,plain,
     ( nil = X2
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[17]c_0_39])).

 cnf(c_0_49,plain,
     ( app(X1,app(nil,X2)) = app(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_40,c_0_41]),c_0_35])])).

 cnf(c_0_50,plain,
     ( app(tl(X1),nil) = tl(X1)
     | nil = X1
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[19]c_0_42,c_0_41]),c_0_35])])).

 cnf(c_0_51,plain,
     ( tl(cons(X1,nil)) = nil
     | ~ ssList(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(spm,[status(thm)],[[20]c_0_43,[21]c_0_44])).

 cnf(c_0_52,plain,
     ( cons(X1,nil) != nil
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[22]c_0_45,c_0_37]),c_0_35])])).

 fof(c_0_53,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).

 cnf(c_0_54,plain,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[23]c_0_46])).

 cnf(c_0_55,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[24]c_0_47])).

 cnf(c_0_56,plain,
     ( app(nil,X1) = nil
     | app(X2,X1) != nil
     | ~ ssList(app(nil,X1))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[25]c_0_48,[26]c_0_49])).

 cnf(c_0_57,plain,
     ( app(nil,nil) = nil
     | ~ ssList(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_50,c_0_51]),[28]c_0_52])).

 cnf(c_0_58,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[29]c_0_53])).

 fof(c_0_59,plain,(
     ! [X3] :
       ( ( ssItem(esk46_1(X3))
         | nil = X3
         | ~ ssList(X3) )
       & ( hd(X3) = esk46_1(X3)
         | nil = X3
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).

 cnf(c_0_60,plain,
     ( app(nil,app(X1,X2)) = app(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[30]c_0_40,c_0_54]),c_0_35])])).

 cnf(c_0_61,plain,
     ( app(X1,app(X2,nil)) = app(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[31]c_0_41,c_0_40]),c_0_35])]),[32]c_0_55])).

 cnf(c_0_62,plain,
     ( app(nil,X1) = nil
     | app(X2,X1) != nil
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[33]c_0_56,c_0_55]),c_0_35])])).

 cnf(c_0_63,plain,
     ( app(nil,nil) = nil
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[34]c_0_57,c_0_58]),c_0_35])])).

 cnf(c_0_64,plain,
     ( nil = X1
     | ssItem(esk46_1(X1))
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[35]c_0_59])).

 cnf(c_0_65,plain,
     ( app(nil,X1) = app(X1,nil)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[36]c_0_60,c_0_61]),c_0_35])])).

 cnf(c_0_66,plain,
     ( app(nil,X1) = nil
     | X1 != nil
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[37]c_0_62,c_0_54]),c_0_35])])).

 cnf(c_0_67,plain,
     ( app(nil,nil) = nil
     | nil = X1
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[38]c_0_63,[39]c_0_64])).

 fof(c_0_68,plain,(
     ! [X4,X5,X7] :
       ( ( ssList(esk7_2(X4,X5))
         | ~ rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( app(esk7_2(X4,X5),X5) = X4
         | ~ rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( ~ ssList(X7)
         | app(X7,X5) != X4
         | rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).

 fof(c_0_69,plain,(
     ! [X4,X5,X7] :
       ( ( ssList(esk6_2(X4,X5))
         | ~ frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( app(X5,esk6_2(X4,X5)) = X4
         | ~ frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( ~ ssList(X7)
         | app(X5,X7) != X4
         | frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])).

 cnf(c_0_70,plain,
     ( ssList(app(nil,X1))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[40]c_0_55,c_0_65]),c_0_35])])).

 cnf(c_0_71,plain,
     ( app(nil,app(X1,X2)) = app(nil,X2)
     | X1 != nil
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[41]c_0_40,c_0_66]),c_0_35])])).

 cnf(c_0_72,plain,
     ( app(nil,nil) = nil
     | app(X1,X2) = nil
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[42]c_0_67,[43]c_0_55])).

 fof(c_0_73,plain,(
     ! [X2] :
       ( ( ~ rearsegP(nil,X2)
         | nil = X2
         | ~ ssList(X2) )
       & ( nil != X2
         | rearsegP(nil,X2)
         | ~ ssList(X2) ) ) ),
     inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).

 cnf(c_0_74,plain,
     ( rearsegP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(X3,X2) != X1
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[44]c_0_68])).

 cnf(c_0_75,plain,
     ( app(X2,esk6_2(X1,X2)) = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ frontsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[45]c_0_69])).

 cnf(c_0_76,plain,
     ( ssList(esk6_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ frontsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[46]c_0_69])).

 fof(c_0_77,plain,(
     ! [X5,X6,X9,X10] :
       ( ( ssList(esk8_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ssList(esk9_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ~ ssList(X9)
         | ~ ssList(X10)
         | app(app(X9,X6),X10) != X5
         | segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])])).

 cnf(c_0_78,plain,
     ( app(X1,app(X2,app(nil,X3))) = app(app(X1,X2),X3)
     | ~ ssList(X2)
     | ~ ssList(X1)
     | ~ ssList(X3) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[47]c_0_40,c_0_49]),c_0_55]),[48]c_0_70])).

 cnf(c_0_79,plain,
     ( app(nil,app(X1,nil)) = app(nil,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[49]c_0_71,c_0_65]),c_0_35])])).

 cnf(c_0_80,plain,
     ( app(nil,nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(ef,[status(thm)],[c_0_72]),c_0_35])])).

 cnf(c_0_81,plain,
     ( nil = X1
     | ~ ssList(X1)
     | ~ rearsegP(nil,X1) ),
     inference(split_conjunct,[status(thm)],[[50]c_0_73])).

 cnf(c_0_82,plain,
     ( rearsegP(X1,esk6_2(X1,X2))
     | ~ frontsegP(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[[51]c_0_74,c_0_75])]),[52]c_0_76])).

 fof(c_0_83,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | ~ frontsegP(X4,X5)
       | ~ frontsegP(X5,X6)
       | frontsegP(X4,X6) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).

 fof(c_0_84,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | frontsegP(X2,nil) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).

 fof(c_0_85,plain,(
     ! [X2] :
       ( ( ~ segmentP(nil,X2)
         | nil = X2
         | ~ ssList(X2) )
       & ( nil != X2
         | segmentP(nil,X2)
         | ~ ssList(X2) ) ) ),
     inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).

 fof(c_0_86,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | ~ segmentP(X4,X5)
       | ~ segmentP(X5,X6)
       | segmentP(X4,X6) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).

 cnf(c_0_87,plain,
     ( segmentP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(app(X3,X2),X4) != X1
     | ~ ssList(X4)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[53]c_0_77])).

 cnf(c_0_88,plain,
     ( app(app(X1,nil),nil) = app(X1,nil)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[54]c_0_78,c_0_79]),c_0_80]),c_0_35])])).

 cnf(c_0_89,plain,
     ( esk6_2(nil,X1) = nil
     | ~ frontsegP(nil,X1)
     | ~ ssList(esk6_2(nil,X1))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[55]c_0_81,c_0_82]),c_0_35])])).

 cnf(c_0_90,plain,
     ( frontsegP(X1,X2)
     | ~ frontsegP(X3,X2)
     | ~ frontsegP(X1,X3)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[56]c_0_83])).

 cnf(c_0_91,plain,
     ( frontsegP(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[57]c_0_84])).

 fof(c_0_92,plain,(
     ! [X2] :
       ( ( ~ frontsegP(nil,X2)
         | nil = X2
         | ~ ssList(X2) )
       & ( nil != X2
         | frontsegP(nil,X2)
         | ~ ssList(X2) ) ) ),
     inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).

 cnf(c_0_93,plain,
     ( nil = X1
     | ~ ssList(X1)
     | ~ segmentP(nil,X1) ),
     inference(split_conjunct,[status(thm)],[[58]c_0_85])).

 cnf(c_0_94,plain,
     ( segmentP(X1,X2)
     | ~ segmentP(X3,X2)
     | ~ segmentP(X1,X3)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[59]c_0_86])).

 cnf(c_0_95,plain,
     ( segmentP(app(app(X1,X2),X3),X2)
     | ~ ssList(app(app(X1,X2),X3))
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(er,[status(thm)],[[60]c_0_87])).

 cnf(c_0_96,plain,
     ( app(app(nil,X1),nil) = app(nil,X1)
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[61]c_0_88,[62]c_0_65])).

 cnf(c_0_97,plain,
     ( esk6_2(nil,X1) = nil
     | ~ frontsegP(nil,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[63]c_0_89,c_0_76]),c_0_35])])).

 cnf(c_0_98,plain,
     ( frontsegP(X1,X2)
     | ~ frontsegP(nil,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[64]c_0_90,c_0_91]),c_0_35])])).

 cnf(c_0_99,plain,
     ( frontsegP(nil,X1)
     | ~ ssList(X1)
     | nil != X1 ),
     inference(split_conjunct,[status(thm)],[[65]c_0_92])).

 cnf(c_0_100,plain,
     ( nil = X1
     | ~ segmentP(nil,X2)
     | ~ segmentP(X2,X1)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[66]c_0_93,c_0_94]),c_0_35])])).

 cnf(c_0_101,plain,
     ( segmentP(app(nil,X1),X1)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[67]c_0_95,c_0_96]),c_0_35])]),[68]c_0_70])).

 cnf(c_0_102,plain,
     ( app(X1,nil) = nil
     | ~ frontsegP(nil,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[69]c_0_75,c_0_97]),c_0_35])])).

 cnf(c_0_103,plain,
     ( frontsegP(X1,X2)
     | nil != X2
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[70]c_0_98,[71]c_0_99])).

 fof(c_0_104,plain,(
     ! [X1] :
       ( ssItem(X1)
      => ~ memberP(nil,X1) ) ),
     inference(fof_simplification,[status(thm)],[[72]c_0_21])).

 cnf(c_0_105,plain,
     ( nil = X1
     | ~ segmentP(nil,app(nil,X1))
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[73]c_0_100,c_0_101]),[74]c_0_70])).

 cnf(c_0_106,plain,
     ( ssList(app(X1,nil))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[75]c_0_55,c_0_65]),c_0_35])])).

 cnf(c_0_107,plain,
     ( app(X1,nil) = nil
     | nil != X1
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[76]c_0_102,c_0_103]),c_0_35])])).

 fof(c_0_108,plain,(
     ! [X4,X5,X6] :
       ( ( ~ memberP(app(X5,X6),X4)
         | memberP(X5,X4)
         | memberP(X6,X4)
         | ~ ssList(X6)
         | ~ ssList(X5)
         | ~ ssItem(X4) )
       & ( ~ memberP(X5,X4)
         | memberP(app(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssList(X5)
         | ~ ssItem(X4) )
       & ( ~ memberP(X6,X4)
         | memberP(app(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssList(X5)
         | ~ ssItem(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).

 fof(c_0_109,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | ~ memberP(nil,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_104])])).

 cnf(c_0_110,plain,
     ( app(X1,nil) = nil
     | ~ segmentP(nil,app(nil,X1))
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[77]c_0_105,c_0_61]),c_0_35])]),[78]c_0_106])).

 cnf(c_0_111,plain,
     ( app(X1,app(nil,X2)) = app(nil,X2)
     | nil != X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[79]c_0_40,c_0_107]),c_0_35])])).

 cnf(c_0_112,plain,
     ( memberP(app(X2,X3),X1)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ memberP(X3,X1) ),
     inference(split_conjunct,[status(thm)],[[80]c_0_108])).

 cnf(c_0_113,plain,
     ( ~ memberP(nil,X1)
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[81]c_0_109])).

 cnf(c_0_114,plain,
     ( app(app(X1,nil),nil) = nil
     | ~ segmentP(nil,app(nil,X1))
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[82]c_0_110,c_0_61]),c_0_35])]),[83]c_0_106])).

 cnf(c_0_115,plain,
     ( app(X1,X2) = X2
     | nil != X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[84]c_0_111,[85]c_0_54])).

 cnf(c_0_116,plain,
     ( X1 != nil
     | ~ memberP(X1,X2)
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[86]c_0_112,c_0_66]),c_0_35])]),[87]c_0_113])).

 cnf(c_0_117,plain,
     ( app(app(X1,nil),nil) = nil
     | ~ segmentP(nil,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[88]c_0_114,c_0_115]),c_0_35])])).

 fof(c_0_118,plain,(
     ! [X4,X5,X6] :
       ( ( ~ memberP(cons(X5,X6),X4)
         | X4 = X5
         | memberP(X6,X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) )
       & ( X4 != X5
         | memberP(cons(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) )
       & ( ~ memberP(X6,X4)
         | memberP(cons(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).

 cnf(c_0_119,plain,
     ( app(X1,X2) != nil
     | ~ memberP(X2,X3)
     | ~ ssList(X2)
     | ~ ssList(X1)
     | ~ ssItem(X3) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[89]c_0_116,c_0_112]),[90]c_0_55])).

 cnf(c_0_120,plain,
     ( memberP(app(X2,X3),X1)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ memberP(X2,X1) ),
     inference(split_conjunct,[status(thm)],[[91]c_0_108])).

 cnf(c_0_121,plain,
     ( app(app(nil,X1),nil) = nil
     | ~ segmentP(nil,X1)
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[92]c_0_117,[93]c_0_65])).

 cnf(c_0_122,plain,
     ( memberP(cons(X2,X3),X1)
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | X1 != X2 ),
     inference(split_conjunct,[status(thm)],[[94]c_0_118])).

 fof(c_0_123,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | hd(cons(X4,X3)) = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).

 fof(c_0_124,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( nil != X2
                       | X2 != X4
                       | X1 != X3
                       | nil = X1
                       | ? [X5] :
                           ( ssItem(X5)
                           & ( ( ~ memberP(X3,X5)
                               & ? [X6] :
                                   ( ssList(X6)
                                   & segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) ) )
                             | ( ! [X6] :
                                   ( ssList(X6)
                                  => ~ segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))) )
                               & memberP(X3,X5) ) ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_25])])).

 cnf(c_0_125,plain,
     ( app(X1,app(X2,X3)) != nil
     | ~ memberP(X2,X4)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssItem(X4) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[95]c_0_119,c_0_120]),[96]c_0_55])).

 cnf(c_0_126,plain,
     ( app(nil,app(X1,nil)) = nil
     | ~ segmentP(nil,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[97]c_0_40,c_0_121]),c_0_35])])).

 cnf(c_0_127,plain,
     ( app(esk7_2(X1,X2),X2) = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ rearsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[98]c_0_68])).

 cnf(c_0_128,plain,
     ( ssList(esk7_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ rearsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[99]c_0_68])).

 cnf(c_0_129,plain,
     ( memberP(cons(X1,X2),X1)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(er,[status(thm)],[[100]c_0_122])).

 cnf(c_0_130,plain,
     ( hd(cons(X1,X2)) = X1
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[101]c_0_123])).

 fof(c_0_131,negated_conjecture,(
     ! [X11,X12] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & nil = esk49_0
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & nil != esk48_0
       & ( memberP(esk50_0,X11)
         | ~ ssList(X12)
         | ~ segmentP(esk51_0,app(app(cons(X11,nil),X12),cons(X11,nil)))
         | ~ ssItem(X11) )
       & ( ssList(esk52_1(X11))
         | ~ memberP(esk50_0,X11)
         | ~ ssItem(X11) )
       & ( segmentP(esk51_0,app(app(cons(X11,nil),esk52_1(X11)),cons(X11,nil)))
         | ~ memberP(esk50_0,X11)
         | ~ ssItem(X11) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_124])])])])])).

 cnf(c_0_132,plain,
     ( ~ segmentP(nil,X1)
     | ~ memberP(X1,X2)
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[102]c_0_125,c_0_126]),c_0_35])])).

 cnf(c_0_133,plain,
     ( memberP(X1,X2)
     | ~ rearsegP(X1,X3)
     | ~ memberP(X3,X2)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[103]c_0_112,c_0_127]),[104]c_0_128])).

 cnf(c_0_134,plain,
     ( memberP(X1,esk5_1(X1))
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[105]c_0_129,c_0_34]),c_0_35])]),[106]c_0_36])).

 cnf(c_0_135,plain,
     ( esk5_1(X1) = hd(X1)
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[107]c_0_130,c_0_34]),c_0_35])]),[108]c_0_36])).

 cnf(c_0_136,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[109]c_0_131])).

 cnf(c_0_137,negated_conjecture,
     ( nil = esk49_0 ),
     inference(split_conjunct,[status(thm)],[[110]c_0_131])).

 cnf(c_0_138,plain,
     ( ~ segmentP(nil,X1)
     | ~ rearsegP(X1,X2)
     | ~ memberP(X2,X3)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ ssItem(X3) ),
     inference(spm,[status(thm)],[[111]c_0_132,[112]c_0_133])).

 cnf(c_0_139,plain,
     ( memberP(X1,hd(X1))
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[113]c_0_134,[114]c_0_135])).

 cnf(c_0_140,plain,
     ( ssItem(hd(X1))
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[115]c_0_36,[116]c_0_135])).

 cnf(c_0_141,negated_conjecture,
     ( segmentP(esk51_0,app(app(cons(X1,nil),esk52_1(X1)),cons(X1,nil)))
     | ~ ssItem(X1)
     | ~ memberP(esk50_0,X1) ),
     inference(split_conjunct,[status(thm)],[[117]c_0_131])).

 cnf(c_0_142,negated_conjecture,
     ( esk51_0 = nil ),
     inference(rw,[status(thm)],[[118]c_0_136,[119]c_0_137])).

 cnf(c_0_143,plain,
     ( ~ segmentP(nil,X1)
     | ~ rearsegP(X1,X2)
     | ~ singletonP(X2)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[120]c_0_138,c_0_139]),[121]c_0_140])).

 cnf(c_0_144,plain,
     ( rearsegP(app(X1,X2),X2)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_74]),[122]c_0_55])).

 cnf(c_0_145,negated_conjecture,
     ( segmentP(nil,app(app(cons(X1,nil),esk52_1(X1)),cons(X1,nil)))
     | ~ memberP(esk50_0,X1)
     | ~ ssItem(X1) ),
     inference(rw,[status(thm)],[[123]c_0_141,[124]c_0_142])).

 cnf(c_0_146,plain,
     ( ~ segmentP(nil,app(X1,X2))
     | ~ singletonP(X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[125]c_0_143,c_0_144]),[126]c_0_55])).

 cnf(c_0_147,negated_conjecture,
     ( segmentP(nil,app(app(X1,esk52_1(esk5_1(X1))),X1))
     | ~ singletonP(X1)
     | ~ memberP(esk50_0,esk5_1(X1))
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[127]c_0_145,c_0_34]),[128]c_0_36])).

 cnf(c_0_148,negated_conjecture,
     ( ~ singletonP(X1)
     | ~ memberP(esk50_0,esk5_1(X1))
     | ~ ssList(app(X1,esk52_1(esk5_1(X1))))
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[129]c_0_146,[130]c_0_147])).

 fof(c_0_149,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssItem(X6)
       | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).

 cnf(c_0_150,negated_conjecture,
     ( ~ singletonP(X1)
     | ~ memberP(esk50_0,esk5_1(X1))
     | ~ ssList(esk52_1(esk5_1(X1)))
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[131]c_0_148,[132]c_0_55])).

 cnf(c_0_151,negated_conjecture,
     ( ssList(esk52_1(X1))
     | ~ ssItem(X1)
     | ~ memberP(esk50_0,X1) ),
     inference(split_conjunct,[status(thm)],[[133]c_0_131])).

 cnf(c_0_152,plain,
     ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[134]c_0_149])).

 cnf(c_0_153,negated_conjecture,
     ( ~ singletonP(X1)
     | ~ memberP(esk50_0,esk5_1(X1))
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[135]c_0_150,c_0_151]),[136]c_0_36])).

 cnf(c_0_154,plain,
     ( ssList(cons(X1,app(X2,X3)))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[137]c_0_55,c_0_152]),[138]c_0_58])).

 cnf(c_0_155,negated_conjecture,
     ( ~ singletonP(X1)
     | ~ memberP(esk50_0,hd(X1))
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[139]c_0_153,[140]c_0_135])).

 cnf(c_0_156,plain,
     ( ssList(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[141]c_0_154,c_0_80]),c_0_35])])).

 fof(c_0_157,plain,(
     ! [X4] :
       ( ( ssList(esk44_1(X4))
         | nil = X4
         | ~ ssList(X4) )
       & ( ssItem(esk45_1(X4))
         | nil = X4
         | ~ ssList(X4) )
       & ( cons(esk45_1(X4),esk44_1(X4)) = X4
         | nil = X4
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])).

 cnf(c_0_158,negated_conjecture,
     ( ~ memberP(esk50_0,hd(cons(X1,nil)))
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[142]c_0_155,c_0_44]),[143]c_0_156])).

 cnf(c_0_159,plain,
     ( nil = X1
     | cons(esk45_1(X1),esk44_1(X1)) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[144]c_0_157])).

 cnf(c_0_160,plain,
     ( nil = X1
     | ssItem(esk45_1(X1))
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[145]c_0_157])).

 cnf(c_0_161,plain,
     ( nil = X1
     | ssList(esk44_1(X1))
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[146]c_0_157])).

 cnf(c_0_162,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[147]c_0_131])).

 cnf(c_0_163,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[148]c_0_131])).

 cnf(c_0_164,negated_conjecture,
     ( nil != esk48_0 ),
     inference(split_conjunct,[status(thm)],[[149]c_0_131])).

 cnf(c_0_165,negated_conjecture,
     ( ~ memberP(esk50_0,X1)
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[150]c_0_158,c_0_130]),c_0_35])])).

 cnf(c_0_166,plain,
     ( nil = X1
     | memberP(X1,esk45_1(X1))
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[151]c_0_129,c_0_159]),c_0_160]),[152]c_0_161])).

 cnf(c_0_167,negated_conjecture,
     ( ssList(esk50_0) ),
     inference(rw,[status(thm)],[[153]c_0_162,[154]c_0_163])).

 cnf(c_0_168,negated_conjecture,
     ( esk50_0 != nil ),
     inference(rw,[status(thm)],[[155]c_0_164,[156]c_0_163])).

 cnf(c_0_169,negated_conjecture,
     ( ~ ssItem(esk45_1(esk50_0)) ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[157]c_0_165,c_0_166]),c_0_167])]),[158]c_0_168])).

 cnf(c_0_170,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[159]c_0_169,c_0_160]),c_0_167])]),[160]c_0_168]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC045+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.22  % Computer   : n115.star.cs.uiowa.edu
 % 0.02/0.22  % Model      : x86_64 x86_64
 % 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.22  % Memory     : 32218.75MB
 % 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.22  % CPULimit   : 300
 % 0.02/0.22  % DateTime   : Tue May 10 20:34:53 CDT 2016
 % 0.02/0.22  % CPUTime    :
 % 0.02/0.22  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 151.10/151.24  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 151.10/151.24  # and selection function SelectCQPrecWNTNp.
 % 151.10/151.24  #
 % 151.10/151.24  # Preprocessing time       : 0.014 s
 % 151.10/151.24  # Presaturation interreduction done
 % 151.40/151.57  # No success with AutoSched0
 % 151.40/151.57  # Trying AutoSched1 for 73 seconds
 % 162.80/162.90  # AutoSched1-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S00
 % 162.80/162.90  # and selection function NoSelection.
 % 162.80/162.90  #
 % 162.80/162.90  # Preprocessing time       : 0.014 s
 % 162.80/162.90  # Presaturation interreduction done
 % 162.80/162.90
 % 162.80/162.90  # Proof found!
 % 162.80/162.90  # SZS status Theorem
 % 162.80/162.90  # SZS output start CNFRefutation
 % 162.80/162.90  fof(c_0_0, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
 % 162.80/162.90  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)<=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax4)).
 % 162.80/162.90  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
 % 162.80/162.90  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
 % 162.80/162.90  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax86)).
 % 162.80/162.90  fof(c_0_5, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 162.80/162.90  fof(c_0_6, axiom, (~(singletonP(nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax39)).
 % 162.80/162.90  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 162.80/162.90  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 162.80/162.90  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 162.80/162.90  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 162.80/162.90  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>?[X2]:(ssItem(X2)&hd(X1)=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax75)).
 % 162.80/162.90  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)<=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax6)).
 % 162.80/162.90  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)<=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax5)).
 % 162.80/162.90  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>(rearsegP(nil,X1)<=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax52)).
 % 162.80/162.90  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)<=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax7)).
 % 162.80/162.90  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>((frontsegP(X1,X2)&frontsegP(X2,X3))=>frontsegP(X1,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax40)).
 % 162.80/162.90  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax45)).
 % 162.80/162.90  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>(segmentP(nil,X1)<=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax58)).
 % 162.80/162.90  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>((segmentP(X1,X2)&segmentP(X2,X3))=>segmentP(X1,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax53)).
 % 162.80/162.90  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>(frontsegP(nil,X1)<=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax46)).
 % 162.80/162.90  fof(c_0_21, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax38)).
 % 162.80/162.90  fof(c_0_22, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(memberP(app(X2,X3),X1)<=>(memberP(X2,X1)|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax36)).
 % 162.80/162.90  fof(c_0_23, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)<=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax37)).
 % 162.80/162.90  fof(c_0_24, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax23)).
 % 162.80/162.90  fof(c_0_25, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((nil!=X2|X2!=X4)|X1!=X3)|nil=X1)|?[X5]:(ssItem(X5)&((~(memberP(X3,X5))&?[X6]:(ssList(X6)&segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil)))))|(![X6]:(ssList(X6)=>~(segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil)))))&memberP(X3,X5)))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 162.80/162.90  fof(c_0_26, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
 % 162.80/162.90  fof(c_0_27, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax20)).
 % 162.80/162.90  fof(c_0_28, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_0])])])).
 % 162.80/162.90  fof(c_0_29, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
 % 162.80/162.90  fof(c_0_30, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
 % 162.80/162.90  fof(c_0_31, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
 % 162.80/162.90  fof(c_0_32, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 162.80/162.90  cnf(c_0_33, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
 % 162.80/162.90  cnf(c_0_34, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
 % 162.80/162.90  cnf(c_0_35, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 162.80/162.90  cnf(c_0_36, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
 % 162.80/162.90  cnf(c_0_37, plain, (singletonP(X1)|~ssList(X1)|cons(X2,nil)!=X1|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
 % 162.80/162.90  fof(c_0_38, plain, (~singletonP(nil)), inference(fof_simplification,[status(thm)],[c_0_6])).
 % 162.80/162.90  fof(c_0_39, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
 % 162.80/162.90  cnf(c_0_40, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
 % 162.80/162.90  cnf(c_0_41, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
 % 162.80/162.90  cnf(c_0_42, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
 % 162.80/162.90  cnf(c_0_43, plain, (tl(X1)=nil|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])]), c_0_36])).
 % 162.80/162.90  cnf(c_0_44, plain, (singletonP(cons(X1,nil))|~ssList(cons(X1,nil))|~ssItem(X1)), inference(er,[status(thm)],[c_0_37])).
 % 162.80/162.90  cnf(c_0_45, plain, (~singletonP(nil)), inference(split_conjunct,[status(thm)],[c_0_38])).
 % 162.80/162.90  fof(c_0_46, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
 % 162.80/162.90  fof(c_0_47, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
 % 162.80/162.90  cnf(c_0_48, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 162.80/162.90  cnf(c_0_49, plain, (app(X1,app(nil,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_50, plain, (app(tl(X1),nil)=tl(X1)|nil=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_41]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_51, plain, (tl(cons(X1,nil))=nil|~ssList(cons(X1,nil))|~ssItem(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
 % 162.80/162.90  cnf(c_0_52, plain, (cons(X1,nil)!=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_37]), c_0_35])])).
 % 162.80/162.90  fof(c_0_53, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
 % 162.80/162.90  cnf(c_0_54, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
 % 162.80/162.90  cnf(c_0_55, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
 % 162.80/162.90  cnf(c_0_56, plain, (app(nil,X1)=nil|app(X2,X1)!=nil|~ssList(app(nil,X1))|~ssList(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
 % 162.80/162.90  cnf(c_0_57, plain, (app(nil,nil)=nil|~ssList(cons(X1,nil))|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
 % 162.80/162.90  cnf(c_0_58, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
 % 162.80/162.90  fof(c_0_59, plain, (![X3]:(((ssItem(esk46_1(X3))|nil=X3)|~ssList(X3))&((hd(X3)=esk46_1(X3)|nil=X3)|~ssList(X3)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).
 % 162.80/162.90  cnf(c_0_60, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_54]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_61, plain, (app(X1,app(X2,nil))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_40]), c_0_35])]), c_0_55])).
 % 162.80/162.90  cnf(c_0_62, plain, (app(nil,X1)=nil|app(X2,X1)!=nil|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_55]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_63, plain, (app(nil,nil)=nil|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_64, plain, (nil=X1|ssItem(esk46_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
 % 162.80/162.90  cnf(c_0_65, plain, (app(nil,X1)=app(X1,nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_66, plain, (app(nil,X1)=nil|X1!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_54]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_67, plain, (app(nil,nil)=nil|nil=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
 % 162.80/162.90  fof(c_0_68, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).
 % 162.80/162.90  fof(c_0_69, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])])).
 % 162.80/162.90  cnf(c_0_70, plain, (ssList(app(nil,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_65]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_71, plain, (app(nil,app(X1,X2))=app(nil,X2)|X1!=nil|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_66]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_72, plain, (app(nil,nil)=nil|app(X1,X2)=nil|~ssList(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_67, c_0_55])).
 % 162.80/162.90  fof(c_0_73, plain, (![X2]:(((~rearsegP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|rearsegP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
 % 162.80/162.90  cnf(c_0_74, plain, (rearsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X3,X2)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_68])).
 % 162.80/162.90  cnf(c_0_75, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
 % 162.80/162.90  cnf(c_0_76, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
 % 162.80/162.90  fof(c_0_77, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])])).
 % 162.80/162.90  cnf(c_0_78, plain, (app(X1,app(X2,app(nil,X3)))=app(app(X1,X2),X3)|~ssList(X2)|~ssList(X1)|~ssList(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_49]), c_0_55]), c_0_70])).
 % 162.80/162.90  cnf(c_0_79, plain, (app(nil,app(X1,nil))=app(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_65]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_80, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(ef,[status(thm)],[c_0_72]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_81, plain, (nil=X1|~ssList(X1)|~rearsegP(nil,X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
 % 162.80/162.90  cnf(c_0_82, plain, (rearsegP(X1,esk6_2(X1,X2))|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75])]), c_0_76])).
 % 162.80/162.90  fof(c_0_83, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|((~frontsegP(X4,X5)|~frontsegP(X5,X6))|frontsegP(X4,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
 % 162.80/162.90  fof(c_0_84, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).
 % 162.80/162.90  fof(c_0_85, plain, (![X2]:(((~segmentP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|segmentP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
 % 162.80/162.90  fof(c_0_86, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|((~segmentP(X4,X5)|~segmentP(X5,X6))|segmentP(X4,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
 % 162.80/162.90  cnf(c_0_87, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
 % 162.80/162.90  cnf(c_0_88, plain, (app(app(X1,nil),nil)=app(X1,nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_89, plain, (esk6_2(nil,X1)=nil|~frontsegP(nil,X1)|~ssList(esk6_2(nil,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_90, plain, (frontsegP(X1,X2)|~frontsegP(X3,X2)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
 % 162.80/162.90  cnf(c_0_91, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
 % 162.80/162.90  fof(c_0_92, plain, (![X2]:(((~frontsegP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|frontsegP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
 % 162.80/162.90  cnf(c_0_93, plain, (nil=X1|~ssList(X1)|~segmentP(nil,X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
 % 162.80/162.90  cnf(c_0_94, plain, (segmentP(X1,X2)|~segmentP(X3,X2)|~segmentP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
 % 162.80/162.90  cnf(c_0_95, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_87])).
 % 162.80/162.90  cnf(c_0_96, plain, (app(app(nil,X1),nil)=app(nil,X1)|~ssList(X1)), inference(spm,[status(thm)],[c_0_88, c_0_65])).
 % 162.80/162.90  cnf(c_0_97, plain, (esk6_2(nil,X1)=nil|~frontsegP(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_76]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_98, plain, (frontsegP(X1,X2)|~frontsegP(nil,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_99, plain, (frontsegP(nil,X1)|~ssList(X1)|nil!=X1), inference(split_conjunct,[status(thm)],[c_0_92])).
 % 162.80/162.90  cnf(c_0_100, plain, (nil=X1|~segmentP(nil,X2)|~segmentP(X2,X1)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_101, plain, (segmentP(app(nil,X1),X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_35])]), c_0_70])).
 % 162.80/162.90  cnf(c_0_102, plain, (app(X1,nil)=nil|~frontsegP(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_97]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_103, plain, (frontsegP(X1,X2)|nil!=X2|~ssList(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
 % 162.80/162.90  fof(c_0_104, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_21])).
 % 162.80/162.90  cnf(c_0_105, plain, (nil=X1|~segmentP(nil,app(nil,X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_70])).
 % 162.80/162.90  cnf(c_0_106, plain, (ssList(app(X1,nil))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_65]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_107, plain, (app(X1,nil)=nil|nil!=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_35])])).
 % 162.80/162.90  fof(c_0_108, plain, (![X4]:![X5]:![X6]:(((((~memberP(app(X5,X6),X4)|(memberP(X5,X4)|memberP(X6,X4)))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&(((((~memberP(X5,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).
 % 162.80/162.90  fof(c_0_109, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_104])])).
 % 162.80/162.90  cnf(c_0_110, plain, (app(X1,nil)=nil|~segmentP(nil,app(nil,X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_61]), c_0_35])]), c_0_106])).
 % 162.80/162.90  cnf(c_0_111, plain, (app(X1,app(nil,X2))=app(nil,X2)|nil!=X1|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_107]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_112, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
 % 162.80/162.90  cnf(c_0_113, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
 % 162.80/162.90  cnf(c_0_114, plain, (app(app(X1,nil),nil)=nil|~segmentP(nil,app(nil,X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_61]), c_0_35])]), c_0_106])).
 % 162.80/162.90  cnf(c_0_115, plain, (app(X1,X2)=X2|nil!=X1|~ssList(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_111, c_0_54])).
 % 162.80/162.90  cnf(c_0_116, plain, (X1!=nil|~memberP(X1,X2)|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_66]), c_0_35])]), c_0_113])).
 % 162.80/162.90  cnf(c_0_117, plain, (app(app(X1,nil),nil)=nil|~segmentP(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_35])])).
 % 162.80/162.90  fof(c_0_118, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).
 % 162.80/162.90  cnf(c_0_119, plain, (app(X1,X2)!=nil|~memberP(X2,X3)|~ssList(X2)|~ssList(X1)|~ssItem(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_112]), c_0_55])).
 % 162.80/162.90  cnf(c_0_120, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
 % 162.80/162.90  cnf(c_0_121, plain, (app(app(nil,X1),nil)=nil|~segmentP(nil,X1)|~ssList(X1)), inference(spm,[status(thm)],[c_0_117, c_0_65])).
 % 162.80/162.90  cnf(c_0_122, plain, (memberP(cons(X2,X3),X1)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_118])).
 % 162.80/162.90  fof(c_0_123, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
 % 162.80/162.90  fof(c_0_124, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((nil!=X2|X2!=X4)|X1!=X3)|nil=X1)|?[X5]:(ssItem(X5)&((~memberP(X3,X5)&?[X6]:(ssList(X6)&segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil)))))|(![X6]:(ssList(X6)=>~segmentP(X4,app(app(cons(X5,nil),X6),cons(X5,nil))))&memberP(X3,X5))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_25])])).
 % 162.80/162.90  cnf(c_0_125, plain, (app(X1,app(X2,X3))!=nil|~memberP(X2,X4)|~ssList(X1)|~ssList(X3)|~ssList(X2)|~ssItem(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_55])).
 % 162.80/162.90  cnf(c_0_126, plain, (app(nil,app(X1,nil))=nil|~segmentP(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_121]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_127, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
 % 162.80/162.90  cnf(c_0_128, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
 % 162.80/162.90  cnf(c_0_129, plain, (memberP(cons(X1,X2),X1)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_122])).
 % 162.80/162.90  cnf(c_0_130, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_123])).
 % 162.80/162.90  fof(c_0_131, negated_conjecture, (![X11]:![X12]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((nil=esk49_0&esk49_0=esk51_0)&esk48_0=esk50_0)&nil!=esk48_0)&(((memberP(esk50_0,X11)|(~ssList(X12)|~segmentP(esk51_0,app(app(cons(X11,nil),X12),cons(X11,nil)))))|~ssItem(X11))&(((ssList(esk52_1(X11))|~memberP(esk50_0,X11))|~ssItem(X11))&((segmentP(esk51_0,app(app(cons(X11,nil),esk52_1(X11)),cons(X11,nil)))|~memberP(esk50_0,X11))|~ssItem(X11)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_124])])])])])).
 % 162.80/162.90  cnf(c_0_132, plain, (~segmentP(nil,X1)|~memberP(X1,X2)|~ssList(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_133, plain, (memberP(X1,X2)|~rearsegP(X1,X3)|~memberP(X3,X2)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_127]), c_0_128])).
 % 162.80/162.90  cnf(c_0_134, plain, (memberP(X1,esk5_1(X1))|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_34]), c_0_35])]), c_0_36])).
 % 162.80/162.90  cnf(c_0_135, plain, (esk5_1(X1)=hd(X1)|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_34]), c_0_35])]), c_0_36])).
 % 162.80/162.90  cnf(c_0_136, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_131])).
 % 162.80/162.90  cnf(c_0_137, negated_conjecture, (nil=esk49_0), inference(split_conjunct,[status(thm)],[c_0_131])).
 % 162.80/162.90  cnf(c_0_138, plain, (~segmentP(nil,X1)|~rearsegP(X1,X2)|~memberP(X2,X3)|~ssList(X1)|~ssList(X2)|~ssItem(X3)), inference(spm,[status(thm)],[c_0_132, c_0_133])).
 % 162.80/162.90  cnf(c_0_139, plain, (memberP(X1,hd(X1))|~singletonP(X1)|~ssList(X1)), inference(spm,[status(thm)],[c_0_134, c_0_135])).
 % 162.80/162.90  cnf(c_0_140, plain, (ssItem(hd(X1))|~singletonP(X1)|~ssList(X1)), inference(spm,[status(thm)],[c_0_36, c_0_135])).
 % 162.80/162.90  cnf(c_0_141, negated_conjecture, (segmentP(esk51_0,app(app(cons(X1,nil),esk52_1(X1)),cons(X1,nil)))|~ssItem(X1)|~memberP(esk50_0,X1)), inference(split_conjunct,[status(thm)],[c_0_131])).
 % 162.80/162.90  cnf(c_0_142, negated_conjecture, (esk51_0=nil), inference(rw,[status(thm)],[c_0_136, c_0_137])).
 % 162.80/162.90  cnf(c_0_143, plain, (~segmentP(nil,X1)|~rearsegP(X1,X2)|~singletonP(X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_140])).
 % 162.80/162.90  cnf(c_0_144, plain, (rearsegP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_74]), c_0_55])).
 % 162.80/162.90  cnf(c_0_145, negated_conjecture, (segmentP(nil,app(app(cons(X1,nil),esk52_1(X1)),cons(X1,nil)))|~memberP(esk50_0,X1)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_141, c_0_142])).
 % 162.80/162.90  cnf(c_0_146, plain, (~segmentP(nil,app(X1,X2))|~singletonP(X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_55])).
 % 162.80/162.90  cnf(c_0_147, negated_conjecture, (segmentP(nil,app(app(X1,esk52_1(esk5_1(X1))),X1))|~singletonP(X1)|~memberP(esk50_0,esk5_1(X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_34]), c_0_36])).
 % 162.80/162.90  cnf(c_0_148, negated_conjecture, (~singletonP(X1)|~memberP(esk50_0,esk5_1(X1))|~ssList(app(X1,esk52_1(esk5_1(X1))))|~ssList(X1)), inference(spm,[status(thm)],[c_0_146, c_0_147])).
 % 162.80/162.90  fof(c_0_149, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
 % 162.80/162.90  cnf(c_0_150, negated_conjecture, (~singletonP(X1)|~memberP(esk50_0,esk5_1(X1))|~ssList(esk52_1(esk5_1(X1)))|~ssList(X1)), inference(spm,[status(thm)],[c_0_148, c_0_55])).
 % 162.80/162.90  cnf(c_0_151, negated_conjecture, (ssList(esk52_1(X1))|~ssItem(X1)|~memberP(esk50_0,X1)), inference(split_conjunct,[status(thm)],[c_0_131])).
 % 162.80/162.90  cnf(c_0_152, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_149])).
 % 162.80/162.90  cnf(c_0_153, negated_conjecture, (~singletonP(X1)|~memberP(esk50_0,esk5_1(X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_36])).
 % 162.80/162.90  cnf(c_0_154, plain, (ssList(cons(X1,app(X2,X3)))|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_152]), c_0_58])).
 % 162.80/162.90  cnf(c_0_155, negated_conjecture, (~singletonP(X1)|~memberP(esk50_0,hd(X1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_153, c_0_135])).
 % 162.80/162.90  cnf(c_0_156, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_80]), c_0_35])])).
 % 162.80/162.90  fof(c_0_157, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])).
 % 162.80/162.90  cnf(c_0_158, negated_conjecture, (~memberP(esk50_0,hd(cons(X1,nil)))|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_44]), c_0_156])).
 % 162.80/162.90  cnf(c_0_159, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_157])).
 % 162.80/162.90  cnf(c_0_160, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_157])).
 % 162.80/162.90  cnf(c_0_161, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_157])).
 % 162.80/162.90  cnf(c_0_162, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_131])).
 % 162.80/162.90  cnf(c_0_163, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_131])).
 % 162.80/162.90  cnf(c_0_164, negated_conjecture, (nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_131])).
 % 162.80/162.90  cnf(c_0_165, negated_conjecture, (~memberP(esk50_0,X1)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_130]), c_0_35])])).
 % 162.80/162.90  cnf(c_0_166, plain, (nil=X1|memberP(X1,esk45_1(X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_159]), c_0_160]), c_0_161])).
 % 162.80/162.90  cnf(c_0_167, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_162, c_0_163])).
 % 162.80/162.90  cnf(c_0_168, negated_conjecture, (esk50_0!=nil), inference(rw,[status(thm)],[c_0_164, c_0_163])).
 % 162.80/162.90  cnf(c_0_169, negated_conjecture, (~ssItem(esk45_1(esk50_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_167])]), c_0_168])).
 % 162.80/162.90  cnf(c_0_170, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_160]), c_0_167])]), c_0_168]), ['proof']).
 % 162.80/162.90  # SZS output end CNFRefutation
 % 162.80/162.90  # Training examples: 0 positive, 0 negative
 % 162.80/162.93
 % 162.80/162.93  # -------------------------------------------------
 % 162.80/162.93  # User time                : 159.072 s
 % 162.80/162.93  # System time              : 3.730 s
 % 162.80/162.93  # Total time               : 162.801 s
 % 162.80/162.93  # Maximum resident set size: 1868 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
