   --------------------------------------------------------------------------

                   TSTP Solution File: SWC066+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC066+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n128.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:30 EDT 2016

 % Result     : Theorem 49.35s
 % Output     : CNFRefutation 49.35s
 % Verified   :
 % Statistics : Number of formulae       :  347 (1046505 expanded)
 %              Number of clauses        :  267 (357975 expanded)
 %              Number of leaves         :   40 (256206 expanded)
 %              Depth                    :   37
 %              Number of atoms          : 1046 (13928142 expanded)
 %              Number of equality atoms :  330 (2823381 expanded)
 %              Maximal formula depth    :   21 (   4 average)
 %              Maximal clause size      :   27 (   2 average)
 %              Maximal term depth       :    5 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ~ segmentP(X4,X3)
                     | ~ totalorderedP(X3)
                     | ? [X5] :
                         ( ssList(X5)
                         & neq(X3,X5)
                         & segmentP(X4,X5)
                         & segmentP(X5,X3)
                         & totalorderedP(X5) )
                     | ( ( nil != X2
                         | nil = X1 )
                       & ( ~ neq(X2,nil)
                         | ? [X6] :
                             ( ssList(X6)
                             & neq(X6,nil)
                             & segmentP(X2,X6)
                             & segmentP(X1,X6) ) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( segmentP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & ? [X4] :
                     ( ssList(X4)
                     & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax7)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( frontsegP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & app(X2,X3) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax5)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( rearsegP(X1,X2)
                => rearsegP(app(X3,X1),X2) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax50)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssList(X1)
      => rearsegP(X1,nil) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax51)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => cons(hd(X1),tl(X1)) = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax78)).

 fof(c_0_8,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => tl(cons(X2,X1)) = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

 fof(c_0_9,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_10,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => nil != cons(X2,X1) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax21)).

 fof(c_0_11,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( singletonP(X1)
       <=> ? [X2] :
             ( ssItem(X2)
             & cons(X2,nil) = X1 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax4)).

 fof(c_0_12,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( rearsegP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & app(X3,X2) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax6)).

 fof(c_0_13,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_14,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(X1,nil) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

 fof(c_0_15,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( frontsegP(X1,X2)
                => frontsegP(app(X1,X3),X2) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax43)).

 fof(c_0_16,axiom,(
     ! [X1] :
       ( ssList(X1)
      => frontsegP(X1,nil) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax45)).

 fof(c_0_17,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( app(X3,X2) = app(X1,X2)
                => X3 = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax79)).

 fof(c_0_18,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(nil,X1) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

 fof(c_0_19,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( app(X2,X3) = app(X2,X1)
                => X3 = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax80)).

 fof(c_0_20,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssList(tl(X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

 fof(c_0_21,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil != X1
            => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax86)).

 fof(c_0_22,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_23,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => hd(cons(X2,X1)) = X2 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax23)).

 fof(c_0_24,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil != X1
            => hd(app(X1,X2)) = hd(X1) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax85)).

 fof(c_0_25,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

 fof(c_0_26,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssItem(hd(X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax22)).

 fof(c_0_27,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( frontsegP(cons(X1,X3),cons(X2,X4))
                   <=> ( X1 = X2
                       & frontsegP(X3,X4) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax44)).

 fof(c_0_28,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil = X1
         | ? [X2] :
             ( ssList(X2)
             & ? [X3] :
                 ( ssItem(X3)
                 & cons(X3,X2) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax20)).

 fof(c_0_29,axiom,(
     ! [X1] :
       ( ssList(X1)
      => segmentP(X1,X1) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax55)).

 fof(c_0_30,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssItem(X3)
              => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

 fof(c_0_31,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_32,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( segmentP(nil,X1)
       <=> nil = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax58)).

 fof(c_0_33,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( segmentP(X1,X2)
                    => segmentP(app(app(X3,X1),X4),X2) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax56)).

 fof(c_0_34,axiom,(
     ! [X1] :
       ( ssList(X1)
      => frontsegP(X1,X1) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax42)).

 fof(c_0_35,axiom,(
     ! [X1] :
       ( ssList(X1)
      => rearsegP(X1,X1) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax49)).

 fof(c_0_36,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( rearsegP(nil,X1)
       <=> nil = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax52)).

 fof(c_0_37,axiom,(
     ! [X1] :
       ( ssList(X1)
      => segmentP(X1,nil) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax57)).

 fof(c_0_38,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => totalorderedP(cons(X1,nil)) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax65)).

 fof(c_0_39,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) != X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax18)).

 fof(c_0_40,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ~ segmentP(X4,X3)
                       | ~ totalorderedP(X3)
                       | ? [X5] :
                           ( ssList(X5)
                           & neq(X3,X5)
                           & segmentP(X4,X5)
                           & segmentP(X5,X3)
                           & totalorderedP(X5) )
                       | ( ( nil != X2
                           | nil = X1 )
                         & ( ~ neq(X2,nil)
                           | ? [X6] :
                               ( ssList(X6)
                               & neq(X6,nil)
                               & segmentP(X2,X6)
                               & segmentP(X1,X6) ) ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_41,negated_conjecture,(
     ! [X11,X12] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & segmentP(esk51_0,esk50_0)
       & totalorderedP(esk50_0)
       & ( ~ ssList(X11)
         | ~ neq(esk50_0,X11)
         | ~ segmentP(esk51_0,X11)
         | ~ segmentP(X11,esk50_0)
         | ~ totalorderedP(X11) )
       & ( neq(esk49_0,nil)
         | nil = esk49_0 )
       & ( ~ ssList(X12)
         | ~ neq(X12,nil)
         | ~ segmentP(esk49_0,X12)
         | ~ segmentP(esk48_0,X12)
         | nil = esk49_0 )
       & ( neq(esk49_0,nil)
         | nil != esk48_0 )
       & ( ~ ssList(X12)
         | ~ neq(X12,nil)
         | ~ segmentP(esk49_0,X12)
         | ~ segmentP(esk48_0,X12)
         | nil != esk48_0 ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])])])])).

 fof(c_0_42,plain,(
     ! [X5,X6,X9,X10] :
       ( ( ssList(esk8_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ssList(esk9_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ~ ssList(X9)
         | ~ ssList(X10)
         | app(app(X9,X6),X10) != X5
         | segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

 cnf(c_0_43,negated_conjecture,
     ( segmentP(esk51_0,esk50_0) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_41])).

 cnf(c_0_44,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_41])).

 cnf(c_0_45,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_41])).

 fof(c_0_46,plain,(
     ! [X4,X5,X7] :
       ( ( ssList(esk6_2(X4,X5))
         | ~ frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( app(X5,esk6_2(X4,X5)) = X4
         | ~ frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( ~ ssList(X7)
         | app(X5,X7) != X4
         | frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).

 fof(c_0_47,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

 fof(c_0_48,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

 cnf(c_0_49,plain,
     ( app(app(esk8_2(X1,X2),X2),esk9_2(X1,X2)) = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ segmentP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[7]c_0_42])).

 cnf(c_0_50,negated_conjecture,
     ( segmentP(esk49_0,esk48_0) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[8]c_0_43,c_0_44]),[9]c_0_45])).

 cnf(c_0_51,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[10]c_0_41])).

 cnf(c_0_52,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_41])).

 cnf(c_0_53,plain,
     ( ssList(esk9_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ segmentP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_42])).

 cnf(c_0_54,plain,
     ( ssList(esk8_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ segmentP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[13]c_0_42])).

 cnf(c_0_55,plain,
     ( frontsegP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(X2,X3) != X1
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_46])).

 cnf(c_0_56,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_47])).

 cnf(c_0_57,plain,
     ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_48])).

 cnf(c_0_58,negated_conjecture,
     ( app(app(esk8_2(esk49_0,esk48_0),esk48_0),esk9_2(esk49_0,esk48_0)) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_49,c_0_50]),c_0_51]),c_0_52])])).

 cnf(c_0_59,negated_conjecture,
     ( ssList(esk9_2(esk49_0,esk48_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_53,c_0_50]),c_0_51]),c_0_52])])).

 cnf(c_0_60,negated_conjecture,
     ( ssList(esk8_2(esk49_0,esk48_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[19]c_0_54,c_0_50]),c_0_51]),c_0_52])])).

 cnf(c_0_61,plain,
     ( frontsegP(app(X1,X2),X1)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_55]),[20]c_0_56])).

 cnf(c_0_62,negated_conjecture,
     ( app(esk8_2(esk49_0,esk48_0),app(esk48_0,esk9_2(esk49_0,esk48_0))) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[21]c_0_57,c_0_58]),c_0_59]),c_0_51]),c_0_60])])).

 fof(c_0_63,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | ~ rearsegP(X4,X5)
       | rearsegP(app(X6,X4),X5) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

 fof(c_0_64,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | rearsegP(X2,nil) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).

 cnf(c_0_65,negated_conjecture,
     ( frontsegP(esk49_0,esk8_2(esk49_0,esk48_0))
     | ~ ssList(app(esk48_0,esk9_2(esk49_0,esk48_0))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[22]c_0_61,c_0_62]),c_0_60])])).

 fof(c_0_66,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | cons(hd(X2),tl(X2)) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

 fof(c_0_67,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | tl(cons(X4,X3)) = X3 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

 fof(c_0_68,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

 fof(c_0_69,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | nil != cons(X4,X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).

 fof(c_0_70,plain,(
     ! [X3,X5] :
       ( ( ssItem(esk5_1(X3))
         | ~ singletonP(X3)
         | ~ ssList(X3) )
       & ( cons(esk5_1(X3),nil) = X3
         | ~ singletonP(X3)
         | ~ ssList(X3) )
       & ( ~ ssItem(X5)
         | cons(X5,nil) != X3
         | singletonP(X3)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).

 fof(c_0_71,plain,(
     ! [X4,X5,X7] :
       ( ( ssList(esk7_2(X4,X5))
         | ~ rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( app(esk7_2(X4,X5),X5) = X4
         | ~ rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( ~ ssList(X7)
         | app(X7,X5) != X4
         | rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).

 cnf(c_0_72,plain,
     ( rearsegP(app(X1,X2),X3)
     | ~ rearsegP(X2,X3)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[23]c_0_63])).

 cnf(c_0_73,plain,
     ( rearsegP(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[24]c_0_64])).

 cnf(c_0_74,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[25]c_0_13])).

 cnf(c_0_75,plain,
     ( app(X2,esk6_2(X1,X2)) = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ frontsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[26]c_0_46])).

 cnf(c_0_76,negated_conjecture,
     ( frontsegP(esk49_0,esk8_2(esk49_0,esk48_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_65,c_0_56]),c_0_59]),c_0_51])])).

 cnf(c_0_77,plain,
     ( ssList(esk6_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ frontsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[28]c_0_46])).

 cnf(c_0_78,plain,
     ( cons(hd(X1),tl(X1)) = X1
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[29]c_0_66])).

 cnf(c_0_79,plain,
     ( tl(cons(X1,X2)) = X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[30]c_0_67])).

 cnf(c_0_80,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[31]c_0_68])).

 cnf(c_0_81,plain,
     ( nil != cons(X1,X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[32]c_0_69])).

 cnf(c_0_82,plain,
     ( singletonP(X1)
     | ~ ssList(X1)
     | cons(X2,nil) != X1
     | ~ ssItem(X2) ),
     inference(split_conjunct,[status(thm)],[[33]c_0_70])).

 cnf(c_0_83,plain,
     ( rearsegP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(X3,X2) != X1
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[34]c_0_71])).

 cnf(c_0_84,plain,
     ( rearsegP(app(X1,X2),nil)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[35]c_0_72,c_0_73]),c_0_74])])).

 cnf(c_0_85,negated_conjecture,
     ( app(esk8_2(esk49_0,esk48_0),esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[36]c_0_75,c_0_76]),c_0_60]),c_0_52])])).

 cnf(c_0_86,negated_conjecture,
     ( ssList(esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[37]c_0_77,c_0_76]),c_0_60]),c_0_52])])).

 cnf(c_0_87,plain,
     ( cons(hd(cons(X1,X2)),X2) = cons(X1,X2)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[38]c_0_78,c_0_79]),c_0_80]),[39]c_0_81])).

 cnf(c_0_88,plain,
     ( cons(esk5_1(X1),nil) = X1
     | ~ ssList(X1)
     | ~ singletonP(X1) ),
     inference(split_conjunct,[status(thm)],[[40]c_0_70])).

 cnf(c_0_89,plain,
     ( ssItem(esk5_1(X1))
     | ~ ssList(X1)
     | ~ singletonP(X1) ),
     inference(split_conjunct,[status(thm)],[[41]c_0_70])).

 cnf(c_0_90,plain,
     ( singletonP(cons(X1,nil))
     | ~ ssList(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(er,[status(thm)],[[42]c_0_82])).

 cnf(c_0_91,plain,
     ( rearsegP(app(X1,X2),X2)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_83]),[43]c_0_56])).

 fof(c_0_92,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(X2,nil) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).

 cnf(c_0_93,plain,
     ( app(esk7_2(X1,X2),X2) = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ rearsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[44]c_0_71])).

 cnf(c_0_94,negated_conjecture,
     ( rearsegP(esk49_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[45]c_0_84,c_0_85]),c_0_86]),c_0_60])])).

 cnf(c_0_95,plain,
     ( ssList(esk7_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ rearsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[46]c_0_71])).

 cnf(c_0_96,plain,
     ( cons(hd(X1),nil) = X1
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[47]c_0_87,c_0_88]),c_0_74])]),[48]c_0_89])).

 cnf(c_0_97,plain,
     ( singletonP(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[49]c_0_90,c_0_80]),c_0_74])])).

 cnf(c_0_98,negated_conjecture,
     ( rearsegP(esk49_0,esk9_2(esk49_0,esk48_0))
     | ~ ssList(app(esk8_2(esk49_0,esk48_0),esk48_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[50]c_0_91,c_0_58]),c_0_59])])).

 cnf(c_0_99,plain,
     ( app(X1,nil) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[51]c_0_92])).

 cnf(c_0_100,negated_conjecture,
     ( app(esk7_2(esk49_0,nil),nil) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[52]c_0_93,c_0_94]),c_0_74]),c_0_52])])).

 cnf(c_0_101,negated_conjecture,
     ( ssList(esk7_2(esk49_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[53]c_0_95,c_0_94]),c_0_74]),c_0_52])])).

 fof(c_0_102,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | ~ frontsegP(X4,X5)
       | frontsegP(app(X4,X6),X5) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).

 fof(c_0_103,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | frontsegP(X2,nil) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).

 cnf(c_0_104,plain,
     ( cons(hd(cons(X1,nil)),nil) = cons(X1,nil)
     | ~ ssList(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(spm,[status(thm)],[[54]c_0_96,[55]c_0_97])).

 fof(c_0_105,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(X6,X5) != app(X4,X5)
       | X6 = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).

 fof(c_0_106,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(nil,X2) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])).

 fof(c_0_107,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(X5,X6) != app(X5,X4)
       | X6 = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).

 fof(c_0_108,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssList(tl(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])).

 fof(c_0_109,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | nil = X3
       | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).

 fof(c_0_110,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).

 cnf(c_0_111,negated_conjecture,
     ( rearsegP(esk49_0,esk9_2(esk49_0,esk48_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[56]c_0_98,c_0_56]),c_0_51]),c_0_60])])).

 cnf(c_0_112,negated_conjecture,
     ( esk7_2(esk49_0,nil) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[57]c_0_99,c_0_100]),c_0_101])])).

 cnf(c_0_113,plain,
     ( frontsegP(app(X1,X2),X3)
     | ~ frontsegP(X1,X3)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[58]c_0_102])).

 cnf(c_0_114,plain,
     ( frontsegP(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[59]c_0_103])).

 cnf(c_0_115,plain,
     ( cons(hd(cons(X1,nil)),nil) = cons(X1,nil)
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[60]c_0_104,c_0_80]),c_0_74])])).

 fof(c_0_116,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | hd(cons(X4,X3)) = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).

 fof(c_0_117,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | nil = X3
       | hd(app(X3,X4)) = hd(X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).

 cnf(c_0_118,plain,
     ( X1 = X2
     | app(X1,X3) != app(X2,X3)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[61]c_0_105])).

 cnf(c_0_119,plain,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[62]c_0_106])).

 cnf(c_0_120,plain,
     ( X1 = X2
     | app(X3,X1) != app(X3,X2)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[63]c_0_107])).

 cnf(c_0_121,plain,
     ( ssList(tl(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[64]c_0_108])).

 cnf(c_0_122,plain,
     ( tl(app(X1,X2)) = app(tl(X1),X2)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[65]c_0_109])).

 cnf(c_0_123,plain,
     ( nil = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[66]c_0_110])).

 cnf(c_0_124,negated_conjecture,
     ( app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),esk9_2(esk49_0,esk48_0)) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[67]c_0_93,c_0_111]),c_0_59]),c_0_52])])).

 cnf(c_0_125,negated_conjecture,
     ( ssList(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[68]c_0_95,c_0_111]),c_0_59]),c_0_52])])).

 cnf(c_0_126,negated_conjecture,
     ( app(esk49_0,nil) = esk49_0 ),
     inference(rw,[status(thm)],[[69]c_0_100,[70]c_0_112])).

 cnf(c_0_127,plain,
     ( frontsegP(app(X1,X2),nil)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[71]c_0_113,c_0_114]),c_0_74])])).

 fof(c_0_128,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).

 cnf(c_0_129,plain,
     ( ssList(cons(X1,nil))
     | ~ ssItem(hd(cons(X1,nil)))
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[72]c_0_80,c_0_115]),c_0_74])])).

 cnf(c_0_130,plain,
     ( hd(cons(X1,X2)) = X1
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[73]c_0_116])).

 cnf(c_0_131,plain,
     ( hd(app(X1,X2)) = hd(X1)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[74]c_0_117])).

 cnf(c_0_132,plain,
     ( X1 = nil
     | app(X1,X2) != X2
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[75]c_0_118,c_0_119]),c_0_74])])).

 cnf(c_0_133,plain,
     ( X1 = nil
     | app(X2,X1) != X2
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[76]c_0_120,c_0_99]),c_0_74])])).

 cnf(c_0_134,plain,
     ( nil = X1
     | ssList(app(tl(X1),X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[77]c_0_121,c_0_122]),c_0_56]),[78]c_0_123])).

 cnf(c_0_135,negated_conjecture,
     ( app(tl(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))),esk9_2(esk49_0,esk48_0)) = tl(esk49_0)
     | esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[79]c_0_122,c_0_124]),c_0_59]),c_0_125])])).

 cnf(c_0_136,negated_conjecture,
     ( app(tl(esk49_0),nil) = tl(esk49_0)
     | esk49_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[80]c_0_122,c_0_126]),c_0_74]),c_0_52])])).

 cnf(c_0_137,negated_conjecture,
     ( frontsegP(esk49_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[81]c_0_127,c_0_85]),c_0_86]),c_0_60])])).

 cnf(c_0_138,plain,
     ( cons(X1,X2) = app(cons(X1,nil),X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[82]c_0_128])).

 cnf(c_0_139,plain,
     ( ssList(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[83]c_0_129,c_0_130]),c_0_74])])).

 cnf(c_0_140,plain,
     ( cons(hd(X1),tl(app(X1,X2))) = app(X1,X2)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[84]c_0_78,c_0_131]),c_0_56]),[85]c_0_123])).

 cnf(c_0_141,negated_conjecture,
     ( app(esk8_2(esk49_0,esk48_0),esk48_0) = nil
     | esk9_2(esk49_0,esk48_0) != esk49_0
     | ~ ssList(app(esk8_2(esk49_0,esk48_0),esk48_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[86]c_0_132,c_0_58]),c_0_59])])).

 cnf(c_0_142,negated_conjecture,
     ( esk9_2(esk49_0,esk48_0) = nil
     | esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)) != esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[87]c_0_133,c_0_124]),c_0_125]),c_0_59])])).

 cnf(c_0_143,negated_conjecture,
     ( esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)) = nil
     | ssList(tl(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[88]c_0_134,c_0_135]),c_0_59]),c_0_125])])).

 cnf(c_0_144,negated_conjecture,
     ( esk49_0 = nil
     | ssList(tl(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[89]c_0_134,c_0_136]),c_0_74]),c_0_52])])).

 cnf(c_0_145,plain,
     ( app(tl(X1),nil) = tl(X1)
     | nil = X1
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[90]c_0_122,c_0_99]),c_0_74])])).

 fof(c_0_146,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssItem(hd(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])).

 cnf(c_0_147,negated_conjecture,
     ( app(nil,esk6_2(esk49_0,nil)) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[91]c_0_75,c_0_137]),c_0_74]),c_0_52])])).

 cnf(c_0_148,negated_conjecture,
     ( ssList(esk6_2(esk49_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[92]c_0_77,c_0_137]),c_0_74]),c_0_52])])).

 fof(c_0_149,plain,(
     ! [X5,X6,X7,X8] :
       ( ( X5 = X6
         | ~ frontsegP(cons(X5,X7),cons(X6,X8))
         | ~ ssList(X8)
         | ~ ssList(X7)
         | ~ ssItem(X6)
         | ~ ssItem(X5) )
       & ( frontsegP(X7,X8)
         | ~ frontsegP(cons(X5,X7),cons(X6,X8))
         | ~ ssList(X8)
         | ~ ssList(X7)
         | ~ ssItem(X6)
         | ~ ssItem(X5) )
       & ( X5 != X6
         | ~ frontsegP(X7,X8)
         | frontsegP(cons(X5,X7),cons(X6,X8))
         | ~ ssList(X8)
         | ~ ssList(X7)
         | ~ ssItem(X6)
         | ~ ssItem(X5) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])).

 fof(c_0_150,plain,(
     ! [X4] :
       ( ( ssList(esk44_1(X4))
         | nil = X4
         | ~ ssList(X4) )
       & ( ssItem(esk45_1(X4))
         | nil = X4
         | ~ ssList(X4) )
       & ( cons(esk45_1(X4),esk44_1(X4)) = X4
         | nil = X4
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])])).

 cnf(c_0_151,plain,
     ( frontsegP(cons(X1,X2),cons(X1,nil))
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[93]c_0_61,c_0_138]),[94]c_0_139])).

 cnf(c_0_152,negated_conjecture,
     ( cons(hd(esk49_0),tl(esk49_0)) = esk49_0
     | esk49_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[95]c_0_140,c_0_126]),c_0_74]),c_0_52])])).

 cnf(c_0_153,negated_conjecture,
     ( app(esk8_2(esk49_0,esk48_0),esk48_0) = nil
     | esk9_2(esk49_0,esk48_0) != esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[96]c_0_141,c_0_56]),c_0_51]),c_0_60])])).

 cnf(c_0_154,negated_conjecture,
     ( esk9_2(esk49_0,esk48_0) = nil
     | ssList(tl(esk49_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[97]c_0_142,c_0_143]),[98]c_0_144])).

 cnf(c_0_155,negated_conjecture,
     ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil
     | esk8_2(esk49_0,esk48_0) != esk49_0
     | ~ ssList(app(esk48_0,esk9_2(esk49_0,esk48_0))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[99]c_0_133,c_0_62]),c_0_60])])).

 cnf(c_0_156,plain,
     ( app(app(tl(X1),X2),nil) = app(tl(X1),X2)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[100]c_0_145,c_0_122]),c_0_56]),[101]c_0_123])).

 cnf(c_0_157,negated_conjecture,
     ( app(tl(esk8_2(esk49_0,esk48_0)),esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))) = tl(esk49_0)
     | esk8_2(esk49_0,esk48_0) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[102]c_0_122,c_0_85]),c_0_86]),c_0_60])])).

 cnf(c_0_158,plain,
     ( ssItem(hd(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[103]c_0_146])).

 cnf(c_0_159,negated_conjecture,
     ( hd(esk8_2(esk49_0,esk48_0)) = hd(esk49_0)
     | esk8_2(esk49_0,esk48_0) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[104]c_0_131,c_0_85]),c_0_86]),c_0_60])])).

 fof(c_0_160,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | segmentP(X2,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])).

 cnf(c_0_161,plain,
     ( segmentP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(app(X3,X2),X4) != X1
     | ~ ssList(X4)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[105]c_0_42])).

 cnf(c_0_162,negated_conjecture,
     ( esk6_2(esk49_0,nil) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[106]c_0_119,c_0_147]),c_0_148])])).

 cnf(c_0_163,plain,
     ( X1 = X2
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | ~ ssList(X4)
     | ~ frontsegP(cons(X1,X3),cons(X2,X4)) ),
     inference(split_conjunct,[status(thm)],[[107]c_0_149])).

 cnf(c_0_164,plain,
     ( nil = X1
     | cons(esk45_1(X1),esk44_1(X1)) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[108]c_0_150])).

 cnf(c_0_165,plain,
     ( nil = X1
     | ssItem(esk45_1(X1))
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[109]c_0_150])).

 cnf(c_0_166,plain,
     ( nil = X1
     | ssList(esk44_1(X1))
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[110]c_0_150])).

 cnf(c_0_167,negated_conjecture,
     ( esk49_0 = nil
     | frontsegP(esk49_0,cons(hd(esk49_0),nil))
     | ~ ssItem(hd(esk49_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[111]c_0_151,c_0_152]),[112]c_0_144])).

 cnf(c_0_168,plain,
     ( nil = X2
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[113]c_0_110])).

 cnf(c_0_169,negated_conjecture,
     ( app(esk8_2(esk49_0,esk48_0),esk48_0) = nil
     | ssList(tl(esk49_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[114]c_0_153,c_0_154]),[115]c_0_144])).

 fof(c_0_170,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssItem(X6)
       | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).

 cnf(c_0_171,negated_conjecture,
     ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil
     | esk8_2(esk49_0,esk48_0) != esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[116]c_0_155,c_0_56]),c_0_59]),c_0_51])])).

 cnf(c_0_172,negated_conjecture,
     ( app(tl(esk49_0),nil) = tl(esk49_0)
     | esk8_2(esk49_0,esk48_0) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[117]c_0_156,c_0_157]),c_0_86]),c_0_60])])).

 cnf(c_0_173,negated_conjecture,
     ( esk8_2(esk49_0,esk48_0) = nil
     | ssItem(hd(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[118]c_0_158,c_0_159]),c_0_60])])).

 cnf(c_0_174,negated_conjecture,
     ( nil = esk49_0
     | ~ segmentP(esk48_0,X1)
     | ~ segmentP(esk49_0,X1)
     | ~ neq(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[119]c_0_41])).

 cnf(c_0_175,plain,
     ( segmentP(X1,X1)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[120]c_0_160])).

 fof(c_0_176,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])])).

 cnf(c_0_177,plain,
     ( segmentP(app(app(X1,X2),X3),X2)
     | ~ ssList(app(app(X1,X2),X3))
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(er,[status(thm)],[[121]c_0_161])).

 cnf(c_0_178,negated_conjecture,
     ( app(nil,esk49_0) = esk49_0 ),
     inference(rw,[status(thm)],[[122]c_0_147,[123]c_0_162])).

 cnf(c_0_179,plain,
     ( esk45_1(X1) = X2
     | nil = X1
     | ~ frontsegP(X1,cons(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[124]c_0_163,c_0_164]),c_0_165]),[125]c_0_166])).

 cnf(c_0_180,negated_conjecture,
     ( esk49_0 = nil
     | frontsegP(esk49_0,cons(hd(esk49_0),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[126]c_0_167,c_0_158]),c_0_52])])).

 cnf(c_0_181,negated_conjecture,
     ( esk48_0 = nil
     | ssList(tl(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[127]c_0_168,c_0_169]),c_0_51]),c_0_60])])).

 cnf(c_0_182,plain,
     ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[128]c_0_170])).

 cnf(c_0_183,negated_conjecture,
     ( app(tl(esk49_0),nil) = tl(esk49_0)
     | app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[129]c_0_171,c_0_172]),[130]c_0_136])).

 cnf(c_0_184,negated_conjecture,
     ( app(app(nil,esk48_0),esk9_2(esk49_0,esk48_0)) = esk49_0
     | ssItem(hd(esk49_0)) ),
     inference(spm,[status(thm)],[[131]c_0_58,[132]c_0_173])).

 cnf(c_0_185,negated_conjecture,
     ( esk49_0 = nil
     | ~ neq(esk48_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[133]c_0_174,c_0_175]),c_0_50]),c_0_51])])).

 cnf(c_0_186,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[134]c_0_176])).

 fof(c_0_187,plain,(
     ! [X2] :
       ( ( ~ segmentP(nil,X2)
         | nil = X2
         | ~ ssList(X2) )
       & ( nil != X2
         | segmentP(nil,X2)
         | ~ ssList(X2) ) ) ),
     inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])).

 cnf(c_0_188,negated_conjecture,
     ( segmentP(app(esk49_0,X1),esk49_0)
     | ~ ssList(app(esk49_0,X1))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[135]c_0_177,c_0_178]),c_0_74]),c_0_52])])).

 cnf(c_0_189,negated_conjecture,
     ( esk45_1(esk49_0) = hd(esk49_0)
     | esk49_0 = nil
     | ~ ssItem(hd(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[136]c_0_179,c_0_180]),c_0_74]),c_0_52])])).

 cnf(c_0_190,negated_conjecture,
     ( esk8_2(esk49_0,esk48_0) = nil
     | ssList(tl(esk49_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[137]c_0_132,c_0_169]),c_0_51]),c_0_60])]),[138]c_0_181])).

 cnf(c_0_191,plain,
     ( cons(X1,app(X2,X3)) != nil
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[139]c_0_123,c_0_182]),c_0_80]),[140]c_0_81])).

 cnf(c_0_192,negated_conjecture,
     ( app(tl(esk49_0),nil) = tl(esk49_0)
     | esk48_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[141]c_0_123,c_0_183]),c_0_59]),c_0_51])])).

 cnf(c_0_193,plain,
     ( cons(hd(app(X1,X2)),app(tl(X1),X2)) = app(X1,X2)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[142]c_0_78,c_0_122]),c_0_56]),[143]c_0_123])).

 cnf(c_0_194,negated_conjecture,
     ( app(nil,esk48_0) = nil
     | ssItem(hd(esk49_0))
     | esk49_0 != nil
     | ~ ssList(app(nil,esk48_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[144]c_0_123,c_0_184]),c_0_59])])).

 cnf(c_0_195,negated_conjecture,
     ( esk48_0 = nil
     | esk49_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[145]c_0_185,c_0_186]),c_0_74]),c_0_51])])).

 fof(c_0_196,plain,(
     ! [X5,X6,X7,X8] :
       ( ~ ssList(X5)
       | ~ ssList(X6)
       | ~ ssList(X7)
       | ~ ssList(X8)
       | ~ segmentP(X5,X6)
       | segmentP(app(app(X7,X5),X8),X6) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])).

 cnf(c_0_197,plain,
     ( segmentP(nil,X1)
     | ~ ssList(X1)
     | nil != X1 ),
     inference(split_conjunct,[status(thm)],[[146]c_0_187])).

 cnf(c_0_198,negated_conjecture,
     ( segmentP(esk49_0,esk49_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[147]c_0_188,c_0_99]),c_0_52]),c_0_74])])).

 cnf(c_0_199,plain,
     ( nil = X1
     | app(X2,app(X3,X1)) != nil
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[148]c_0_168,c_0_57]),[149]c_0_56])).

 cnf(c_0_200,negated_conjecture,
     ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = esk49_0
     | ssItem(hd(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[150]c_0_184,c_0_119]),c_0_51])])).

 cnf(c_0_201,negated_conjecture,
     ( esk45_1(esk49_0) = hd(esk49_0)
     | esk49_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[151]c_0_189,c_0_158]),c_0_52])])).

 cnf(c_0_202,negated_conjecture,
     ( X1 = esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))
     | app(X1,esk9_2(esk49_0,esk48_0)) != esk49_0
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[152]c_0_118,c_0_124]),c_0_59]),c_0_125])])).

 cnf(c_0_203,negated_conjecture,
     ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil
     | ssList(tl(esk49_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[153]c_0_171,c_0_190]),[154]c_0_144])).

 cnf(c_0_204,negated_conjecture,
     ( esk48_0 = nil
     | cons(X1,tl(esk49_0)) != nil
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[155]c_0_191,c_0_192]),c_0_74])]),[156]c_0_181])).

 cnf(c_0_205,negated_conjecture,
     ( cons(hd(esk49_0),tl(esk49_0)) = esk49_0
     | esk8_2(esk49_0,esk48_0) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[157]c_0_193,c_0_157]),c_0_85]),c_0_85]),c_0_86]),c_0_60])])).

 cnf(c_0_206,negated_conjecture,
     ( esk48_0 = nil
     | ssItem(hd(esk49_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[158]c_0_194,c_0_119]),c_0_51])]),[159]c_0_195])).

 cnf(c_0_207,plain,
     ( segmentP(app(app(X1,X2),X3),X4)
     | ~ segmentP(X2,X4)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssList(X4)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[160]c_0_196])).

 cnf(c_0_208,plain,
     ( segmentP(nil,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_197]),c_0_74])])).

 fof(c_0_209,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | frontsegP(X2,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])).

 cnf(c_0_210,negated_conjecture,
     ( frontsegP(esk49_0,esk49_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[161]c_0_61,c_0_126]),c_0_74]),c_0_52])])).

 cnf(c_0_211,plain,
     ( frontsegP(app(app(X1,X2),X3),X1)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[162]c_0_113,c_0_61]),[163]c_0_56])).

 cnf(c_0_212,negated_conjecture,
     ( app(app(esk8_2(esk49_0,esk49_0),esk49_0),esk9_2(esk49_0,esk49_0)) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[164]c_0_49,c_0_198]),c_0_52])])).

 cnf(c_0_213,negated_conjecture,
     ( ssList(esk8_2(esk49_0,esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[165]c_0_54,c_0_198]),c_0_52])])).

 cnf(c_0_214,negated_conjecture,
     ( ssList(esk9_2(esk49_0,esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[166]c_0_53,c_0_198]),c_0_52])])).

 cnf(c_0_215,negated_conjecture,
     ( esk6_2(esk49_0,esk8_2(esk49_0,esk48_0)) = nil
     | app(X1,esk49_0) != nil
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[167]c_0_199,c_0_85]),c_0_86]),c_0_60])])).

 cnf(c_0_216,negated_conjecture,
     ( app(esk8_2(esk49_0,esk48_0),esk49_0) = esk49_0
     | ssItem(hd(esk49_0)) ),
     inference(spm,[status(thm)],[[168]c_0_62,[169]c_0_200])).

 cnf(c_0_217,negated_conjecture,
     ( esk49_0 = nil
     | ssItem(hd(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[170]c_0_165,c_0_201]),c_0_52])])).

 cnf(c_0_218,negated_conjecture,
     ( esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)) = esk48_0
     | ssList(tl(esk49_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[171]c_0_202,c_0_203]),c_0_51])]),[172]c_0_144])).

 cnf(c_0_219,negated_conjecture,
     ( esk8_2(esk49_0,esk48_0) = nil
     | esk48_0 = nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[173]c_0_204,c_0_205]),c_0_206]),[174]c_0_195])).

 cnf(c_0_220,plain,
     ( segmentP(app(app(X1,nil),X2),nil)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[175]c_0_207,c_0_208]),c_0_74])])).

 cnf(c_0_221,plain,
     ( frontsegP(X1,X1)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[176]c_0_209])).

 cnf(c_0_222,negated_conjecture,
     ( app(esk49_0,esk6_2(esk49_0,esk49_0)) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[177]c_0_75,c_0_210]),c_0_52])])).

 cnf(c_0_223,negated_conjecture,
     ( ssList(esk6_2(esk49_0,esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[178]c_0_77,c_0_210]),c_0_52])])).

 cnf(c_0_224,negated_conjecture,
     ( frontsegP(esk49_0,esk8_2(esk49_0,esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[179]c_0_211,c_0_212]),c_0_213]),c_0_214]),c_0_52])])).

 cnf(c_0_225,negated_conjecture,
     ( esk6_2(esk49_0,esk8_2(esk49_0,esk48_0)) = nil
     | ssItem(hd(esk49_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[180]c_0_215,c_0_216]),c_0_60])]),[181]c_0_217])).

 cnf(c_0_226,negated_conjecture,
     ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = esk49_0
     | ssList(tl(esk49_0)) ),
     inference(spm,[status(thm)],[[182]c_0_124,[183]c_0_218])).

 cnf(c_0_227,negated_conjecture,
     ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil
     | esk48_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[184]c_0_171,c_0_219]),[185]c_0_195])).

 cnf(c_0_228,negated_conjecture,
     ( segmentP(app(esk49_0,X1),nil)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[186]c_0_220,c_0_126]),c_0_52])])).

 cnf(c_0_229,plain,
     ( app(X1,esk6_2(X1,X1)) = X1
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[187]c_0_75,[188]c_0_221])).

 cnf(c_0_230,negated_conjecture,
     ( esk6_2(esk49_0,esk49_0) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[189]c_0_133,c_0_222]),c_0_52]),c_0_223])])).

 fof(c_0_231,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | rearsegP(X2,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])).

 fof(c_0_232,plain,(
     ! [X2] :
       ( ( ~ rearsegP(nil,X2)
         | nil = X2
         | ~ ssList(X2) )
       & ( nil != X2
         | rearsegP(nil,X2)
         | ~ ssList(X2) ) ) ),
     inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).

 cnf(c_0_233,negated_conjecture,
     ( app(esk8_2(esk49_0,esk49_0),app(esk49_0,esk9_2(esk49_0,esk49_0))) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[190]c_0_57,c_0_212]),c_0_214]),c_0_52]),c_0_213])])).

 cnf(c_0_234,negated_conjecture,
     ( app(esk8_2(esk49_0,esk49_0),esk6_2(esk49_0,esk8_2(esk49_0,esk49_0))) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[191]c_0_75,c_0_224]),c_0_213]),c_0_52])])).

 cnf(c_0_235,negated_conjecture,
     ( ssList(esk6_2(esk49_0,esk8_2(esk49_0,esk49_0))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[192]c_0_77,c_0_224]),c_0_213]),c_0_52])])).

 cnf(c_0_236,negated_conjecture,
     ( app(esk8_2(esk49_0,esk48_0),nil) = esk49_0
     | ssItem(hd(esk49_0)) ),
     inference(spm,[status(thm)],[[193]c_0_85,[194]c_0_225])).

 cnf(c_0_237,negated_conjecture,
     ( app(esk8_2(esk49_0,esk48_0),esk49_0) = esk49_0
     | ssList(tl(esk49_0)) ),
     inference(spm,[status(thm)],[[195]c_0_62,[196]c_0_226])).

 cnf(c_0_238,negated_conjecture,
     ( esk48_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[197]c_0_123,c_0_227]),c_0_59]),c_0_51])])).

 cnf(c_0_239,negated_conjecture,
     ( segmentP(esk49_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[198]c_0_228,c_0_229]),c_0_230]),c_0_74]),c_0_52])])).

 cnf(c_0_240,plain,
     ( rearsegP(X1,X1)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[199]c_0_231])).

 cnf(c_0_241,plain,
     ( rearsegP(nil,X1)
     | ~ ssList(X1)
     | nil != X1 ),
     inference(split_conjunct,[status(thm)],[[200]c_0_232])).

 cnf(c_0_242,negated_conjecture,
     ( segmentP(app(esk49_0,X1),esk9_2(esk49_0,esk48_0))
     | ~ ssList(app(esk49_0,X1))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[201]c_0_177,c_0_124]),c_0_125]),c_0_59])])).

 cnf(c_0_243,negated_conjecture,
     ( app(esk49_0,esk9_2(esk49_0,esk49_0)) = nil
     | esk8_2(esk49_0,esk49_0) != esk49_0
     | ~ ssList(app(esk49_0,esk9_2(esk49_0,esk49_0))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[202]c_0_133,c_0_233]),c_0_213])])).

 cnf(c_0_244,negated_conjecture,
     ( hd(esk8_2(esk49_0,esk49_0)) = hd(esk49_0)
     | esk8_2(esk49_0,esk49_0) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[203]c_0_131,c_0_234]),c_0_235]),c_0_213])])).

 cnf(c_0_245,negated_conjecture,
     ( frontsegP(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[204]c_0_61,c_0_124]),c_0_59]),c_0_125])])).

 cnf(c_0_246,negated_conjecture,
     ( esk8_2(esk49_0,esk48_0) = esk49_0
     | ssItem(hd(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[205]c_0_99,c_0_236]),c_0_60])])).

 cnf(c_0_247,negated_conjecture,
     ( esk6_2(esk49_0,esk8_2(esk49_0,esk49_0)) = nil
     | app(X1,esk49_0) != nil
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[206]c_0_199,c_0_234]),c_0_235]),c_0_213])])).

 cnf(c_0_248,negated_conjecture,
     ( app(esk8_2(esk49_0,nil),esk49_0) = esk49_0
     | ssList(tl(esk49_0)) ),
     inference(rw,[status(thm)],[[207]c_0_237,[208]c_0_238])).

 cnf(c_0_249,negated_conjecture,
     ( ssList(esk8_2(esk49_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[209]c_0_54,c_0_239]),c_0_74]),c_0_52])])).

 cnf(c_0_250,plain,
     ( app(esk7_2(X1,X1),X1) = X1
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[210]c_0_93,[211]c_0_240])).

 cnf(c_0_251,plain,
     ( ssList(esk7_2(X1,X1))
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[212]c_0_95,[213]c_0_240])).

 cnf(c_0_252,plain,
     ( rearsegP(nil,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_241]),c_0_74])])).

 cnf(c_0_253,plain,
     ( rearsegP(cons(X1,X2),X2)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[214]c_0_91,c_0_138]),[215]c_0_139])).

 cnf(c_0_254,negated_conjecture,
     ( segmentP(app(esk49_0,X1),esk9_2(esk49_0,esk48_0))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[216]c_0_242,c_0_56]),c_0_52])])).

 cnf(c_0_255,negated_conjecture,
     ( app(esk49_0,esk9_2(esk49_0,esk49_0)) = nil
     | esk8_2(esk49_0,esk49_0) != esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[217]c_0_243,c_0_56]),c_0_214]),c_0_52])])).

 cnf(c_0_256,negated_conjecture,
     ( esk8_2(esk49_0,esk49_0) = nil
     | ssItem(hd(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[218]c_0_158,c_0_244]),c_0_213])])).

 cnf(c_0_257,negated_conjecture,
     ( X1 = esk9_2(esk49_0,esk48_0)
     | app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),X1) != esk49_0
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[219]c_0_120,c_0_124]),c_0_125]),c_0_59])])).

 cnf(c_0_258,negated_conjecture,
     ( app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)))) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[220]c_0_75,c_0_245]),c_0_125]),c_0_52])])).

 cnf(c_0_259,negated_conjecture,
     ( ssList(esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[221]c_0_77,c_0_245]),c_0_125]),c_0_52])])).

 cnf(c_0_260,negated_conjecture,
     ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil
     | ssItem(hd(esk49_0)) ),
     inference(spm,[status(thm)],[[222]c_0_171,[223]c_0_246])).

 cnf(c_0_261,negated_conjecture,
     ( esk8_2(esk49_0,esk49_0) = nil
     | esk6_2(esk49_0,esk8_2(esk49_0,esk49_0)) != esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[224]c_0_132,c_0_234]),c_0_235]),c_0_213])])).

 cnf(c_0_262,negated_conjecture,
     ( esk6_2(esk49_0,esk8_2(esk49_0,esk49_0)) = nil
     | ssList(tl(esk49_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[225]c_0_247,c_0_248]),c_0_249])]),[226]c_0_144])).

 cnf(c_0_263,plain,
     ( X1 = esk7_2(X2,X2)
     | app(X1,X2) != X2
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[227]c_0_118,c_0_250]),[228]c_0_251])).

 cnf(c_0_264,plain,
     ( app(esk7_2(nil,nil),nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[229]c_0_93,c_0_252]),c_0_74])])).

 cnf(c_0_265,plain,
     ( ssList(esk7_2(nil,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[230]c_0_95,c_0_252]),c_0_74])])).

 cnf(c_0_266,negated_conjecture,
     ( esk49_0 = nil
     | rearsegP(esk49_0,tl(esk49_0))
     | ~ ssItem(hd(esk49_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[231]c_0_253,c_0_152]),[232]c_0_144])).

 cnf(c_0_267,negated_conjecture,
     ( nil != esk48_0
     | ~ segmentP(esk48_0,X1)
     | ~ segmentP(esk49_0,X1)
     | ~ neq(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[233]c_0_41])).

 cnf(c_0_268,negated_conjecture,
     ( segmentP(app(esk49_0,X1),esk9_2(esk49_0,nil))
     | ~ ssList(X1) ),
     inference(rw,[status(thm)],[[234]c_0_254,[235]c_0_238])).

 cnf(c_0_269,negated_conjecture,
     ( app(esk49_0,esk9_2(esk49_0,esk49_0)) = nil
     | ssItem(hd(esk49_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[236]c_0_255,c_0_256]),[237]c_0_217])).

 cnf(c_0_270,negated_conjecture,
     ( segmentP(esk49_0,esk9_2(esk49_0,esk48_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[238]c_0_242,c_0_99]),c_0_52]),c_0_74])])).

 cnf(c_0_271,negated_conjecture,
     ( esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))) = esk9_2(esk49_0,esk48_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[239]c_0_257,c_0_258]),c_0_259])])).

 cnf(c_0_272,negated_conjecture,
     ( esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)) = esk48_0
     | ssItem(hd(esk49_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[240]c_0_202,c_0_260]),c_0_51])]),[241]c_0_217])).

 cnf(c_0_273,negated_conjecture,
     ( esk8_2(esk49_0,esk49_0) = nil
     | ssList(tl(esk49_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[242]c_0_261,c_0_262]),[243]c_0_144])).

 cnf(c_0_274,plain,
     ( app(X1,app(nil,X2)) = app(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[244]c_0_57,c_0_99]),c_0_74])])).

 cnf(c_0_275,negated_conjecture,
     ( app(esk8_2(esk49_0,nil),app(nil,esk9_2(esk49_0,nil))) = esk49_0 ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[245]c_0_62,c_0_238]),c_0_238]),[246]c_0_238])).

 cnf(c_0_276,negated_conjecture,
     ( ssList(esk9_2(esk49_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[247]c_0_53,c_0_239]),c_0_74]),c_0_52])])).

 cnf(c_0_277,negated_conjecture,
     ( esk8_2(esk49_0,esk48_0) = esk7_2(esk48_0,esk48_0)
     | ssList(tl(esk49_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[248]c_0_263,c_0_169]),c_0_51]),c_0_60])]),[249]c_0_181])).

 cnf(c_0_278,plain,
     ( esk7_2(nil,nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[250]c_0_123,c_0_264]),c_0_74]),c_0_265])])).

 cnf(c_0_279,negated_conjecture,
     ( esk49_0 = nil
     | rearsegP(esk49_0,tl(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[251]c_0_266,c_0_158]),c_0_52])])).

 cnf(c_0_280,plain,
     ( app(esk7_2(X1,nil),nil) = X1
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[252]c_0_93,c_0_73]),c_0_74])])).

 cnf(c_0_281,plain,
     ( ssList(esk7_2(X1,nil))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[253]c_0_95,c_0_73]),c_0_74])])).

 cnf(c_0_282,negated_conjecture,
     ( ~ segmentP(nil,X1)
     | ~ segmentP(esk49_0,X1)
     | ~ ssList(X1)
     | ~ neq(X1,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[254]c_0_267,c_0_238]),c_0_238])])).

 cnf(c_0_283,negated_conjecture,
     ( segmentP(nil,esk9_2(esk49_0,nil))
     | ssItem(hd(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[255]c_0_268,c_0_269]),c_0_214])])).

 cnf(c_0_284,negated_conjecture,
     ( segmentP(esk49_0,esk9_2(esk49_0,nil)) ),
     inference(rw,[status(thm)],[[256]c_0_270,[257]c_0_238])).

 cnf(c_0_285,negated_conjecture,
     ( esk9_2(esk49_0,esk48_0) = esk6_2(esk49_0,esk48_0)
     | ssItem(hd(esk49_0)) ),
     inference(spm,[status(thm)],[[258]c_0_271,[259]c_0_272])).

 cnf(c_0_286,negated_conjecture,
     ( neq(esk49_0,nil)
     | nil != esk48_0 ),
     inference(split_conjunct,[status(thm)],[[260]c_0_41])).

 cnf(c_0_287,negated_conjecture,
     ( app(esk49_0,esk9_2(esk49_0,esk49_0)) = nil
     | ssList(tl(esk49_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[261]c_0_255,c_0_273]),[262]c_0_144])).

 cnf(c_0_288,negated_conjecture,
     ( app(nil,esk9_2(esk49_0,esk48_0)) = esk49_0
     | ssList(tl(esk49_0)) ),
     inference(spm,[status(thm)],[[263]c_0_124,[264]c_0_143])).

 cnf(c_0_289,plain,
     ( nil = X1
     | rearsegP(app(X1,X2),tl(app(X1,X2)))
     | ~ ssList(tl(app(X1,X2)))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[265]c_0_253,c_0_140]),[266]c_0_158])).

 cnf(c_0_290,negated_conjecture,
     ( app(esk8_2(esk49_0,nil),esk9_2(esk49_0,nil)) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[267]c_0_274,c_0_275]),c_0_276]),c_0_249])])).

 cnf(c_0_291,negated_conjecture,
     ( esk8_2(esk49_0,nil) = nil
     | ssList(tl(esk49_0)) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[268]c_0_277,c_0_238]),c_0_238]),c_0_238]),[269]c_0_278])).

 cnf(c_0_292,plain,
     ( X1 = cons(X2,nil)
     | app(X1,X3) != cons(X2,X3)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[270]c_0_118,c_0_138]),[271]c_0_139])).

 cnf(c_0_293,negated_conjecture,
     ( app(esk7_2(esk49_0,tl(esk49_0)),tl(esk49_0)) = esk49_0
     | esk49_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[272]c_0_93,c_0_279]),c_0_52])]),[273]c_0_144])).

 cnf(c_0_294,negated_conjecture,
     ( esk49_0 = nil
     | ssList(esk7_2(esk49_0,tl(esk49_0))) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[274]c_0_95,c_0_279]),c_0_52])]),[275]c_0_144])).

 cnf(c_0_295,plain,
     ( app(tl(esk7_2(X1,nil)),nil) = tl(X1)
     | esk7_2(X1,nil) = nil
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[276]c_0_122,c_0_280]),c_0_74])]),[277]c_0_281])).

 cnf(c_0_296,negated_conjecture,
     ( ssItem(hd(esk49_0))
     | ~ neq(esk9_2(esk49_0,nil),nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[278]c_0_282,c_0_283]),c_0_284]),c_0_276])])).

 cnf(c_0_297,negated_conjecture,
     ( esk9_2(esk49_0,nil) = esk49_0
     | ssItem(hd(esk49_0)) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[279]c_0_285,c_0_238]),c_0_238]),[280]c_0_162])).

 cnf(c_0_298,negated_conjecture,
     ( neq(esk49_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[281]c_0_286,c_0_238])])).

 cnf(c_0_299,negated_conjecture,
     ( esk8_2(esk49_0,esk48_0) = nil
     | esk49_0 != nil
     | ~ ssList(app(esk48_0,esk9_2(esk49_0,esk48_0))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[282]c_0_123,c_0_62]),c_0_60])])).

 cnf(c_0_300,negated_conjecture,
     ( segmentP(nil,esk9_2(esk49_0,nil))
     | ssList(tl(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[283]c_0_268,c_0_287]),c_0_214])])).

 cnf(c_0_301,negated_conjecture,
     ( esk9_2(esk49_0,esk48_0) = esk49_0
     | ssList(tl(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[284]c_0_119,c_0_288]),c_0_59])])).

 cnf(c_0_302,negated_conjecture,
     ( esk8_2(esk49_0,nil) = nil
     | rearsegP(esk49_0,tl(esk49_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[285]c_0_289,c_0_290]),c_0_276]),c_0_249])]),[286]c_0_291])).

 cnf(c_0_303,negated_conjecture,
     ( esk7_2(esk49_0,tl(esk49_0)) = cons(X1,nil)
     | esk49_0 = nil
     | cons(X1,tl(esk49_0)) != esk49_0
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[287]c_0_292,c_0_293]),c_0_294]),[288]c_0_144])).

 cnf(c_0_304,plain,
     ( cons(hd(app(esk7_2(X1,nil),nil)),tl(X1)) = app(esk7_2(X1,nil),nil)
     | esk7_2(X1,nil) = nil
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[289]c_0_193,c_0_295]),c_0_74])]),[290]c_0_281])).

 cnf(c_0_305,negated_conjecture,
     ( ssItem(hd(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[291]c_0_296,c_0_297]),c_0_298])])).

 cnf(c_0_306,negated_conjecture,
     ( esk8_2(esk49_0,esk48_0) = nil
     | esk49_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[292]c_0_299,c_0_56]),c_0_59]),c_0_51])])).

 cnf(c_0_307,negated_conjecture,
     ( ~ totalorderedP(X1)
     | ~ segmentP(X1,esk50_0)
     | ~ segmentP(esk51_0,X1)
     | ~ neq(esk50_0,X1)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[293]c_0_41])).

 fof(c_0_308,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | segmentP(X2,nil) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])).

 cnf(c_0_309,plain,
     ( segmentP(app(app(X1,X2),X3),X2)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[294]c_0_207,[295]c_0_175])).

 cnf(c_0_310,negated_conjecture,
     ( ssList(tl(esk49_0))
     | ~ neq(esk9_2(esk49_0,nil),nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[296]c_0_282,c_0_300]),c_0_284]),c_0_276])])).

 cnf(c_0_311,negated_conjecture,
     ( esk9_2(esk49_0,nil) = esk49_0
     | ssList(tl(esk49_0)) ),
     inference(rw,[status(thm)],[[297]c_0_301,[298]c_0_238])).

 cnf(c_0_312,negated_conjecture,
     ( esk8_2(esk49_0,nil) = nil
     | ssList(esk7_2(esk49_0,tl(esk49_0))) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[299]c_0_95,c_0_302]),c_0_52])]),[300]c_0_291])).

 cnf(c_0_313,negated_conjecture,
     ( esk7_2(esk49_0,tl(esk49_0)) = cons(hd(esk49_0),nil)
     | esk49_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[301]c_0_303,c_0_304]),c_0_112]),c_0_126]),c_0_112]),c_0_112]),c_0_126]),c_0_112]),c_0_126]),c_0_305]),c_0_52])])).

 cnf(c_0_314,negated_conjecture,
     ( esk8_2(esk49_0,nil) = nil
     | esk49_0 != nil ),
     inference(rw,[status(thm)],[[302]c_0_306,[303]c_0_238])).

 cnf(c_0_315,negated_conjecture,
     ( ~ totalorderedP(X1)
     | ~ segmentP(X1,esk48_0)
     | ~ segmentP(esk49_0,X1)
     | ~ ssList(X1)
     | ~ neq(esk48_0,X1) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[304]c_0_307,c_0_45]),c_0_45]),[305]c_0_44])).

 cnf(c_0_316,plain,
     ( segmentP(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[306]c_0_308])).

 cnf(c_0_317,plain,
     ( segmentP(app(X1,X2),X1)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[307]c_0_309,c_0_250]),[308]c_0_251])).

 cnf(c_0_318,negated_conjecture,
     ( ssList(tl(esk49_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[309]c_0_310,c_0_311]),c_0_298])])).

 cnf(c_0_319,negated_conjecture,
     ( app(nil,esk9_2(esk49_0,nil)) = nil
     | esk8_2(esk49_0,nil) != esk49_0 ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[310]c_0_171,c_0_238]),c_0_238]),[311]c_0_238])).

 cnf(c_0_320,negated_conjecture,
     ( esk8_2(esk49_0,nil) = nil
     | ssList(cons(hd(esk49_0),nil)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[312]c_0_312,c_0_313]),[313]c_0_314])).

 cnf(c_0_321,negated_conjecture,
     ( esk49_0 = nil
     | ssList(cons(hd(esk49_0),nil)) ),
     inference(spm,[status(thm)],[[314]c_0_294,[315]c_0_313])).

 cnf(c_0_322,negated_conjecture,
     ( X1 = esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))
     | app(esk8_2(esk49_0,esk48_0),X1) != esk49_0
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[316]c_0_120,c_0_85]),c_0_60]),c_0_86])])).

 cnf(c_0_323,negated_conjecture,
     ( ~ totalorderedP(X1)
     | ~ segmentP(esk49_0,X1)
     | ~ ssList(X1)
     | ~ neq(nil,X1) ),
     inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[317]c_0_315,c_0_238]),c_0_238]),[318]c_0_316])).

 cnf(c_0_324,negated_conjecture,
     ( esk49_0 = nil
     | segmentP(esk49_0,esk7_2(esk49_0,tl(esk49_0))) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[319]c_0_317,c_0_293]),c_0_318])]),[320]c_0_294])).

 cnf(c_0_325,plain,
     ( segmentP(app(X1,X2),X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[321]c_0_309,c_0_99]),c_0_74])]),[322]c_0_56])).

 cnf(c_0_326,negated_conjecture,
     ( app(nil,esk9_2(esk49_0,nil)) = nil
     | ssList(cons(hd(esk49_0),nil)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[323]c_0_319,c_0_320]),[324]c_0_321])).

 cnf(c_0_327,negated_conjecture,
     ( X1 = esk6_2(esk49_0,esk8_2(esk49_0,nil))
     | app(esk8_2(esk49_0,nil),X1) != esk49_0
     | ~ ssList(X1) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[325]c_0_322,c_0_238]),[326]c_0_238])).

 cnf(c_0_328,negated_conjecture,
     ( esk49_0 = nil
     | ~ totalorderedP(esk7_2(esk49_0,tl(esk49_0)))
     | ~ neq(nil,esk7_2(esk49_0,tl(esk49_0))) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[327]c_0_323,c_0_324]),[328]c_0_294])).

 fof(c_0_329,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | totalorderedP(cons(X2,nil)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])).

 cnf(c_0_330,negated_conjecture,
     ( segmentP(nil,esk9_2(esk49_0,nil))
     | ssList(cons(hd(esk49_0),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[329]c_0_325,c_0_326]),c_0_276]),c_0_74])])).

 cnf(c_0_331,negated_conjecture,
     ( esk6_2(esk49_0,esk8_2(esk49_0,nil)) = esk9_2(esk49_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[330]c_0_327,c_0_290]),c_0_276])])).

 cnf(c_0_332,negated_conjecture,
     ( esk49_0 = nil
     | ~ totalorderedP(cons(hd(esk49_0),nil))
     | ~ neq(nil,cons(hd(esk49_0),nil)) ),
     inference(spm,[status(thm)],[[331]c_0_328,[332]c_0_313])).

 cnf(c_0_333,plain,
     ( totalorderedP(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[333]c_0_329])).

 cnf(c_0_334,negated_conjecture,
     ( ssList(cons(hd(esk49_0),nil))
     | ~ neq(esk9_2(esk49_0,nil),nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[334]c_0_282,c_0_330]),c_0_284]),c_0_276])])).

 cnf(c_0_335,negated_conjecture,
     ( esk9_2(esk49_0,nil) = esk49_0
     | ssList(cons(hd(esk49_0),nil)) ),
     inference(rw,[status(thm)],[inference(spm,[status(thm)],[[335]c_0_331,c_0_320]),[336]c_0_162])).

 cnf(c_0_336,negated_conjecture,
     ( totalorderedP(esk50_0) ),
     inference(split_conjunct,[status(thm)],[[337]c_0_41])).

 fof(c_0_337,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) != X3 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_39])])])).

 cnf(c_0_338,negated_conjecture,
     ( esk49_0 = nil
     | ~ neq(nil,cons(hd(esk49_0),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[338]c_0_332,c_0_333]),c_0_305])])).

 cnf(c_0_339,negated_conjecture,
     ( ssList(cons(hd(esk49_0),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[339]c_0_334,c_0_335]),c_0_298])])).

 cnf(c_0_340,negated_conjecture,
     ( totalorderedP(esk48_0) ),
     inference(rw,[status(thm)],[[340]c_0_336,[341]c_0_45])).

 cnf(c_0_341,plain,
     ( cons(X1,X2) != X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[342]c_0_337])).

 cnf(c_0_342,negated_conjecture,
     ( cons(hd(esk49_0),nil) = nil
     | esk49_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[343]c_0_338,c_0_186]),c_0_339]),c_0_74])])).

 cnf(c_0_343,negated_conjecture,
     ( ~ neq(esk48_0,esk48_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[344]c_0_315,c_0_175]),c_0_340]),c_0_50]),c_0_51])])).

 cnf(c_0_344,negated_conjecture,
     ( esk49_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[345]c_0_341,c_0_342]),c_0_74]),c_0_305])])).

 cnf(c_0_345,negated_conjecture,
     ( ~ neq(nil,nil) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[346]c_0_343,c_0_238]),[347]c_0_238])).

 cnf(c_0_346,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[[348]c_0_298,c_0_344]),[349]c_0_345]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC066+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.23  % Computer   : n128.star.cs.uiowa.edu
 % 0.03/0.23  % Model      : x86_64 x86_64
 % 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.23  % Memory     : 32218.75MB
 % 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.23  % CPULimit   : 300
 % 0.03/0.23  % DateTime   : Tue May 10 20:35:07 CDT 2016
 % 0.03/0.23  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 49.35/49.55  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 49.35/49.55  # and selection function SelectCQPrecWNTNp.
 % 49.35/49.55  #
 % 49.35/49.55  # Preprocessing time       : 0.014 s
 % 49.35/49.55  # Presaturation interreduction done
 % 49.35/49.55
 % 49.35/49.55  # Proof found!
 % 49.35/49.55  # SZS status Theorem
 % 49.35/49.55  # SZS output start CNFRefutation
 % 49.35/49.55  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~(segmentP(X4,X3)))|~(totalorderedP(X3)))|?[X5]:((((ssList(X5)&neq(X3,X5))&segmentP(X4,X5))&segmentP(X5,X3))&totalorderedP(X5)))|((nil!=X2|nil=X1)&(~(neq(X2,nil))|?[X6]:(((ssList(X6)&neq(X6,nil))&segmentP(X2,X6))&segmentP(X1,X6)))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 49.35/49.55  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)<=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax7)).
 % 49.35/49.55  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)<=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax5)).
 % 49.35/49.55  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 49.35/49.55  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
 % 49.35/49.55  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(rearsegP(X1,X2)=>rearsegP(app(X3,X1),X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax50)).
 % 49.35/49.55  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>rearsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax51)).
 % 49.35/49.55  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax78)).
 % 49.35/49.55  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
 % 49.35/49.55  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 49.35/49.55  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>nil!=cons(X2,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax21)).
 % 49.35/49.55  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)<=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax4)).
 % 49.35/49.55  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)<=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax6)).
 % 49.35/49.55  fof(c_0_13, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 49.35/49.55  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
 % 49.35/49.55  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(frontsegP(X1,X2)=>frontsegP(app(X1,X3),X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax43)).
 % 49.35/49.55  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax45)).
 % 49.35/49.55  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax79)).
 % 49.35/49.55  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 49.35/49.55  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X2,X3)=app(X2,X1)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax80)).
 % 49.35/49.55  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
 % 49.35/49.55  fof(c_0_21, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax86)).
 % 49.35/49.55  fof(c_0_22, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 49.35/49.55  fof(c_0_23, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax23)).
 % 49.35/49.55  fof(c_0_24, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>hd(app(X1,X2))=hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax85)).
 % 49.35/49.55  fof(c_0_25, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
 % 49.35/49.55  fof(c_0_26, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax22)).
 % 49.35/49.55  fof(c_0_27, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(frontsegP(cons(X1,X3),cons(X2,X4))<=>(X1=X2&frontsegP(X3,X4)))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax44)).
 % 49.35/49.55  fof(c_0_28, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax20)).
 % 49.35/49.55  fof(c_0_29, axiom, (![X1]:(ssList(X1)=>segmentP(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax55)).
 % 49.35/49.55  fof(c_0_30, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
 % 49.35/49.55  fof(c_0_31, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 49.35/49.55  fof(c_0_32, axiom, (![X1]:(ssList(X1)=>(segmentP(nil,X1)<=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax58)).
 % 49.35/49.55  fof(c_0_33, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(segmentP(X1,X2)=>segmentP(app(app(X3,X1),X4),X2))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax56)).
 % 49.35/49.55  fof(c_0_34, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax42)).
 % 49.35/49.55  fof(c_0_35, axiom, (![X1]:(ssList(X1)=>rearsegP(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax49)).
 % 49.35/49.55  fof(c_0_36, axiom, (![X1]:(ssList(X1)=>(rearsegP(nil,X1)<=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax52)).
 % 49.35/49.55  fof(c_0_37, axiom, (![X1]:(ssList(X1)=>segmentP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax57)).
 % 49.35/49.55  fof(c_0_38, axiom, (![X1]:(ssItem(X1)=>totalorderedP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax65)).
 % 49.35/49.55  fof(c_0_39, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax18)).
 % 49.35/49.55  fof(c_0_40, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~segmentP(X4,X3))|~totalorderedP(X3))|?[X5]:((((ssList(X5)&neq(X3,X5))&segmentP(X4,X5))&segmentP(X5,X3))&totalorderedP(X5)))|((nil!=X2|nil=X1)&(~neq(X2,nil)|?[X6]:(((ssList(X6)&neq(X6,nil))&segmentP(X2,X6))&segmentP(X1,X6))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 49.35/49.55  fof(c_0_41, negated_conjecture, (![X11]:![X12]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((((esk49_0=esk51_0&esk48_0=esk50_0)&segmentP(esk51_0,esk50_0))&totalorderedP(esk50_0))&((((~ssList(X11)|~neq(esk50_0,X11))|~segmentP(esk51_0,X11))|~segmentP(X11,esk50_0))|~totalorderedP(X11)))&(((neq(esk49_0,nil)|nil=esk49_0)&((((~ssList(X12)|~neq(X12,nil))|~segmentP(esk49_0,X12))|~segmentP(esk48_0,X12))|nil=esk49_0))&((neq(esk49_0,nil)|nil!=esk48_0)&((((~ssList(X12)|~neq(X12,nil))|~segmentP(esk49_0,X12))|~segmentP(esk48_0,X12))|nil!=esk48_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])])])])).
 % 49.35/49.55  fof(c_0_42, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
 % 49.35/49.55  cnf(c_0_43, negated_conjecture, (segmentP(esk51_0,esk50_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 49.35/49.55  cnf(c_0_44, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 49.35/49.55  cnf(c_0_45, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 49.35/49.55  fof(c_0_46, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
 % 49.35/49.55  fof(c_0_47, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
 % 49.35/49.55  fof(c_0_48, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 49.35/49.55  cnf(c_0_49, plain, (app(app(esk8_2(X1,X2),X2),esk9_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
 % 49.35/49.55  cnf(c_0_50, negated_conjecture, (segmentP(esk49_0,esk48_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
 % 49.35/49.55  cnf(c_0_51, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 49.35/49.55  cnf(c_0_52, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 49.35/49.55  cnf(c_0_53, plain, (ssList(esk9_2(X1,X2))|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
 % 49.35/49.55  cnf(c_0_54, plain, (ssList(esk8_2(X1,X2))|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
 % 49.35/49.55  cnf(c_0_55, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_46])).
 % 49.35/49.55  cnf(c_0_56, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
 % 49.35/49.55  cnf(c_0_57, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
 % 49.35/49.55  cnf(c_0_58, negated_conjecture, (app(app(esk8_2(esk49_0,esk48_0),esk48_0),esk9_2(esk49_0,esk48_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_59, negated_conjecture, (ssList(esk9_2(esk49_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_50]), c_0_51]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_60, negated_conjecture, (ssList(esk8_2(esk49_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_50]), c_0_51]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_61, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_55]), c_0_56])).
 % 49.35/49.55  cnf(c_0_62, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),app(esk48_0,esk9_2(esk49_0,esk48_0)))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_51]), c_0_60])])).
 % 49.35/49.55  fof(c_0_63, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~rearsegP(X4,X5)|rearsegP(app(X6,X4),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
 % 49.35/49.55  fof(c_0_64, plain, (![X2]:(~ssList(X2)|rearsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
 % 49.35/49.55  cnf(c_0_65, negated_conjecture, (frontsegP(esk49_0,esk8_2(esk49_0,esk48_0))|~ssList(app(esk48_0,esk9_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_60])])).
 % 49.35/49.55  fof(c_0_66, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
 % 49.35/49.55  fof(c_0_67, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
 % 49.35/49.55  fof(c_0_68, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
 % 49.35/49.55  fof(c_0_69, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|nil!=cons(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
 % 49.35/49.55  fof(c_0_70, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
 % 49.35/49.55  fof(c_0_71, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).
 % 49.35/49.55  cnf(c_0_72, plain, (rearsegP(app(X1,X2),X3)|~rearsegP(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
 % 49.35/49.55  cnf(c_0_73, plain, (rearsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
 % 49.35/49.55  cnf(c_0_74, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 49.35/49.55  cnf(c_0_75, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
 % 49.35/49.55  cnf(c_0_76, negated_conjecture, (frontsegP(esk49_0,esk8_2(esk49_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_56]), c_0_59]), c_0_51])])).
 % 49.35/49.55  cnf(c_0_77, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
 % 49.35/49.55  cnf(c_0_78, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
 % 49.35/49.55  cnf(c_0_79, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
 % 49.35/49.55  cnf(c_0_80, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
 % 49.35/49.55  cnf(c_0_81, plain, (nil!=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
 % 49.35/49.55  cnf(c_0_82, plain, (singletonP(X1)|~ssList(X1)|cons(X2,nil)!=X1|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
 % 49.35/49.55  cnf(c_0_83, plain, (rearsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X3,X2)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_71])).
 % 49.35/49.55  cnf(c_0_84, plain, (rearsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_85, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk6_2(esk49_0,esk8_2(esk49_0,esk48_0)))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_60]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_86, negated_conjecture, (ssList(esk6_2(esk49_0,esk8_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_76]), c_0_60]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_87, plain, (cons(hd(cons(X1,X2)),X2)=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80]), c_0_81])).
 % 49.35/49.55  cnf(c_0_88, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
 % 49.35/49.55  cnf(c_0_89, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
 % 49.35/49.55  cnf(c_0_90, plain, (singletonP(cons(X1,nil))|~ssList(cons(X1,nil))|~ssItem(X1)), inference(er,[status(thm)],[c_0_82])).
 % 49.35/49.55  cnf(c_0_91, plain, (rearsegP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_83]), c_0_56])).
 % 49.35/49.55  fof(c_0_92, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).
 % 49.35/49.55  cnf(c_0_93, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
 % 49.35/49.55  cnf(c_0_94, negated_conjecture, (rearsegP(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_95, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
 % 49.35/49.55  cnf(c_0_96, plain, (cons(hd(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_74])]), c_0_89])).
 % 49.35/49.55  cnf(c_0_97, plain, (singletonP(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_80]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_98, negated_conjecture, (rearsegP(esk49_0,esk9_2(esk49_0,esk48_0))|~ssList(app(esk8_2(esk49_0,esk48_0),esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_58]), c_0_59])])).
 % 49.35/49.55  cnf(c_0_99, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
 % 49.35/49.55  cnf(c_0_100, negated_conjecture, (app(esk7_2(esk49_0,nil),nil)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_74]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_101, negated_conjecture, (ssList(esk7_2(esk49_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_94]), c_0_74]), c_0_52])])).
 % 49.35/49.55  fof(c_0_102, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~frontsegP(X4,X5)|frontsegP(app(X4,X6),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
 % 49.35/49.55  fof(c_0_103, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).
 % 49.35/49.55  cnf(c_0_104, plain, (cons(hd(cons(X1,nil)),nil)=cons(X1,nil)|~ssList(cons(X1,nil))|~ssItem(X1)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
 % 49.35/49.55  fof(c_0_105, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
 % 49.35/49.55  fof(c_0_106, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])).
 % 49.35/49.55  fof(c_0_107, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X5,X6)!=app(X5,X4)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
 % 49.35/49.55  fof(c_0_108, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])).
 % 49.35/49.55  fof(c_0_109, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
 % 49.35/49.55  fof(c_0_110, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).
 % 49.35/49.55  cnf(c_0_111, negated_conjecture, (rearsegP(esk49_0,esk9_2(esk49_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_56]), c_0_51]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_112, negated_conjecture, (esk7_2(esk49_0,nil)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101])])).
 % 49.35/49.55  cnf(c_0_113, plain, (frontsegP(app(X1,X2),X3)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
 % 49.35/49.55  cnf(c_0_114, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
 % 49.35/49.55  cnf(c_0_115, plain, (cons(hd(cons(X1,nil)),nil)=cons(X1,nil)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_80]), c_0_74])])).
 % 49.35/49.55  fof(c_0_116, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).
 % 49.35/49.55  fof(c_0_117, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|hd(app(X3,X4))=hd(X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
 % 49.35/49.55  cnf(c_0_118, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_105])).
 % 49.35/49.55  cnf(c_0_119, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
 % 49.35/49.55  cnf(c_0_120, plain, (X1=X2|app(X3,X1)!=app(X3,X2)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
 % 49.35/49.55  cnf(c_0_121, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
 % 49.35/49.55  cnf(c_0_122, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
 % 49.35/49.55  cnf(c_0_123, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_110])).
 % 49.35/49.55  cnf(c_0_124, negated_conjecture, (app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),esk9_2(esk49_0,esk48_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_111]), c_0_59]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_125, negated_conjecture, (ssList(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_111]), c_0_59]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_126, negated_conjecture, (app(esk49_0,nil)=esk49_0), inference(rw,[status(thm)],[c_0_100, c_0_112])).
 % 49.35/49.55  cnf(c_0_127, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_74])])).
 % 49.35/49.55  fof(c_0_128, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
 % 49.35/49.55  cnf(c_0_129, plain, (ssList(cons(X1,nil))|~ssItem(hd(cons(X1,nil)))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_115]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_130, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_116])).
 % 49.35/49.55  cnf(c_0_131, plain, (hd(app(X1,X2))=hd(X1)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_117])).
 % 49.35/49.55  cnf(c_0_132, plain, (X1=nil|app(X1,X2)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_133, plain, (X1=nil|app(X2,X1)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_99]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_134, plain, (nil=X1|ssList(app(tl(X1),X2))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_56]), c_0_123])).
 % 49.35/49.55  cnf(c_0_135, negated_conjecture, (app(tl(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))),esk9_2(esk49_0,esk48_0))=tl(esk49_0)|esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_124]), c_0_59]), c_0_125])])).
 % 49.35/49.55  cnf(c_0_136, negated_conjecture, (app(tl(esk49_0),nil)=tl(esk49_0)|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_126]), c_0_74]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_137, negated_conjecture, (frontsegP(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_85]), c_0_86]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_138, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_128])).
 % 49.35/49.55  cnf(c_0_139, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_140, plain, (cons(hd(X1),tl(app(X1,X2)))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_131]), c_0_56]), c_0_123])).
 % 49.35/49.55  cnf(c_0_141, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk48_0)=nil|esk9_2(esk49_0,esk48_0)!=esk49_0|~ssList(app(esk8_2(esk49_0,esk48_0),esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_58]), c_0_59])])).
 % 49.35/49.55  cnf(c_0_142, negated_conjecture, (esk9_2(esk49_0,esk48_0)=nil|esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_124]), c_0_125]), c_0_59])])).
 % 49.35/49.55  cnf(c_0_143, negated_conjecture, (esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))=nil|ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_59]), c_0_125])])).
 % 49.35/49.55  cnf(c_0_144, negated_conjecture, (esk49_0=nil|ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_136]), c_0_74]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_145, plain, (app(tl(X1),nil)=tl(X1)|nil=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_99]), c_0_74])])).
 % 49.35/49.55  fof(c_0_146, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])).
 % 49.35/49.55  cnf(c_0_147, negated_conjecture, (app(nil,esk6_2(esk49_0,nil))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_137]), c_0_74]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_148, negated_conjecture, (ssList(esk6_2(esk49_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_137]), c_0_74]), c_0_52])])).
 % 49.35/49.55  fof(c_0_149, plain, (![X5]:![X6]:![X7]:![X8]:(((((((X5=X6|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5))&(((((frontsegP(X7,X8)|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))&((((((X5!=X6|~frontsegP(X7,X8))|frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])).
 % 49.35/49.55  fof(c_0_150, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])])).
 % 49.35/49.55  cnf(c_0_151, plain, (frontsegP(cons(X1,X2),cons(X1,nil))|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_138]), c_0_139])).
 % 49.35/49.55  cnf(c_0_152, negated_conjecture, (cons(hd(esk49_0),tl(esk49_0))=esk49_0|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_126]), c_0_74]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_153, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk48_0)=nil|esk9_2(esk49_0,esk48_0)!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_56]), c_0_51]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_154, negated_conjecture, (esk9_2(esk49_0,esk48_0)=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143]), c_0_144])).
 % 49.35/49.55  cnf(c_0_155, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=nil|esk8_2(esk49_0,esk48_0)!=esk49_0|~ssList(app(esk48_0,esk9_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_62]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_156, plain, (app(app(tl(X1),X2),nil)=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_122]), c_0_56]), c_0_123])).
 % 49.35/49.55  cnf(c_0_157, negated_conjecture, (app(tl(esk8_2(esk49_0,esk48_0)),esk6_2(esk49_0,esk8_2(esk49_0,esk48_0)))=tl(esk49_0)|esk8_2(esk49_0,esk48_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_85]), c_0_86]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_158, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_146])).
 % 49.35/49.55  cnf(c_0_159, negated_conjecture, (hd(esk8_2(esk49_0,esk48_0))=hd(esk49_0)|esk8_2(esk49_0,esk48_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_85]), c_0_86]), c_0_60])])).
 % 49.35/49.55  fof(c_0_160, plain, (![X2]:(~ssList(X2)|segmentP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])).
 % 49.35/49.55  cnf(c_0_161, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
 % 49.35/49.55  cnf(c_0_162, negated_conjecture, (esk6_2(esk49_0,nil)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_147]), c_0_148])])).
 % 49.35/49.55  cnf(c_0_163, plain, (X1=X2|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~ssList(X4)|~frontsegP(cons(X1,X3),cons(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_149])).
 % 49.35/49.55  cnf(c_0_164, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_150])).
 % 49.35/49.55  cnf(c_0_165, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_150])).
 % 49.35/49.55  cnf(c_0_166, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_150])).
 % 49.35/49.55  cnf(c_0_167, negated_conjecture, (esk49_0=nil|frontsegP(esk49_0,cons(hd(esk49_0),nil))|~ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_144])).
 % 49.35/49.55  cnf(c_0_168, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_110])).
 % 49.35/49.55  cnf(c_0_169, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk48_0)=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154]), c_0_144])).
 % 49.35/49.55  fof(c_0_170, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
 % 49.35/49.55  cnf(c_0_171, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=nil|esk8_2(esk49_0,esk48_0)!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_56]), c_0_59]), c_0_51])])).
 % 49.35/49.55  cnf(c_0_172, negated_conjecture, (app(tl(esk49_0),nil)=tl(esk49_0)|esk8_2(esk49_0,esk48_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_86]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_173, negated_conjecture, (esk8_2(esk49_0,esk48_0)=nil|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_174, negated_conjecture, (nil=esk49_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 49.35/49.55  cnf(c_0_175, plain, (segmentP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_160])).
 % 49.35/49.55  fof(c_0_176, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])])).
 % 49.35/49.55  cnf(c_0_177, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_161])).
 % 49.35/49.55  cnf(c_0_178, negated_conjecture, (app(nil,esk49_0)=esk49_0), inference(rw,[status(thm)],[c_0_147, c_0_162])).
 % 49.35/49.55  cnf(c_0_179, plain, (esk45_1(X1)=X2|nil=X1|~frontsegP(X1,cons(X2,X3))|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164]), c_0_165]), c_0_166])).
 % 49.35/49.55  cnf(c_0_180, negated_conjecture, (esk49_0=nil|frontsegP(esk49_0,cons(hd(esk49_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_158]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_181, negated_conjecture, (esk48_0=nil|ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169]), c_0_51]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_182, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_170])).
 % 49.35/49.55  cnf(c_0_183, negated_conjecture, (app(tl(esk49_0),nil)=tl(esk49_0)|app(esk48_0,esk9_2(esk49_0,esk48_0))=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_136])).
 % 49.35/49.55  cnf(c_0_184, negated_conjecture, (app(app(nil,esk48_0),esk9_2(esk49_0,esk48_0))=esk49_0|ssItem(hd(esk49_0))), inference(spm,[status(thm)],[c_0_58, c_0_173])).
 % 49.35/49.55  cnf(c_0_185, negated_conjecture, (esk49_0=nil|~neq(esk48_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_175]), c_0_50]), c_0_51])])).
 % 49.35/49.55  cnf(c_0_186, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_176])).
 % 49.35/49.55  fof(c_0_187, plain, (![X2]:(((~segmentP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|segmentP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])).
 % 49.35/49.55  cnf(c_0_188, negated_conjecture, (segmentP(app(esk49_0,X1),esk49_0)|~ssList(app(esk49_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_178]), c_0_74]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_189, negated_conjecture, (esk45_1(esk49_0)=hd(esk49_0)|esk49_0=nil|~ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_74]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_190, negated_conjecture, (esk8_2(esk49_0,esk48_0)=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_169]), c_0_51]), c_0_60])]), c_0_181])).
 % 49.35/49.55  cnf(c_0_191, plain, (cons(X1,app(X2,X3))!=nil|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_182]), c_0_80]), c_0_81])).
 % 49.35/49.55  cnf(c_0_192, negated_conjecture, (app(tl(esk49_0),nil)=tl(esk49_0)|esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_183]), c_0_59]), c_0_51])])).
 % 49.35/49.55  cnf(c_0_193, plain, (cons(hd(app(X1,X2)),app(tl(X1),X2))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_122]), c_0_56]), c_0_123])).
 % 49.35/49.55  cnf(c_0_194, negated_conjecture, (app(nil,esk48_0)=nil|ssItem(hd(esk49_0))|esk49_0!=nil|~ssList(app(nil,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_184]), c_0_59])])).
 % 49.35/49.55  cnf(c_0_195, negated_conjecture, (esk48_0=nil|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_186]), c_0_74]), c_0_51])])).
 % 49.35/49.55  fof(c_0_196, plain, (![X5]:![X6]:![X7]:![X8]:(~ssList(X5)|(~ssList(X6)|(~ssList(X7)|(~ssList(X8)|(~segmentP(X5,X6)|segmentP(app(app(X7,X5),X8),X6))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])).
 % 49.35/49.55  cnf(c_0_197, plain, (segmentP(nil,X1)|~ssList(X1)|nil!=X1), inference(split_conjunct,[status(thm)],[c_0_187])).
 % 49.35/49.55  cnf(c_0_198, negated_conjecture, (segmentP(esk49_0,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_99]), c_0_52]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_199, plain, (nil=X1|app(X2,app(X3,X1))!=nil|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_57]), c_0_56])).
 % 49.35/49.55  cnf(c_0_200, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=esk49_0|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_119]), c_0_51])])).
 % 49.35/49.55  cnf(c_0_201, negated_conjecture, (esk45_1(esk49_0)=hd(esk49_0)|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_158]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_202, negated_conjecture, (X1=esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))|app(X1,esk9_2(esk49_0,esk48_0))!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_124]), c_0_59]), c_0_125])])).
 % 49.35/49.55  cnf(c_0_203, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_190]), c_0_144])).
 % 49.35/49.55  cnf(c_0_204, negated_conjecture, (esk48_0=nil|cons(X1,tl(esk49_0))!=nil|~ssItem(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192]), c_0_74])]), c_0_181])).
 % 49.35/49.55  cnf(c_0_205, negated_conjecture, (cons(hd(esk49_0),tl(esk49_0))=esk49_0|esk8_2(esk49_0,esk48_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_157]), c_0_85]), c_0_85]), c_0_86]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_206, negated_conjecture, (esk48_0=nil|ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_119]), c_0_51])]), c_0_195])).
 % 49.35/49.55  cnf(c_0_207, plain, (segmentP(app(app(X1,X2),X3),X4)|~segmentP(X2,X4)|~ssList(X3)|~ssList(X1)|~ssList(X4)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_196])).
 % 49.35/49.55  cnf(c_0_208, plain, (segmentP(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_197]), c_0_74])])).
 % 49.35/49.55  fof(c_0_209, plain, (![X2]:(~ssList(X2)|frontsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])).
 % 49.35/49.55  cnf(c_0_210, negated_conjecture, (frontsegP(esk49_0,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_126]), c_0_74]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_211, plain, (frontsegP(app(app(X1,X2),X3),X1)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_61]), c_0_56])).
 % 49.35/49.55  cnf(c_0_212, negated_conjecture, (app(app(esk8_2(esk49_0,esk49_0),esk49_0),esk9_2(esk49_0,esk49_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_198]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_213, negated_conjecture, (ssList(esk8_2(esk49_0,esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_198]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_214, negated_conjecture, (ssList(esk9_2(esk49_0,esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_198]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_215, negated_conjecture, (esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))=nil|app(X1,esk49_0)!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_85]), c_0_86]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_216, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk49_0)=esk49_0|ssItem(hd(esk49_0))), inference(spm,[status(thm)],[c_0_62, c_0_200])).
 % 49.35/49.55  cnf(c_0_217, negated_conjecture, (esk49_0=nil|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_201]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_218, negated_conjecture, (esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))=esk48_0|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203]), c_0_51])]), c_0_144])).
 % 49.35/49.55  cnf(c_0_219, negated_conjecture, (esk8_2(esk49_0,esk48_0)=nil|esk48_0=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_205]), c_0_206]), c_0_195])).
 % 49.35/49.55  cnf(c_0_220, plain, (segmentP(app(app(X1,nil),X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_221, plain, (frontsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_209])).
 % 49.35/49.55  cnf(c_0_222, negated_conjecture, (app(esk49_0,esk6_2(esk49_0,esk49_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_210]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_223, negated_conjecture, (ssList(esk6_2(esk49_0,esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_210]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_224, negated_conjecture, (frontsegP(esk49_0,esk8_2(esk49_0,esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_212]), c_0_213]), c_0_214]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_225, negated_conjecture, (esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))=nil|ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216]), c_0_60])]), c_0_217])).
 % 49.35/49.55  cnf(c_0_226, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=esk49_0|ssList(tl(esk49_0))), inference(spm,[status(thm)],[c_0_124, c_0_218])).
 % 49.35/49.55  cnf(c_0_227, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=nil|esk48_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_219]), c_0_195])).
 % 49.35/49.55  cnf(c_0_228, negated_conjecture, (segmentP(app(esk49_0,X1),nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_126]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_229, plain, (app(X1,esk6_2(X1,X1))=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_75, c_0_221])).
 % 49.35/49.55  cnf(c_0_230, negated_conjecture, (esk6_2(esk49_0,esk49_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_222]), c_0_52]), c_0_223])])).
 % 49.35/49.55  fof(c_0_231, plain, (![X2]:(~ssList(X2)|rearsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])).
 % 49.35/49.55  fof(c_0_232, plain, (![X2]:(((~rearsegP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|rearsegP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).
 % 49.35/49.55  cnf(c_0_233, negated_conjecture, (app(esk8_2(esk49_0,esk49_0),app(esk49_0,esk9_2(esk49_0,esk49_0)))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_212]), c_0_214]), c_0_52]), c_0_213])])).
 % 49.35/49.55  cnf(c_0_234, negated_conjecture, (app(esk8_2(esk49_0,esk49_0),esk6_2(esk49_0,esk8_2(esk49_0,esk49_0)))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_224]), c_0_213]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_235, negated_conjecture, (ssList(esk6_2(esk49_0,esk8_2(esk49_0,esk49_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_224]), c_0_213]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_236, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),nil)=esk49_0|ssItem(hd(esk49_0))), inference(spm,[status(thm)],[c_0_85, c_0_225])).
 % 49.35/49.55  cnf(c_0_237, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk49_0)=esk49_0|ssList(tl(esk49_0))), inference(spm,[status(thm)],[c_0_62, c_0_226])).
 % 49.35/49.55  cnf(c_0_238, negated_conjecture, (esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_227]), c_0_59]), c_0_51])])).
 % 49.35/49.55  cnf(c_0_239, negated_conjecture, (segmentP(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_229]), c_0_230]), c_0_74]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_240, plain, (rearsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_231])).
 % 49.35/49.55  cnf(c_0_241, plain, (rearsegP(nil,X1)|~ssList(X1)|nil!=X1), inference(split_conjunct,[status(thm)],[c_0_232])).
 % 49.35/49.55  cnf(c_0_242, negated_conjecture, (segmentP(app(esk49_0,X1),esk9_2(esk49_0,esk48_0))|~ssList(app(esk49_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_124]), c_0_125]), c_0_59])])).
 % 49.35/49.55  cnf(c_0_243, negated_conjecture, (app(esk49_0,esk9_2(esk49_0,esk49_0))=nil|esk8_2(esk49_0,esk49_0)!=esk49_0|~ssList(app(esk49_0,esk9_2(esk49_0,esk49_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_233]), c_0_213])])).
 % 49.35/49.55  cnf(c_0_244, negated_conjecture, (hd(esk8_2(esk49_0,esk49_0))=hd(esk49_0)|esk8_2(esk49_0,esk49_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_234]), c_0_235]), c_0_213])])).
 % 49.35/49.55  cnf(c_0_245, negated_conjecture, (frontsegP(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_124]), c_0_59]), c_0_125])])).
 % 49.35/49.55  cnf(c_0_246, negated_conjecture, (esk8_2(esk49_0,esk48_0)=esk49_0|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_236]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_247, negated_conjecture, (esk6_2(esk49_0,esk8_2(esk49_0,esk49_0))=nil|app(X1,esk49_0)!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_234]), c_0_235]), c_0_213])])).
 % 49.35/49.55  cnf(c_0_248, negated_conjecture, (app(esk8_2(esk49_0,nil),esk49_0)=esk49_0|ssList(tl(esk49_0))), inference(rw,[status(thm)],[c_0_237, c_0_238])).
 % 49.35/49.55  cnf(c_0_249, negated_conjecture, (ssList(esk8_2(esk49_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_239]), c_0_74]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_250, plain, (app(esk7_2(X1,X1),X1)=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_93, c_0_240])).
 % 49.35/49.55  cnf(c_0_251, plain, (ssList(esk7_2(X1,X1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_95, c_0_240])).
 % 49.35/49.55  cnf(c_0_252, plain, (rearsegP(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_241]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_253, plain, (rearsegP(cons(X1,X2),X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_138]), c_0_139])).
 % 49.35/49.55  cnf(c_0_254, negated_conjecture, (segmentP(app(esk49_0,X1),esk9_2(esk49_0,esk48_0))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_56]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_255, negated_conjecture, (app(esk49_0,esk9_2(esk49_0,esk49_0))=nil|esk8_2(esk49_0,esk49_0)!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_56]), c_0_214]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_256, negated_conjecture, (esk8_2(esk49_0,esk49_0)=nil|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_244]), c_0_213])])).
 % 49.35/49.55  cnf(c_0_257, negated_conjecture, (X1=esk9_2(esk49_0,esk48_0)|app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),X1)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_124]), c_0_125]), c_0_59])])).
 % 49.35/49.55  cnf(c_0_258, negated_conjecture, (app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_245]), c_0_125]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_259, negated_conjecture, (ssList(esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_245]), c_0_125]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_260, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=nil|ssItem(hd(esk49_0))), inference(spm,[status(thm)],[c_0_171, c_0_246])).
 % 49.35/49.55  cnf(c_0_261, negated_conjecture, (esk8_2(esk49_0,esk49_0)=nil|esk6_2(esk49_0,esk8_2(esk49_0,esk49_0))!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_234]), c_0_235]), c_0_213])])).
 % 49.35/49.55  cnf(c_0_262, negated_conjecture, (esk6_2(esk49_0,esk8_2(esk49_0,esk49_0))=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_248]), c_0_249])]), c_0_144])).
 % 49.35/49.55  cnf(c_0_263, plain, (X1=esk7_2(X2,X2)|app(X1,X2)!=X2|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_250]), c_0_251])).
 % 49.35/49.55  cnf(c_0_264, plain, (app(esk7_2(nil,nil),nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_252]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_265, plain, (ssList(esk7_2(nil,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_252]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_266, negated_conjecture, (esk49_0=nil|rearsegP(esk49_0,tl(esk49_0))|~ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_152]), c_0_144])).
 % 49.35/49.55  cnf(c_0_267, negated_conjecture, (nil!=esk48_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 49.35/49.55  cnf(c_0_268, negated_conjecture, (segmentP(app(esk49_0,X1),esk9_2(esk49_0,nil))|~ssList(X1)), inference(rw,[status(thm)],[c_0_254, c_0_238])).
 % 49.35/49.55  cnf(c_0_269, negated_conjecture, (app(esk49_0,esk9_2(esk49_0,esk49_0))=nil|ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_256]), c_0_217])).
 % 49.35/49.55  cnf(c_0_270, negated_conjecture, (segmentP(esk49_0,esk9_2(esk49_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_99]), c_0_52]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_271, negated_conjecture, (esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)))=esk9_2(esk49_0,esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_258]), c_0_259])])).
 % 49.35/49.55  cnf(c_0_272, negated_conjecture, (esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))=esk48_0|ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_260]), c_0_51])]), c_0_217])).
 % 49.35/49.55  cnf(c_0_273, negated_conjecture, (esk8_2(esk49_0,esk49_0)=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_262]), c_0_144])).
 % 49.35/49.55  cnf(c_0_274, plain, (app(X1,app(nil,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_99]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_275, negated_conjecture, (app(esk8_2(esk49_0,nil),app(nil,esk9_2(esk49_0,nil)))=esk49_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_238]), c_0_238]), c_0_238])).
 % 49.35/49.55  cnf(c_0_276, negated_conjecture, (ssList(esk9_2(esk49_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_239]), c_0_74]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_277, negated_conjecture, (esk8_2(esk49_0,esk48_0)=esk7_2(esk48_0,esk48_0)|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_169]), c_0_51]), c_0_60])]), c_0_181])).
 % 49.35/49.55  cnf(c_0_278, plain, (esk7_2(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_264]), c_0_74]), c_0_265])])).
 % 49.35/49.55  cnf(c_0_279, negated_conjecture, (esk49_0=nil|rearsegP(esk49_0,tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_158]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_280, plain, (app(esk7_2(X1,nil),nil)=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_73]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_281, plain, (ssList(esk7_2(X1,nil))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_73]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_282, negated_conjecture, (~segmentP(nil,X1)|~segmentP(esk49_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_267, c_0_238]), c_0_238])])).
 % 49.35/49.55  cnf(c_0_283, negated_conjecture, (segmentP(nil,esk9_2(esk49_0,nil))|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_269]), c_0_214])])).
 % 49.35/49.55  cnf(c_0_284, negated_conjecture, (segmentP(esk49_0,esk9_2(esk49_0,nil))), inference(rw,[status(thm)],[c_0_270, c_0_238])).
 % 49.35/49.55  cnf(c_0_285, negated_conjecture, (esk9_2(esk49_0,esk48_0)=esk6_2(esk49_0,esk48_0)|ssItem(hd(esk49_0))), inference(spm,[status(thm)],[c_0_271, c_0_272])).
 % 49.35/49.55  cnf(c_0_286, negated_conjecture, (neq(esk49_0,nil)|nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 49.35/49.55  cnf(c_0_287, negated_conjecture, (app(esk49_0,esk9_2(esk49_0,esk49_0))=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_273]), c_0_144])).
 % 49.35/49.55  cnf(c_0_288, negated_conjecture, (app(nil,esk9_2(esk49_0,esk48_0))=esk49_0|ssList(tl(esk49_0))), inference(spm,[status(thm)],[c_0_124, c_0_143])).
 % 49.35/49.55  cnf(c_0_289, plain, (nil=X1|rearsegP(app(X1,X2),tl(app(X1,X2)))|~ssList(tl(app(X1,X2)))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_140]), c_0_158])).
 % 49.35/49.55  cnf(c_0_290, negated_conjecture, (app(esk8_2(esk49_0,nil),esk9_2(esk49_0,nil))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_275]), c_0_276]), c_0_249])])).
 % 49.35/49.55  cnf(c_0_291, negated_conjecture, (esk8_2(esk49_0,nil)=nil|ssList(tl(esk49_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_277, c_0_238]), c_0_238]), c_0_238]), c_0_278])).
 % 49.35/49.55  cnf(c_0_292, plain, (X1=cons(X2,nil)|app(X1,X3)!=cons(X2,X3)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_138]), c_0_139])).
 % 49.35/49.55  cnf(c_0_293, negated_conjecture, (app(esk7_2(esk49_0,tl(esk49_0)),tl(esk49_0))=esk49_0|esk49_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_279]), c_0_52])]), c_0_144])).
 % 49.35/49.55  cnf(c_0_294, negated_conjecture, (esk49_0=nil|ssList(esk7_2(esk49_0,tl(esk49_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_279]), c_0_52])]), c_0_144])).
 % 49.35/49.55  cnf(c_0_295, plain, (app(tl(esk7_2(X1,nil)),nil)=tl(X1)|esk7_2(X1,nil)=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_280]), c_0_74])]), c_0_281])).
 % 49.35/49.55  cnf(c_0_296, negated_conjecture, (ssItem(hd(esk49_0))|~neq(esk9_2(esk49_0,nil),nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_283]), c_0_284]), c_0_276])])).
 % 49.35/49.55  cnf(c_0_297, negated_conjecture, (esk9_2(esk49_0,nil)=esk49_0|ssItem(hd(esk49_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_285, c_0_238]), c_0_238]), c_0_162])).
 % 49.35/49.55  cnf(c_0_298, negated_conjecture, (neq(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_286, c_0_238])])).
 % 49.35/49.55  cnf(c_0_299, negated_conjecture, (esk8_2(esk49_0,esk48_0)=nil|esk49_0!=nil|~ssList(app(esk48_0,esk9_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_62]), c_0_60])])).
 % 49.35/49.55  cnf(c_0_300, negated_conjecture, (segmentP(nil,esk9_2(esk49_0,nil))|ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_287]), c_0_214])])).
 % 49.35/49.55  cnf(c_0_301, negated_conjecture, (esk9_2(esk49_0,esk48_0)=esk49_0|ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_288]), c_0_59])])).
 % 49.35/49.55  cnf(c_0_302, negated_conjecture, (esk8_2(esk49_0,nil)=nil|rearsegP(esk49_0,tl(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_290]), c_0_276]), c_0_249])]), c_0_291])).
 % 49.35/49.55  cnf(c_0_303, negated_conjecture, (esk7_2(esk49_0,tl(esk49_0))=cons(X1,nil)|esk49_0=nil|cons(X1,tl(esk49_0))!=esk49_0|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_293]), c_0_294]), c_0_144])).
 % 49.35/49.55  cnf(c_0_304, plain, (cons(hd(app(esk7_2(X1,nil),nil)),tl(X1))=app(esk7_2(X1,nil),nil)|esk7_2(X1,nil)=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_295]), c_0_74])]), c_0_281])).
 % 49.35/49.55  cnf(c_0_305, negated_conjecture, (ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_297]), c_0_298])])).
 % 49.35/49.55  cnf(c_0_306, negated_conjecture, (esk8_2(esk49_0,esk48_0)=nil|esk49_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_56]), c_0_59]), c_0_51])])).
 % 49.35/49.55  cnf(c_0_307, negated_conjecture, (~totalorderedP(X1)|~segmentP(X1,esk50_0)|~segmentP(esk51_0,X1)|~neq(esk50_0,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 49.35/49.55  fof(c_0_308, plain, (![X2]:(~ssList(X2)|segmentP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])).
 % 49.35/49.55  cnf(c_0_309, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(spm,[status(thm)],[c_0_207, c_0_175])).
 % 49.35/49.55  cnf(c_0_310, negated_conjecture, (ssList(tl(esk49_0))|~neq(esk9_2(esk49_0,nil),nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_300]), c_0_284]), c_0_276])])).
 % 49.35/49.55  cnf(c_0_311, negated_conjecture, (esk9_2(esk49_0,nil)=esk49_0|ssList(tl(esk49_0))), inference(rw,[status(thm)],[c_0_301, c_0_238])).
 % 49.35/49.55  cnf(c_0_312, negated_conjecture, (esk8_2(esk49_0,nil)=nil|ssList(esk7_2(esk49_0,tl(esk49_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_302]), c_0_52])]), c_0_291])).
 % 49.35/49.55  cnf(c_0_313, negated_conjecture, (esk7_2(esk49_0,tl(esk49_0))=cons(hd(esk49_0),nil)|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_304]), c_0_112]), c_0_126]), c_0_112]), c_0_112]), c_0_126]), c_0_112]), c_0_126]), c_0_305]), c_0_52])])).
 % 49.35/49.55  cnf(c_0_314, negated_conjecture, (esk8_2(esk49_0,nil)=nil|esk49_0!=nil), inference(rw,[status(thm)],[c_0_306, c_0_238])).
 % 49.35/49.55  cnf(c_0_315, negated_conjecture, (~totalorderedP(X1)|~segmentP(X1,esk48_0)|~segmentP(esk49_0,X1)|~ssList(X1)|~neq(esk48_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_307, c_0_45]), c_0_45]), c_0_44])).
 % 49.35/49.55  cnf(c_0_316, plain, (segmentP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_308])).
 % 49.35/49.55  cnf(c_0_317, plain, (segmentP(app(X1,X2),X1)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_250]), c_0_251])).
 % 49.35/49.55  cnf(c_0_318, negated_conjecture, (ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_311]), c_0_298])])).
 % 49.35/49.55  cnf(c_0_319, negated_conjecture, (app(nil,esk9_2(esk49_0,nil))=nil|esk8_2(esk49_0,nil)!=esk49_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_171, c_0_238]), c_0_238]), c_0_238])).
 % 49.35/49.55  cnf(c_0_320, negated_conjecture, (esk8_2(esk49_0,nil)=nil|ssList(cons(hd(esk49_0),nil))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_313]), c_0_314])).
 % 49.35/49.55  cnf(c_0_321, negated_conjecture, (esk49_0=nil|ssList(cons(hd(esk49_0),nil))), inference(spm,[status(thm)],[c_0_294, c_0_313])).
 % 49.35/49.55  cnf(c_0_322, negated_conjecture, (X1=esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))|app(esk8_2(esk49_0,esk48_0),X1)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_85]), c_0_60]), c_0_86])])).
 % 49.35/49.55  cnf(c_0_323, negated_conjecture, (~totalorderedP(X1)|~segmentP(esk49_0,X1)|~ssList(X1)|~neq(nil,X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_315, c_0_238]), c_0_238]), c_0_316])).
 % 49.35/49.55  cnf(c_0_324, negated_conjecture, (esk49_0=nil|segmentP(esk49_0,esk7_2(esk49_0,tl(esk49_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_293]), c_0_318])]), c_0_294])).
 % 49.35/49.55  cnf(c_0_325, plain, (segmentP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_99]), c_0_74])]), c_0_56])).
 % 49.35/49.55  cnf(c_0_326, negated_conjecture, (app(nil,esk9_2(esk49_0,nil))=nil|ssList(cons(hd(esk49_0),nil))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_320]), c_0_321])).
 % 49.35/49.55  cnf(c_0_327, negated_conjecture, (X1=esk6_2(esk49_0,esk8_2(esk49_0,nil))|app(esk8_2(esk49_0,nil),X1)!=esk49_0|~ssList(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_322, c_0_238]), c_0_238])).
 % 49.35/49.55  cnf(c_0_328, negated_conjecture, (esk49_0=nil|~totalorderedP(esk7_2(esk49_0,tl(esk49_0)))|~neq(nil,esk7_2(esk49_0,tl(esk49_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_324]), c_0_294])).
 % 49.35/49.55  fof(c_0_329, plain, (![X2]:(~ssItem(X2)|totalorderedP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])).
 % 49.35/49.55  cnf(c_0_330, negated_conjecture, (segmentP(nil,esk9_2(esk49_0,nil))|ssList(cons(hd(esk49_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326]), c_0_276]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_331, negated_conjecture, (esk6_2(esk49_0,esk8_2(esk49_0,nil))=esk9_2(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_290]), c_0_276])])).
 % 49.35/49.55  cnf(c_0_332, negated_conjecture, (esk49_0=nil|~totalorderedP(cons(hd(esk49_0),nil))|~neq(nil,cons(hd(esk49_0),nil))), inference(spm,[status(thm)],[c_0_328, c_0_313])).
 % 49.35/49.55  cnf(c_0_333, plain, (totalorderedP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_329])).
 % 49.35/49.55  cnf(c_0_334, negated_conjecture, (ssList(cons(hd(esk49_0),nil))|~neq(esk9_2(esk49_0,nil),nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_330]), c_0_284]), c_0_276])])).
 % 49.35/49.55  cnf(c_0_335, negated_conjecture, (esk9_2(esk49_0,nil)=esk49_0|ssList(cons(hd(esk49_0),nil))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_331, c_0_320]), c_0_162])).
 % 49.35/49.55  cnf(c_0_336, negated_conjecture, (totalorderedP(esk50_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 49.35/49.55  fof(c_0_337, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_39])])])).
 % 49.35/49.55  cnf(c_0_338, negated_conjecture, (esk49_0=nil|~neq(nil,cons(hd(esk49_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_333]), c_0_305])])).
 % 49.35/49.55  cnf(c_0_339, negated_conjecture, (ssList(cons(hd(esk49_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_335]), c_0_298])])).
 % 49.35/49.55  cnf(c_0_340, negated_conjecture, (totalorderedP(esk48_0)), inference(rw,[status(thm)],[c_0_336, c_0_45])).
 % 49.35/49.55  cnf(c_0_341, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_337])).
 % 49.35/49.55  cnf(c_0_342, negated_conjecture, (cons(hd(esk49_0),nil)=nil|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_186]), c_0_339]), c_0_74])])).
 % 49.35/49.55  cnf(c_0_343, negated_conjecture, (~neq(esk48_0,esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_175]), c_0_340]), c_0_50]), c_0_51])])).
 % 49.35/49.55  cnf(c_0_344, negated_conjecture, (esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_342]), c_0_74]), c_0_305])])).
 % 49.35/49.55  cnf(c_0_345, negated_conjecture, (~neq(nil,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_343, c_0_238]), c_0_238])).
 % 49.35/49.55  cnf(c_0_346, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_298, c_0_344]), c_0_345]), ['proof']).
 % 49.35/49.55  # SZS output end CNFRefutation
 % 49.35/49.55  # Training examples: 0 positive, 0 negative
 % 49.44/49.66
 % 49.44/49.66  # -------------------------------------------------
 % 49.44/49.66  # User time                : 48.377 s
 % 49.44/49.66  # System time              : 1.080 s
 % 49.44/49.66  # Total time               : 49.457 s
 % 49.44/49.66  # Maximum resident set size: 1876 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
