   --------------------------------------------------------------------------

                   TSTP Solution File: SWC030+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC030+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n120.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:22 EDT 2016

 % Result     : Theorem 0.03s
 % Output     : CNFRefutation 0.03s
 % Verified   :
 % Statistics : Number of formulae       :   41 ( 461 expanded)
 %              Number of clauses        :   31 ( 173 expanded)
 %              Number of leaves         :    5 ( 104 expanded)
 %              Depth                    :   14
 %              Number of atoms          :  126 (4291 expanded)
 %              Number of equality atoms :   83 (2637 expanded)
 %              Maximal formula depth    :   19 (   4 average)
 %              Maximal clause size      :   18 (   2 average)
 %              Maximal term depth       :    2 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ! [X5] :
                         ( ssList(X5)
                        => ! [X6] :
                             ( ssList(X6)
                            => ( app(X5,X6) != X4
                               | app(X6,X5) != X3 ) ) )
                     | ( ( nil != X2
                         | nil = X1 )
                       & ( nil != X1
                         | nil = X2 ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(X1,nil) = X1 ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax84)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(nil,X1) = X1 ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax28)).

 fof(c_0_4,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_5,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ! [X5] :
                           ( ssList(X5)
                          => ! [X6] :
                               ( ssList(X6)
                              => ( app(X5,X6) != X4
                                 | app(X6,X5) != X3 ) ) )
                       | ( ( nil != X2
                           | nil = X1 )
                         & ( nil != X1
                           | nil = X2 ) ) ) ) ) ) ) ),
     inference(assume_negation,[status(cth)],[[4]c_0_0])).

 fof(c_0_6,negated_conjecture,
     ( ssList(esk48_0)
     & ssList(esk49_0)
     & ssList(esk50_0)
     & ssList(esk51_0)
     & esk49_0 = esk51_0
     & esk48_0 = esk50_0
     & ssList(esk52_0)
     & ssList(esk53_0)
     & app(esk52_0,esk53_0) = esk51_0
     & app(esk53_0,esk52_0) = esk50_0
     & ( nil = esk48_0
       | nil = esk49_0 )
     & ( nil != esk49_0
       | nil = esk49_0 )
     & ( nil = esk48_0
       | nil != esk48_0 )
     & ( nil != esk49_0
       | nil != esk48_0 ) ),
     inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).

 fof(c_0_7,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).

 cnf(c_0_8,negated_conjecture,
     ( app(esk52_0,esk53_0) = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_6])).

 cnf(c_0_9,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_6])).

 cnf(c_0_10,plain,
     ( nil = X2
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[7]c_0_7])).

 cnf(c_0_11,negated_conjecture,
     ( app(esk52_0,esk53_0) = esk49_0 ),
     inference(rw,[status(thm)],[[8]c_0_8,[9]c_0_9])).

 cnf(c_0_12,negated_conjecture,
     ( ssList(esk53_0) ),
     inference(split_conjunct,[status(thm)],[[10]c_0_6])).

 cnf(c_0_13,negated_conjecture,
     ( ssList(esk52_0) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_6])).

 cnf(c_0_14,negated_conjecture,
     ( esk53_0 = nil
     | esk49_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[12]c_0_10,c_0_11]),c_0_12]),c_0_13])])).

 cnf(c_0_15,negated_conjecture,
     ( nil = esk49_0
     | nil = esk48_0 ),
     inference(split_conjunct,[status(thm)],[[13]c_0_6])).

 cnf(c_0_16,negated_conjecture,
     ( app(esk53_0,esk52_0) = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[14]c_0_6])).

 cnf(c_0_17,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[15]c_0_6])).

 fof(c_0_18,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(X2,nil) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).

 cnf(c_0_19,negated_conjecture,
     ( esk48_0 = nil
     | esk53_0 = nil ),
     inference(spm,[status(thm)],[[16]c_0_14,[17]c_0_15])).

 fof(c_0_20,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(nil,X2) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

 cnf(c_0_21,negated_conjecture,
     ( app(esk53_0,esk52_0) = esk48_0 ),
     inference(rw,[status(thm)],[[18]c_0_16,[19]c_0_17])).

 cnf(c_0_22,plain,
     ( app(X1,nil) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[20]c_0_18])).

 cnf(c_0_23,negated_conjecture,
     ( app(esk52_0,nil) = esk49_0
     | esk48_0 = nil ),
     inference(spm,[status(thm)],[[21]c_0_11,[22]c_0_19])).

 cnf(c_0_24,plain,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[23]c_0_20])).

 cnf(c_0_25,negated_conjecture,
     ( app(nil,esk52_0) = esk48_0
     | esk48_0 = nil ),
     inference(spm,[status(thm)],[[24]c_0_21,[25]c_0_19])).

 cnf(c_0_26,negated_conjecture,
     ( esk48_0 = nil
     | esk52_0 = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_22,c_0_23]),c_0_13])])).

 cnf(c_0_27,negated_conjecture,
     ( esk48_0 = nil
     | esk52_0 = esk48_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_24,c_0_25]),c_0_13])])).

 cnf(c_0_28,negated_conjecture,
     ( esk49_0 = esk48_0
     | esk48_0 = nil ),
     inference(spm,[status(thm)],[[28]c_0_26,[29]c_0_27])).

 cnf(c_0_29,negated_conjecture,
     ( esk52_0 = nil
     | esk48_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[30]c_0_10,c_0_21]),c_0_13]),c_0_12])])).

 cnf(c_0_30,negated_conjecture,
     ( esk48_0 = nil ),
     inference(spm,[status(thm)],[[31]c_0_15,[32]c_0_28])).

 cnf(c_0_31,negated_conjecture,
     ( esk52_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[33]c_0_29,c_0_30])])).

 cnf(c_0_32,plain,
     ( nil = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[34]c_0_7])).

 cnf(c_0_33,negated_conjecture,
     ( app(esk53_0,nil) = nil ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[35]c_0_21,c_0_30]),[36]c_0_31])).

 cnf(c_0_34,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[37]c_0_4])).

 cnf(c_0_35,plain,
     ( nil = app(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != X1
     | nil != X2 ),
     inference(split_conjunct,[status(thm)],[[38]c_0_7])).

 cnf(c_0_36,negated_conjecture,
     ( nil != esk48_0
     | nil != esk49_0 ),
     inference(split_conjunct,[status(thm)],[[39]c_0_6])).

 cnf(c_0_37,negated_conjecture,
     ( esk53_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[40]c_0_32,c_0_33]),c_0_34]),c_0_12])])).

 cnf(c_0_38,plain,
     ( app(nil,nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_35])]),c_0_34])])).

 cnf(c_0_39,negated_conjecture,
     ( esk49_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[41]c_0_36,c_0_30])])).

 cnf(c_0_40,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[42]c_0_11,c_0_31]),c_0_37]),c_0_38]),[43]c_0_39]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC030+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.22  % Computer   : n120.star.cs.uiowa.edu
 % 0.03/0.22  % Model      : x86_64 x86_64
 % 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.22  % Memory     : 32218.75MB
 % 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.23  % CPULimit   : 300
 % 0.03/0.23  % DateTime   : Tue May 10 20:34:38 CDT 2016
 % 0.03/0.23  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.03/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.03/0.25  # and selection function SelectCQPrecWNTNp.
 % 0.03/0.25  #
 % 0.03/0.25  # Preprocessing time       : 0.014 s
 % 0.03/0.25  # Presaturation interreduction done
 % 0.03/0.25
 % 0.03/0.25  # Proof found!
 % 0.03/0.25  # SZS status Theorem
 % 0.03/0.25  # SZS output start CNFRefutation
 % 0.03/0.25  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(X5,X6)!=X4|app(X6,X5)!=X3))))|((nil!=X2|nil=X1)&(nil!=X1|nil=X2)))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
 % 0.03/0.25  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 0.03/0.25  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax84)).
 % 0.03/0.25  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 0.03/0.25  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.03/0.25  fof(c_0_5, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(X5,X6)!=X4|app(X6,X5)!=X3))))|((nil!=X2|nil=X1)&(nil!=X1|nil=X2))))))))), inference(assume_negation,[status(cth)],[c_0_0])).
 % 0.03/0.25  fof(c_0_6, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((esk49_0=esk51_0&esk48_0=esk50_0)&(ssList(esk52_0)&(ssList(esk53_0)&(app(esk52_0,esk53_0)=esk51_0&app(esk53_0,esk52_0)=esk50_0))))&(((nil=esk48_0|nil=esk49_0)&(nil!=esk49_0|nil=esk49_0))&((nil=esk48_0|nil!=esk48_0)&(nil!=esk49_0|nil!=esk48_0))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
 % 0.03/0.25  fof(c_0_7, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
 % 0.03/0.25  cnf(c_0_8, negated_conjecture, (app(esk52_0,esk53_0)=esk51_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.03/0.25  cnf(c_0_9, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.03/0.25  cnf(c_0_10, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.03/0.25  cnf(c_0_11, negated_conjecture, (app(esk52_0,esk53_0)=esk49_0), inference(rw,[status(thm)],[c_0_8, c_0_9])).
 % 0.03/0.25  cnf(c_0_12, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.03/0.25  cnf(c_0_13, negated_conjecture, (ssList(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.03/0.25  cnf(c_0_14, negated_conjecture, (esk53_0=nil|esk49_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13])])).
 % 0.03/0.25  cnf(c_0_15, negated_conjecture, (nil=esk49_0|nil=esk48_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.03/0.25  cnf(c_0_16, negated_conjecture, (app(esk53_0,esk52_0)=esk50_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.03/0.25  cnf(c_0_17, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.03/0.25  fof(c_0_18, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
 % 0.03/0.25  cnf(c_0_19, negated_conjecture, (esk48_0=nil|esk53_0=nil), inference(spm,[status(thm)],[c_0_14, c_0_15])).
 % 0.03/0.25  fof(c_0_20, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
 % 0.03/0.25  cnf(c_0_21, negated_conjecture, (app(esk53_0,esk52_0)=esk48_0), inference(rw,[status(thm)],[c_0_16, c_0_17])).
 % 0.03/0.25  cnf(c_0_22, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
 % 0.03/0.25  cnf(c_0_23, negated_conjecture, (app(esk52_0,nil)=esk49_0|esk48_0=nil), inference(spm,[status(thm)],[c_0_11, c_0_19])).
 % 0.03/0.25  cnf(c_0_24, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
 % 0.03/0.25  cnf(c_0_25, negated_conjecture, (app(nil,esk52_0)=esk48_0|esk48_0=nil), inference(spm,[status(thm)],[c_0_21, c_0_19])).
 % 0.03/0.25  cnf(c_0_26, negated_conjecture, (esk48_0=nil|esk52_0=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_13])])).
 % 0.03/0.25  cnf(c_0_27, negated_conjecture, (esk48_0=nil|esk52_0=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_13])])).
 % 0.03/0.25  cnf(c_0_28, negated_conjecture, (esk49_0=esk48_0|esk48_0=nil), inference(spm,[status(thm)],[c_0_26, c_0_27])).
 % 0.03/0.25  cnf(c_0_29, negated_conjecture, (esk52_0=nil|esk48_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_21]), c_0_13]), c_0_12])])).
 % 0.03/0.25  cnf(c_0_30, negated_conjecture, (esk48_0=nil), inference(spm,[status(thm)],[c_0_15, c_0_28])).
 % 0.03/0.25  cnf(c_0_31, negated_conjecture, (esk52_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30])])).
 % 0.03/0.25  cnf(c_0_32, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.03/0.25  cnf(c_0_33, negated_conjecture, (app(esk53_0,nil)=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_30]), c_0_31])).
 % 0.03/0.25  cnf(c_0_34, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
 % 0.03/0.25  cnf(c_0_35, plain, (nil=app(X1,X2)|~ssList(X1)|~ssList(X2)|nil!=X1|nil!=X2), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.03/0.25  cnf(c_0_36, negated_conjecture, (nil!=esk48_0|nil!=esk49_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.03/0.25  cnf(c_0_37, negated_conjecture, (esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_12])])).
 % 0.03/0.25  cnf(c_0_38, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_35])]), c_0_34])])).
 % 0.03/0.25  cnf(c_0_39, negated_conjecture, (esk49_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_30])])).
 % 0.03/0.25  cnf(c_0_40, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_31]), c_0_37]), c_0_38]), c_0_39]), ['proof']).
 % 0.03/0.25  # SZS output end CNFRefutation
 % 0.03/0.25  # Training examples: 0 positive, 0 negative
 % 0.03/0.25
 % 0.03/0.25  # -------------------------------------------------
 % 0.03/0.25  # User time                : 0.023 s
 % 0.03/0.25  # System time              : 0.005 s
 % 0.03/0.25  # Total time               : 0.029 s
 % 0.03/0.25  # Maximum resident set size: 1872 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
