   --------------------------------------------------------------------------

                   TSTP Solution File: SWC005+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC005+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n110.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:15 EDT 2016

 % Result     : Theorem 0.03s
 % Output     : CNFRefutation 0.03s
 % Verified   :
 % Statistics : Number of formulae       :   59 ( 183 expanded)
 %              Number of clauses        :   39 (  76 expanded)
 %              Number of leaves         :   10 (  45 expanded)
 %              Depth                    :   19
 %              Number of atoms          :  224 (1981 expanded)
 %              Number of equality atoms :   77 ( 608 expanded)
 %              Maximal formula depth    :   23 (   5 average)
 %              Maximal clause size      :   34 (   3 average)
 %              Maximal term depth       :    3 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ( ( ~ neq(X2,nil)
                         | ? [X5] :
                             ( ssList(X5)
                             & ? [X6] :
                                 ( ssList(X6)
                                 & ? [X7] :
                                     ( ssList(X7)
                                     & app(app(X5,X6),X7) = X2
                                     & app(X5,X7) = X1
                                     & neq(X6,nil) ) ) )
                         | ? [X8] :
                             ( ssList(X8)
                             & X3 != X8
                             & tl(X4) = X8
                             & neq(nil,X4) ) )
                       & ( ~ neq(X2,nil)
                         | neq(X4,nil) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(nil,X1) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

 fof(c_0_2,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssList(tl(X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => cons(hd(X1),tl(X1)) = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax78)).

 fof(c_0_8,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssItem(hd(X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax22)).

 fof(c_0_9,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) != X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax18)).

 fof(c_0_10,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ( ( ~ neq(X2,nil)
                           | ? [X5] :
                               ( ssList(X5)
                               & ? [X6] :
                                   ( ssList(X6)
                                   & ? [X7] :
                                       ( ssList(X7)
                                       & app(app(X5,X6),X7) = X2
                                       & app(X5,X7) = X1
                                       & neq(X6,nil) ) ) )
                           | ? [X8] :
                               ( ssList(X8)
                               & X3 != X8
                               & tl(X4) = X8
                               & neq(nil,X4) ) )
                         & ( ~ neq(X2,nil)
                           | neq(X4,nil) ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_11,negated_conjecture,(
     ! [X13,X14,X15,X16] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ( neq(esk49_0,nil)
         | neq(esk49_0,nil) )
       & ( ~ neq(esk51_0,nil)
         | neq(esk49_0,nil) )
       & ( neq(esk49_0,nil)
         | ~ ssList(X13)
         | ~ ssList(X14)
         | ~ ssList(X15)
         | app(app(X13,X14),X15) != esk49_0
         | app(X13,X15) != esk48_0
         | ~ neq(X14,nil) )
       & ( ~ neq(esk51_0,nil)
         | ~ ssList(X13)
         | ~ ssList(X14)
         | ~ ssList(X15)
         | app(app(X13,X14),X15) != esk49_0
         | app(X13,X15) != esk48_0
         | ~ neq(X14,nil) )
       & ( neq(esk49_0,nil)
         | ~ ssList(X16)
         | esk50_0 = X16
         | tl(esk51_0) != X16
         | ~ neq(nil,esk51_0) )
       & ( ~ neq(esk51_0,nil)
         | ~ ssList(X16)
         | esk50_0 = X16
         | tl(esk51_0) != X16
         | ~ neq(nil,esk51_0) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])])).

 cnf(c_0_12,negated_conjecture,
     ( neq(esk49_0,nil)
     | neq(esk49_0,nil) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_11])).

 cnf(c_0_13,negated_conjecture,
     ( neq(esk49_0,nil) ),
     inference(cn,[status(thm)],[[5]c_0_12])).

 cnf(c_0_14,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_11])).

 cnf(c_0_15,negated_conjecture,
     ( ~ neq(X1,nil)
     | app(X2,X3) != esk48_0
     | app(app(X2,X1),X3) != esk49_0
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[7]c_0_11])).

 cnf(c_0_16,negated_conjecture,
     ( neq(esk51_0,nil) ),
     inference(rw,[status(thm)],[[8]c_0_13,[9]c_0_14])).

 fof(c_0_17,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(nil,X2) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).

 cnf(c_0_18,negated_conjecture,
     ( app(app(X1,X2),X3) != esk51_0
     | app(X1,X3) != esk48_0
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ neq(X2,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[10]c_0_15,c_0_14]),c_0_16])])).

 cnf(c_0_19,plain,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_17])).

 cnf(c_0_20,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_2])).

 cnf(c_0_21,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[13]c_0_11])).

 cnf(c_0_22,negated_conjecture,
     ( esk50_0 = X1
     | ~ neq(nil,esk51_0)
     | tl(esk51_0) != X1
     | ~ ssList(X1)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_11])).

 cnf(c_0_23,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[15]c_0_11])).

 fof(c_0_24,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssList(tl(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

 cnf(c_0_25,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_11])).

 cnf(c_0_26,negated_conjecture,
     ( app(app(nil,X1),esk48_0) != esk51_0
     | ~ ssList(X1)
     | ~ neq(X1,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_18,c_0_19]),c_0_20])])]),c_0_21])])).

 fof(c_0_27,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

 cnf(c_0_28,negated_conjecture,
     ( tl(esk51_0) = esk48_0
     | ~ ssList(tl(esk51_0))
     | ~ neq(nil,esk51_0) ),
     inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[18]c_0_22,c_0_23]),c_0_16])])])).

 cnf(c_0_29,plain,
     ( ssList(tl(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_24])).

 cnf(c_0_30,negated_conjecture,
     ( ssList(esk51_0) ),
     inference(rw,[status(thm)],[[20]c_0_25,[21]c_0_14])).

 fof(c_0_31,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).

 cnf(c_0_32,negated_conjecture,
     ( app(X1,esk48_0) != esk51_0
     | ~ ssList(X1)
     | ~ neq(X1,nil) ),
     inference(spm,[status(thm)],[[22]c_0_26,[23]c_0_19])).

 cnf(c_0_33,plain,
     ( cons(X1,X2) = app(cons(X1,nil),X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[24]c_0_27])).

 fof(c_0_34,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).

 fof(c_0_35,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | cons(hd(X2),tl(X2)) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

 cnf(c_0_36,negated_conjecture,
     ( tl(esk51_0) = esk48_0
     | esk51_0 = nil
     | ~ neq(nil,esk51_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_28,c_0_29]),c_0_30])])).

 cnf(c_0_37,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[26]c_0_31])).

 cnf(c_0_38,negated_conjecture,
     ( cons(X1,esk48_0) != esk51_0
     | ~ ssList(cons(X1,nil))
     | ~ neq(cons(X1,nil),nil)
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_32,c_0_33]),c_0_21])])).

 cnf(c_0_39,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[28]c_0_34])).

 cnf(c_0_40,plain,
     ( cons(hd(X1),tl(X1)) = X1
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[29]c_0_35])).

 cnf(c_0_41,negated_conjecture,
     ( tl(esk51_0) = esk48_0
     | esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[30]c_0_36,c_0_37]),c_0_30]),c_0_20])])).

 cnf(c_0_42,negated_conjecture,
     ( cons(X1,esk48_0) != esk51_0
     | ~ neq(cons(X1,nil),nil)
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[31]c_0_38,c_0_39]),c_0_20])])).

 cnf(c_0_43,negated_conjecture,
     ( cons(hd(esk51_0),esk48_0) = esk51_0
     | esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[32]c_0_40,c_0_41]),c_0_30])])).

 fof(c_0_44,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssItem(hd(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).

 cnf(c_0_45,negated_conjecture,
     ( esk51_0 = nil
     | ~ neq(cons(hd(esk51_0),nil),nil)
     | ~ ssItem(hd(esk51_0)) ),
     inference(spm,[status(thm)],[[33]c_0_42,[34]c_0_43])).

 cnf(c_0_46,plain,
     ( ssItem(hd(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[35]c_0_44])).

 cnf(c_0_47,negated_conjecture,
     ( esk51_0 = nil
     | ~ neq(cons(hd(esk51_0),nil),nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[36]c_0_45,c_0_46]),c_0_30])])).

 cnf(c_0_48,negated_conjecture,
     ( cons(hd(esk51_0),nil) = nil
     | esk51_0 = nil
     | ~ ssList(cons(hd(esk51_0),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[37]c_0_47,c_0_37]),c_0_20])])).

 fof(c_0_49,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) != X3 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

 cnf(c_0_50,negated_conjecture,
     ( cons(hd(esk51_0),nil) = nil
     | esk51_0 = nil
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[38]c_0_48,c_0_39]),c_0_20])])).

 cnf(c_0_51,plain,
     ( cons(X1,X2) != X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[39]c_0_49])).

 cnf(c_0_52,negated_conjecture,
     ( cons(hd(esk51_0),nil) = nil
     | esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[40]c_0_50,c_0_46]),c_0_30])])).

 cnf(c_0_53,negated_conjecture,
     ( esk51_0 = nil
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[41]c_0_51,c_0_52]),c_0_20])])).

 cnf(c_0_54,plain,
     ( ~ ssList(X1)
     | ~ ssList(X2)
     | X1 != X2
     | ~ neq(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[42]c_0_31])).

 cnf(c_0_55,negated_conjecture,
     ( esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[43]c_0_53,c_0_46]),c_0_30])])).

 cnf(c_0_56,plain,
     ( ~ ssList(X1)
     | ~ neq(X1,X1) ),
     inference(er,[status(thm)],[[44]c_0_54])).

 cnf(c_0_57,negated_conjecture,
     ( neq(nil,nil) ),
     inference(rw,[status(thm)],[[45]c_0_16,[46]c_0_55])).

 cnf(c_0_58,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[47]c_0_56,c_0_57]),c_0_20])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.02/0.03  % Problem    : SWC005+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.02/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.23  % Computer   : n110.star.cs.uiowa.edu
 % 0.03/0.23  % Model      : x86_64 x86_64
 % 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.23  % Memory     : 32218.75MB
 % 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.23  % CPULimit   : 300
 % 0.03/0.23  % DateTime   : Tue May 10 20:34:23 CDT 2016
 % 0.03/0.23  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.03/0.27  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.03/0.27  # and selection function SelectCQPrecWNTNp.
 % 0.03/0.27  #
 % 0.03/0.27  # Preprocessing time       : 0.014 s
 % 0.03/0.27  # Presaturation interreduction done
 % 0.03/0.27
 % 0.03/0.27  # Proof found!
 % 0.03/0.27  # SZS status Theorem
 % 0.03/0.27  # SZS output start CNFRefutation
 % 0.03/0.27  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((~(neq(X2,nil))|?[X5]:(ssList(X5)&?[X6]:(ssList(X6)&?[X7]:(((ssList(X7)&app(app(X5,X6),X7)=X2)&app(X5,X7)=X1)&neq(X6,nil)))))|?[X8]:(((ssList(X8)&X3!=X8)&tl(X4)=X8)&neq(nil,X4)))&(~(neq(X2,nil))|neq(X4,nil))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.03/0.27  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 0.03/0.27  fof(c_0_2, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.03/0.27  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
 % 0.03/0.27  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
 % 0.03/0.27  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 0.03/0.27  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.03/0.27  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax78)).
 % 0.03/0.27  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax22)).
 % 0.03/0.27  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax18)).
 % 0.03/0.27  fof(c_0_10, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((~neq(X2,nil)|?[X5]:(ssList(X5)&?[X6]:(ssList(X6)&?[X7]:(((ssList(X7)&app(app(X5,X6),X7)=X2)&app(X5,X7)=X1)&neq(X6,nil)))))|?[X8]:(((ssList(X8)&X3!=X8)&tl(X4)=X8)&neq(nil,X4)))&(~neq(X2,nil)|neq(X4,nil)))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.03/0.27  fof(c_0_11, negated_conjecture, (![X13]:![X14]:![X15]:![X16]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((esk49_0=esk51_0&esk48_0=esk50_0)&((((neq(esk49_0,nil)|neq(esk49_0,nil))&(~neq(esk51_0,nil)|neq(esk49_0,nil)))&((neq(esk49_0,nil)|(~ssList(X13)|(~ssList(X14)|(((~ssList(X15)|app(app(X13,X14),X15)!=esk49_0)|app(X13,X15)!=esk48_0)|~neq(X14,nil)))))&(~neq(esk51_0,nil)|(~ssList(X13)|(~ssList(X14)|(((~ssList(X15)|app(app(X13,X14),X15)!=esk49_0)|app(X13,X15)!=esk48_0)|~neq(X14,nil)))))))&((neq(esk49_0,nil)|(((~ssList(X16)|esk50_0=X16)|tl(esk51_0)!=X16)|~neq(nil,esk51_0)))&(~neq(esk51_0,nil)|(((~ssList(X16)|esk50_0=X16)|tl(esk51_0)!=X16)|~neq(nil,esk51_0))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])])).
 % 0.03/0.27  cnf(c_0_12, negated_conjecture, (neq(esk49_0,nil)|neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.03/0.27  cnf(c_0_13, negated_conjecture, (neq(esk49_0,nil)), inference(cn,[status(thm)],[c_0_12])).
 % 0.03/0.27  cnf(c_0_14, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.03/0.27  cnf(c_0_15, negated_conjecture, (~neq(X1,nil)|app(X2,X3)!=esk48_0|app(app(X2,X1),X3)!=esk49_0|~ssList(X3)|~ssList(X1)|~ssList(X2)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.03/0.27  cnf(c_0_16, negated_conjecture, (neq(esk51_0,nil)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
 % 0.03/0.27  fof(c_0_17, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
 % 0.03/0.27  cnf(c_0_18, negated_conjecture, (app(app(X1,X2),X3)!=esk51_0|app(X1,X3)!=esk48_0|~ssList(X3)|~ssList(X1)|~ssList(X2)|~neq(X2,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_14]), c_0_16])])).
 % 0.03/0.27  cnf(c_0_19, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.03/0.27  cnf(c_0_20, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.27  cnf(c_0_21, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.03/0.27  cnf(c_0_22, negated_conjecture, (esk50_0=X1|~neq(nil,esk51_0)|tl(esk51_0)!=X1|~ssList(X1)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.03/0.27  cnf(c_0_23, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.03/0.27  fof(c_0_24, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
 % 0.03/0.27  cnf(c_0_25, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.03/0.27  cnf(c_0_26, negated_conjecture, (app(app(nil,X1),esk48_0)!=esk51_0|~ssList(X1)|~neq(X1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])])]), c_0_21])])).
 % 0.03/0.27  fof(c_0_27, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 0.03/0.27  cnf(c_0_28, negated_conjecture, (tl(esk51_0)=esk48_0|~ssList(tl(esk51_0))|~neq(nil,esk51_0)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23]), c_0_16])])])).
 % 0.03/0.27  cnf(c_0_29, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
 % 0.03/0.27  cnf(c_0_30, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_25, c_0_14])).
 % 0.03/0.27  fof(c_0_31, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
 % 0.03/0.27  cnf(c_0_32, negated_conjecture, (app(X1,esk48_0)!=esk51_0|~ssList(X1)|~neq(X1,nil)), inference(spm,[status(thm)],[c_0_26, c_0_19])).
 % 0.03/0.27  cnf(c_0_33, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
 % 0.03/0.27  fof(c_0_34, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
 % 0.03/0.27  fof(c_0_35, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
 % 0.03/0.27  cnf(c_0_36, negated_conjecture, (tl(esk51_0)=esk48_0|esk51_0=nil|~neq(nil,esk51_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
 % 0.03/0.27  cnf(c_0_37, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
 % 0.03/0.27  cnf(c_0_38, negated_conjecture, (cons(X1,esk48_0)!=esk51_0|~ssList(cons(X1,nil))|~neq(cons(X1,nil),nil)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_21])])).
 % 0.03/0.27  cnf(c_0_39, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
 % 0.03/0.27  cnf(c_0_40, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
 % 0.03/0.27  cnf(c_0_41, negated_conjecture, (tl(esk51_0)=esk48_0|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_30]), c_0_20])])).
 % 0.03/0.27  cnf(c_0_42, negated_conjecture, (cons(X1,esk48_0)!=esk51_0|~neq(cons(X1,nil),nil)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_20])])).
 % 0.03/0.27  cnf(c_0_43, negated_conjecture, (cons(hd(esk51_0),esk48_0)=esk51_0|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_30])])).
 % 0.03/0.27  fof(c_0_44, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
 % 0.03/0.27  cnf(c_0_45, negated_conjecture, (esk51_0=nil|~neq(cons(hd(esk51_0),nil),nil)|~ssItem(hd(esk51_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
 % 0.03/0.27  cnf(c_0_46, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
 % 0.03/0.27  cnf(c_0_47, negated_conjecture, (esk51_0=nil|~neq(cons(hd(esk51_0),nil),nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_30])])).
 % 0.03/0.27  cnf(c_0_48, negated_conjecture, (cons(hd(esk51_0),nil)=nil|esk51_0=nil|~ssList(cons(hd(esk51_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_37]), c_0_20])])).
 % 0.03/0.27  fof(c_0_49, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
 % 0.03/0.27  cnf(c_0_50, negated_conjecture, (cons(hd(esk51_0),nil)=nil|esk51_0=nil|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_39]), c_0_20])])).
 % 0.03/0.27  cnf(c_0_51, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
 % 0.03/0.27  cnf(c_0_52, negated_conjecture, (cons(hd(esk51_0),nil)=nil|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_46]), c_0_30])])).
 % 0.03/0.27  cnf(c_0_53, negated_conjecture, (esk51_0=nil|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_20])])).
 % 0.03/0.27  cnf(c_0_54, plain, (~ssList(X1)|~ssList(X2)|X1!=X2|~neq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
 % 0.03/0.27  cnf(c_0_55, negated_conjecture, (esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_46]), c_0_30])])).
 % 0.03/0.27  cnf(c_0_56, plain, (~ssList(X1)|~neq(X1,X1)), inference(er,[status(thm)],[c_0_54])).
 % 0.03/0.27  cnf(c_0_57, negated_conjecture, (neq(nil,nil)), inference(rw,[status(thm)],[c_0_16, c_0_55])).
 % 0.03/0.27  cnf(c_0_58, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_20])]), ['proof']).
 % 0.03/0.27  # SZS output end CNFRefutation
 % 0.03/0.27  # Training examples: 0 positive, 0 negative
 % 0.03/0.27
 % 0.03/0.27  # -------------------------------------------------
 % 0.03/0.27  # User time                : 0.037 s
 % 0.03/0.27  # System time              : 0.002 s
 % 0.03/0.27  # Total time               : 0.039 s
 % 0.03/0.27  # Maximum resident set size: 1876 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
