   --------------------------------------------------------------------------

                   TSTP Solution File: SWC030-1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC030-1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n136.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:22 EDT 2016

 % Result     : Unsatisfiable 0.03s
 % Output     : CNFRefutation 0.03s
 % Verified   :
 % Statistics : Number of clauses        :   30 ( 172 expanded)
 %              Number of leaves         :   13 ( 103 expanded)
 %              Depth                    :   11
 %              Number of atoms          :   49 ( 266 expanded)
 %              Number of equality atoms :   39 ( 204 expanded)
 %              Maximal clause size      :    5 (   2 average)
 %              Maximal term depth       :    2 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 cnf(c_0_0,negated_conjecture,
     ( app(sk5,sk6) = sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

 cnf(c_0_1,negated_conjecture,
     ( sk2 = sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

 cnf(c_0_2,axiom,
     ( nil = X2
     | app(X1,X2) != nil
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause119)).

 cnf(c_0_3,negated_conjecture,
     ( ssList(sk6) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

 cnf(c_0_4,negated_conjecture,
     ( ssList(sk5) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

 cnf(c_0_5,negated_conjecture,
     ( nil = sk2
     | nil = sk1 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_11)).

 cnf(c_0_6,negated_conjecture,
     ( app(sk6,sk5) = sk3 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_10)).

 cnf(c_0_7,negated_conjecture,
     ( sk1 = sk3 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

 cnf(c_0_8,axiom,
     ( app(X1,nil) = X1
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause73)).

 cnf(c_0_9,axiom,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause74)).

 cnf(c_0_10,axiom,
     ( app(X1,X2) = nil
     | nil != X1
     | nil != X2
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause145)).

 cnf(c_0_11,axiom,
     ( ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

 cnf(c_0_12,negated_conjecture,
     ( nil != sk1
     | nil != sk2 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_14)).

 cnf(c_0_13,negated_conjecture,
     ( app(sk5,sk6) = sk2 ),
     inference(rw,[status(thm)],[[4]c_0_0,[5]c_0_1])).

 cnf(c_0_14,negated_conjecture,
     ( sk6 = nil
     | sk2 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[6]c_0_2,c_0_13]),c_0_3]),c_0_4])])).

 cnf(c_0_15,negated_conjecture,
     ( sk1 = nil
     | sk6 = nil ),
     inference(spm,[status(thm)],[[7]c_0_14,[8]c_0_5])).

 cnf(c_0_16,negated_conjecture,
     ( app(sk6,sk5) = sk1 ),
     inference(rw,[status(thm)],[[9]c_0_6,[10]c_0_7])).

 cnf(c_0_17,negated_conjecture,
     ( app(sk5,nil) = sk2
     | sk1 = nil ),
     inference(spm,[status(thm)],[[11]c_0_13,[12]c_0_15])).

 cnf(c_0_18,negated_conjecture,
     ( app(nil,sk5) = sk1
     | sk1 = nil ),
     inference(spm,[status(thm)],[[13]c_0_16,[14]c_0_15])).

 cnf(c_0_19,negated_conjecture,
     ( sk1 = nil
     | sk5 = sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_8,c_0_17]),c_0_4])])).

 cnf(c_0_20,negated_conjecture,
     ( sk1 = nil
     | sk5 = sk1 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[16]c_0_9,c_0_18]),c_0_4])])).

 cnf(c_0_21,negated_conjecture,
     ( sk2 = sk1
     | sk1 = nil ),
     inference(spm,[status(thm)],[[17]c_0_19,[18]c_0_20])).

 cnf(c_0_22,negated_conjecture,
     ( sk5 = nil
     | sk1 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[19]c_0_2,c_0_16]),c_0_4]),c_0_3])])).

 cnf(c_0_23,negated_conjecture,
     ( sk1 = nil ),
     inference(spm,[status(thm)],[[20]c_0_5,[21]c_0_21])).

 cnf(c_0_24,negated_conjecture,
     ( sk5 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[22]c_0_22,c_0_23])])).

 cnf(c_0_25,negated_conjecture,
     ( app(sk6,nil) = nil ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[23]c_0_16,c_0_23]),[24]c_0_24])).

 cnf(c_0_26,negated_conjecture,
     ( sk6 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_8,c_0_25]),c_0_3])])).

 cnf(c_0_27,plain,
     ( app(nil,nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_10])]),c_0_11])])).

 cnf(c_0_28,negated_conjecture,
     ( sk2 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[26]c_0_12,c_0_23])])).

 cnf(c_0_29,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[27]c_0_13,c_0_24]),c_0_26]),c_0_27]),[28]c_0_28]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC030-1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.23  % Computer   : n136.star.cs.uiowa.edu
 % 0.03/0.23  % Model      : x86_64 x86_64
 % 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.23  % Memory     : 32218.75MB
 % 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.23  % CPULimit   : 300
 % 0.03/0.23  % DateTime   : Tue May 10 20:34:38 CDT 2016
 % 0.03/0.23  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.03/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
 % 0.03/0.25  # and selection function SelectCQArNTEqFirst.
 % 0.03/0.25  #
 % 0.03/0.25  # Preprocessing time       : 0.009 s
 % 0.03/0.25  # Presaturation interreduction done
 % 0.03/0.25
 % 0.03/0.25  # Proof found!
 % 0.03/0.25  # SZS status Unsatisfiable
 % 0.03/0.25  # SZS output start CNFRefutation
 % 0.03/0.25  cnf(c_0_0, negated_conjecture, (app(sk5,sk6)=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
 % 0.03/0.25  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
 % 0.03/0.25  cnf(c_0_2, axiom, (nil=X2|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause119)).
 % 0.03/0.25  cnf(c_0_3, negated_conjecture, (ssList(sk6)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
 % 0.03/0.25  cnf(c_0_4, negated_conjecture, (ssList(sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
 % 0.03/0.25  cnf(c_0_5, negated_conjecture, (nil=sk2|nil=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_11)).
 % 0.03/0.25  cnf(c_0_6, negated_conjecture, (app(sk6,sk5)=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_10)).
 % 0.03/0.25  cnf(c_0_7, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
 % 0.03/0.25  cnf(c_0_8, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause73)).
 % 0.03/0.25  cnf(c_0_9, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause74)).
 % 0.03/0.25  cnf(c_0_10, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause145)).
 % 0.03/0.25  cnf(c_0_11, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
 % 0.03/0.25  cnf(c_0_12, negated_conjecture, (nil!=sk1|nil!=sk2), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_14)).
 % 0.03/0.25  cnf(c_0_13, negated_conjecture, (app(sk5,sk6)=sk2), inference(rw,[status(thm)],[c_0_0, c_0_1])).
 % 0.03/0.25  cnf(c_0_14, negated_conjecture, (sk6=nil|sk2!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_13]), c_0_3]), c_0_4])])).
 % 0.03/0.25  cnf(c_0_15, negated_conjecture, (sk1=nil|sk6=nil), inference(spm,[status(thm)],[c_0_14, c_0_5])).
 % 0.03/0.25  cnf(c_0_16, negated_conjecture, (app(sk6,sk5)=sk1), inference(rw,[status(thm)],[c_0_6, c_0_7])).
 % 0.03/0.25  cnf(c_0_17, negated_conjecture, (app(sk5,nil)=sk2|sk1=nil), inference(spm,[status(thm)],[c_0_13, c_0_15])).
 % 0.03/0.25  cnf(c_0_18, negated_conjecture, (app(nil,sk5)=sk1|sk1=nil), inference(spm,[status(thm)],[c_0_16, c_0_15])).
 % 0.03/0.25  cnf(c_0_19, negated_conjecture, (sk1=nil|sk5=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_17]), c_0_4])])).
 % 0.03/0.25  cnf(c_0_20, negated_conjecture, (sk1=nil|sk5=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_18]), c_0_4])])).
 % 0.03/0.25  cnf(c_0_21, negated_conjecture, (sk2=sk1|sk1=nil), inference(spm,[status(thm)],[c_0_19, c_0_20])).
 % 0.03/0.25  cnf(c_0_22, negated_conjecture, (sk5=nil|sk1!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_16]), c_0_4]), c_0_3])])).
 % 0.03/0.25  cnf(c_0_23, negated_conjecture, (sk1=nil), inference(spm,[status(thm)],[c_0_5, c_0_21])).
 % 0.03/0.25  cnf(c_0_24, negated_conjecture, (sk5=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23])])).
 % 0.03/0.25  cnf(c_0_25, negated_conjecture, (app(sk6,nil)=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_23]), c_0_24])).
 % 0.03/0.25  cnf(c_0_26, negated_conjecture, (sk6=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_25]), c_0_3])])).
 % 0.03/0.25  cnf(c_0_27, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_10])]), c_0_11])])).
 % 0.03/0.25  cnf(c_0_28, negated_conjecture, (sk2!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_23])])).
 % 0.03/0.25  cnf(c_0_29, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_24]), c_0_26]), c_0_27]), c_0_28]), ['proof']).
 % 0.03/0.25  # SZS output end CNFRefutation
 % 0.03/0.25  # Training examples: 0 positive, 0 negative
 % 0.03/0.25
 % 0.03/0.25  # -------------------------------------------------
 % 0.03/0.25  # User time                : 0.019 s
 % 0.03/0.25  # System time              : 0.004 s
 % 0.03/0.25  # Total time               : 0.023 s
 % 0.03/0.25  # Maximum resident set size: 1856 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
