   --------------------------------------------------------------------------

                   TSTP Solution File: SWC053+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC053+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n158.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:27 EDT 2016

 % Result     : Theorem 0.06s
 % Output     : CNFRefutation 0.06s
 % Verified   :
 % Statistics : Number of formulae       :  126 (2049 expanded)
 %              Number of clauses        :   84 ( 754 expanded)
 %              Number of leaves         :   21 ( 503 expanded)
 %              Depth                    :   21
 %              Number of atoms          :  488 (26641 expanded)
 %              Number of equality atoms :  162 (9154 expanded)
 %              Maximal formula depth    :   24 (   5 average)
 %              Maximal clause size      :   36 (   3 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ( ! [X5] :
                           ( ssList(X5)
                          => ( ~ neq(X5,nil)
                             | ~ frontsegP(X4,X5)
                             | ~ frontsegP(X3,X5) ) )
                       & ( nil != X4
                         | nil != X3 ) )
                     | ( ( nil != X2
                         | nil = X1 )
                       & ( ~ neq(X2,nil)
                         | ? [X6] :
                             ( ssList(X6)
                             & neq(X6,nil)
                             & segmentP(X2,X6)
                             & segmentP(X1,X6) ) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( frontsegP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & app(X2,X3) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax5)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil != X1
            => hd(app(X1,X2)) = hd(X1) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax85)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => cons(hd(X1),tl(X1)) = X1 ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax78)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => tl(cons(X2,X1)) = X1 ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax25)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => nil != cons(X2,X1) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax21)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssItem(hd(X1)) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax22)).

 fof(c_0_8,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( frontsegP(cons(X1,X3),cons(X2,X4))
                   <=> ( X1 = X2
                       & frontsegP(X3,X4) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax44)).

 fof(c_0_9,axiom,(
     ! [X1] :
       ( ssList(X1)
      => frontsegP(X1,X1) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax42)).

 fof(c_0_10,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssList(tl(X1)) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax24)).

 fof(c_0_11,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( segmentP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & ? [X4] :
                     ( ssList(X4)
                     & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax7)).

 fof(c_0_12,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(nil,X1) = X1 ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax28)).

 fof(c_0_13,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_14,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( frontsegP(X1,X2)
                => frontsegP(app(X1,X3),X2) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax43)).

 fof(c_0_15,axiom,(
     ! [X1] :
       ( ssList(X1)
      => frontsegP(X1,nil) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax45)).

 fof(c_0_16,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_17,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( ( frontsegP(X1,X2)
               & frontsegP(X2,X1) )
            => X1 = X2 ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax41)).

 fof(c_0_18,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_19,axiom,(
     ! [X1] :
       ( ssList(X1)
      => segmentP(X1,nil) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax57)).

 fof(c_0_20,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_21,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ( ! [X5] :
                             ( ssList(X5)
                            => ( ~ neq(X5,nil)
                               | ~ frontsegP(X4,X5)
                               | ~ frontsegP(X3,X5) ) )
                         & ( nil != X4
                           | nil != X3 ) )
                       | ( ( nil != X2
                           | nil = X1 )
                         & ( ~ neq(X2,nil)
                           | ? [X6] :
                               ( ssList(X6)
                               & neq(X6,nil)
                               & segmentP(X2,X6)
                               & segmentP(X1,X6) ) ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_22,negated_conjecture,(
     ! [X12] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ( nil = esk51_0
         | ssList(esk52_0) )
       & ( nil = esk50_0
         | ssList(esk52_0) )
       & ( nil = esk51_0
         | neq(esk52_0,nil) )
       & ( nil = esk50_0
         | neq(esk52_0,nil) )
       & ( nil = esk51_0
         | frontsegP(esk51_0,esk52_0) )
       & ( nil = esk50_0
         | frontsegP(esk51_0,esk52_0) )
       & ( nil = esk51_0
         | frontsegP(esk50_0,esk52_0) )
       & ( nil = esk50_0
         | frontsegP(esk50_0,esk52_0) )
       & ( neq(esk49_0,nil)
         | nil = esk49_0 )
       & ( ~ ssList(X12)
         | ~ neq(X12,nil)
         | ~ segmentP(esk49_0,X12)
         | ~ segmentP(esk48_0,X12)
         | nil = esk49_0 )
       & ( neq(esk49_0,nil)
         | nil != esk48_0 )
       & ( ~ ssList(X12)
         | ~ neq(X12,nil)
         | ~ segmentP(esk49_0,X12)
         | ~ segmentP(esk48_0,X12)
         | nil != esk48_0 ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).

 fof(c_0_23,plain,(
     ! [X4,X5,X7] :
       ( ( ssList(esk6_2(X4,X5))
         | ~ frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( app(X5,esk6_2(X4,X5)) = X4
         | ~ frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( ~ ssList(X7)
         | app(X5,X7) != X4
         | frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

 cnf(c_0_24,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_22])).

 cnf(c_0_25,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_22])).

 fof(c_0_26,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | nil = X3
       | hd(app(X3,X4)) = hd(X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

 cnf(c_0_27,plain,
     ( app(X2,esk6_2(X1,X2)) = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ frontsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[6]c_0_23])).

 cnf(c_0_28,negated_conjecture,
     ( frontsegP(esk51_0,esk52_0)
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[7]c_0_22])).

 cnf(c_0_29,negated_conjecture,
     ( ssList(esk51_0) ),
     inference(rw,[status(thm)],[[8]c_0_24,[9]c_0_25])).

 cnf(c_0_30,negated_conjecture,
     ( ssList(esk52_0)
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[10]c_0_22])).

 cnf(c_0_31,plain,
     ( ssList(esk6_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ frontsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_23])).

 fof(c_0_32,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | cons(hd(X2),tl(X2)) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

 fof(c_0_33,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | tl(cons(X4,X3)) = X3 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

 fof(c_0_34,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

 fof(c_0_35,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | nil != cons(X4,X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).

 cnf(c_0_36,plain,
     ( hd(app(X1,X2)) = hd(X1)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_26])).

 cnf(c_0_37,negated_conjecture,
     ( app(esk52_0,esk6_2(esk51_0,esk52_0)) = esk51_0
     | esk50_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[13]c_0_27,c_0_28]),c_0_29])]),[14]c_0_30])).

 cnf(c_0_38,negated_conjecture,
     ( esk50_0 = nil
     | ssList(esk6_2(esk51_0,esk52_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_31,c_0_28]),c_0_29])]),[16]c_0_30])).

 fof(c_0_39,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssItem(hd(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

 fof(c_0_40,plain,(
     ! [X5,X6,X7,X8] :
       ( ( X5 = X6
         | ~ frontsegP(cons(X5,X7),cons(X6,X8))
         | ~ ssList(X8)
         | ~ ssList(X7)
         | ~ ssItem(X6)
         | ~ ssItem(X5) )
       & ( frontsegP(X7,X8)
         | ~ frontsegP(cons(X5,X7),cons(X6,X8))
         | ~ ssList(X8)
         | ~ ssList(X7)
         | ~ ssItem(X6)
         | ~ ssItem(X5) )
       & ( X5 != X6
         | ~ frontsegP(X7,X8)
         | frontsegP(cons(X5,X7),cons(X6,X8))
         | ~ ssList(X8)
         | ~ ssList(X7)
         | ~ ssItem(X6)
         | ~ ssItem(X5) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).

 cnf(c_0_41,plain,
     ( cons(hd(X1),tl(X1)) = X1
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[17]c_0_32])).

 cnf(c_0_42,plain,
     ( tl(cons(X1,X2)) = X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[18]c_0_33])).

 cnf(c_0_43,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_34])).

 cnf(c_0_44,plain,
     ( nil != cons(X1,X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[20]c_0_35])).

 cnf(c_0_45,negated_conjecture,
     ( hd(esk52_0) = hd(esk51_0)
     | esk50_0 = nil
     | esk52_0 = nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[21]c_0_36,c_0_37]),c_0_30]),[22]c_0_38])).

 cnf(c_0_46,plain,
     ( ssItem(hd(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[23]c_0_39])).

 cnf(c_0_47,plain,
     ( frontsegP(cons(X1,X3),cons(X2,X4))
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | ~ ssList(X4)
     | ~ frontsegP(X3,X4)
     | X1 != X2 ),
     inference(split_conjunct,[status(thm)],[[24]c_0_40])).

 fof(c_0_48,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | frontsegP(X2,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).

 cnf(c_0_49,plain,
     ( cons(hd(cons(X1,X2)),X2) = cons(X1,X2)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_41,c_0_42]),c_0_43]),[26]c_0_44])).

 cnf(c_0_50,negated_conjecture,
     ( cons(hd(esk51_0),tl(esk52_0)) = esk52_0
     | esk50_0 = nil
     | esk52_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_41,c_0_45]),[28]c_0_30])).

 cnf(c_0_51,negated_conjecture,
     ( esk50_0 = nil
     | esk52_0 = nil
     | ssItem(hd(esk51_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[29]c_0_46,c_0_45]),[30]c_0_30])).

 fof(c_0_52,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssList(tl(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).

 cnf(c_0_53,plain,
     ( frontsegP(cons(X1,X2),cons(X1,X3))
     | ~ frontsegP(X2,X3)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(er,[status(thm)],[[31]c_0_47])).

 cnf(c_0_54,plain,
     ( frontsegP(X1,X1)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[32]c_0_48])).

 cnf(c_0_55,negated_conjecture,
     ( cons(hd(esk52_0),tl(esk52_0)) = esk52_0
     | esk52_0 = nil
     | esk50_0 = nil
     | ~ ssList(tl(esk52_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[33]c_0_49,c_0_50]),[34]c_0_51])).

 cnf(c_0_56,plain,
     ( ssList(tl(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[35]c_0_52])).

 cnf(c_0_57,plain,
     ( frontsegP(cons(X1,X2),cons(X1,X2))
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(spm,[status(thm)],[[36]c_0_53,[37]c_0_54])).

 cnf(c_0_58,negated_conjecture,
     ( cons(hd(esk52_0),tl(esk52_0)) = esk52_0
     | esk50_0 = nil
     | esk52_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[38]c_0_55,c_0_56]),[39]c_0_30])).

 fof(c_0_59,plain,(
     ! [X5,X6,X9,X10] :
       ( ( ssList(esk8_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ssList(esk9_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ~ ssList(X9)
         | ~ ssList(X10)
         | app(app(X9,X6),X10) != X5
         | segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).

 cnf(c_0_60,negated_conjecture,
     ( esk52_0 = nil
     | esk50_0 = nil
     | frontsegP(esk52_0,esk52_0)
     | ~ ssList(tl(esk52_0))
     | ~ ssItem(hd(esk52_0)) ),
     inference(spm,[status(thm)],[[40]c_0_57,[41]c_0_58])).

 cnf(c_0_61,plain,
     ( segmentP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(app(X3,X2),X4) != X1
     | ~ ssList(X4)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[42]c_0_59])).

 fof(c_0_62,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(nil,X2) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).

 fof(c_0_63,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).

 cnf(c_0_64,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[43]c_0_22])).

 cnf(c_0_65,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[44]c_0_22])).

 fof(c_0_66,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | ~ frontsegP(X4,X5)
       | frontsegP(app(X4,X6),X5) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).

 fof(c_0_67,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | frontsegP(X2,nil) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).

 cnf(c_0_68,negated_conjecture,
     ( esk50_0 = nil
     | esk52_0 = nil
     | frontsegP(esk52_0,esk52_0)
     | ~ ssItem(hd(esk52_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[45]c_0_60,c_0_56]),[46]c_0_30])).

 cnf(c_0_69,plain,
     ( segmentP(app(app(X1,X2),X3),X2)
     | ~ ssList(app(app(X1,X2),X3))
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(er,[status(thm)],[[47]c_0_61])).

 cnf(c_0_70,plain,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[48]c_0_62])).

 cnf(c_0_71,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[49]c_0_16])).

 cnf(c_0_72,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[50]c_0_63])).

 cnf(c_0_73,negated_conjecture,
     ( frontsegP(esk50_0,esk52_0)
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[51]c_0_22])).

 cnf(c_0_74,negated_conjecture,
     ( ssList(esk50_0) ),
     inference(rw,[status(thm)],[[52]c_0_64,[53]c_0_65])).

 cnf(c_0_75,negated_conjecture,
     ( frontsegP(esk51_0,esk52_0)
     | nil = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[54]c_0_22])).

 cnf(c_0_76,negated_conjecture,
     ( ssList(esk52_0)
     | nil = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[55]c_0_22])).

 cnf(c_0_77,plain,
     ( frontsegP(app(X1,X2),X3)
     | ~ frontsegP(X1,X3)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[56]c_0_66])).

 cnf(c_0_78,plain,
     ( frontsegP(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[57]c_0_67])).

 cnf(c_0_79,negated_conjecture,
     ( esk52_0 = nil
     | esk50_0 = nil
     | frontsegP(esk52_0,esk52_0) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[58]c_0_68,c_0_46]),[59]c_0_30])).

 cnf(c_0_80,negated_conjecture,
     ( nil = esk49_0
     | ~ segmentP(esk48_0,X1)
     | ~ segmentP(esk49_0,X1)
     | ~ neq(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[60]c_0_22])).

 cnf(c_0_81,plain,
     ( segmentP(app(X1,X2),X1)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[61]c_0_69,c_0_70]),c_0_71])]),[62]c_0_72])).

 cnf(c_0_82,negated_conjecture,
     ( app(esk52_0,esk6_2(esk50_0,esk52_0)) = esk50_0
     | esk50_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[63]c_0_27,c_0_73]),c_0_74])]),[64]c_0_30])).

 cnf(c_0_83,negated_conjecture,
     ( esk50_0 = nil
     | ssList(esk6_2(esk50_0,esk52_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[65]c_0_31,c_0_73]),c_0_74])]),[66]c_0_30])).

 cnf(c_0_84,negated_conjecture,
     ( app(esk52_0,esk6_2(esk51_0,esk52_0)) = esk51_0
     | esk51_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[67]c_0_27,c_0_75]),c_0_29])]),[68]c_0_76])).

 cnf(c_0_85,negated_conjecture,
     ( esk51_0 = nil
     | ssList(esk6_2(esk51_0,esk52_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[69]c_0_31,c_0_75]),c_0_29])]),[70]c_0_76])).

 fof(c_0_86,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ~ frontsegP(X3,X4)
       | ~ frontsegP(X4,X3)
       | X3 = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).

 cnf(c_0_87,plain,
     ( frontsegP(app(X1,X2),nil)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[71]c_0_77,c_0_78]),c_0_71])])).

 cnf(c_0_88,negated_conjecture,
     ( app(esk52_0,esk6_2(esk52_0,esk52_0)) = esk52_0
     | esk50_0 = nil
     | esk52_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[72]c_0_27,c_0_79]),[73]c_0_30])).

 cnf(c_0_89,negated_conjecture,
     ( esk50_0 = nil
     | esk52_0 = nil
     | ssList(esk6_2(esk52_0,esk52_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[74]c_0_31,c_0_79]),[75]c_0_30])).

 cnf(c_0_90,negated_conjecture,
     ( esk51_0 = nil
     | ~ segmentP(esk50_0,X1)
     | ~ segmentP(esk51_0,X1)
     | ~ ssList(X1)
     | ~ neq(X1,nil) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[76]c_0_80,c_0_25]),c_0_65]),[77]c_0_25])).

 cnf(c_0_91,negated_conjecture,
     ( esk50_0 = nil
     | segmentP(esk50_0,esk52_0) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[78]c_0_81,c_0_82]),c_0_30]),[79]c_0_83])).

 cnf(c_0_92,negated_conjecture,
     ( neq(esk52_0,nil)
     | nil = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[80]c_0_22])).

 cnf(c_0_93,negated_conjecture,
     ( esk51_0 = nil
     | segmentP(esk51_0,esk52_0) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[81]c_0_81,c_0_84]),c_0_76]),[82]c_0_85])).

 fof(c_0_94,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])).

 cnf(c_0_95,plain,
     ( X1 = X2
     | ~ frontsegP(X2,X1)
     | ~ frontsegP(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[83]c_0_86])).

 cnf(c_0_96,negated_conjecture,
     ( esk52_0 = nil
     | esk50_0 = nil
     | frontsegP(esk52_0,nil) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[84]c_0_87,c_0_88]),c_0_30]),[85]c_0_89])).

 cnf(c_0_97,negated_conjecture,
     ( frontsegP(esk50_0,esk52_0)
     | nil = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[86]c_0_22])).

 cnf(c_0_98,negated_conjecture,
     ( esk50_0 = nil
     | esk51_0 = nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[87]c_0_90,c_0_91]),c_0_92]),c_0_76]),[88]c_0_93])).

 cnf(c_0_99,plain,
     ( nil = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[89]c_0_94])).

 cnf(c_0_100,negated_conjecture,
     ( nil != esk48_0
     | ~ segmentP(esk48_0,X1)
     | ~ segmentP(esk49_0,X1)
     | ~ neq(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[90]c_0_22])).

 fof(c_0_101,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | segmentP(X2,nil) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).

 fof(c_0_102,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])).

 cnf(c_0_103,negated_conjecture,
     ( esk50_0 = nil
     | esk52_0 = nil
     | ~ frontsegP(nil,esk52_0) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[91]c_0_95,c_0_96]),c_0_71])]),[92]c_0_30])).

 cnf(c_0_104,negated_conjecture,
     ( esk51_0 = nil
     | frontsegP(nil,esk52_0) ),
     inference(spm,[status(thm)],[[93]c_0_97,[94]c_0_98])).

 cnf(c_0_105,negated_conjecture,
     ( esk50_0 = nil
     | esk52_0 = nil
     | esk51_0 != nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[95]c_0_99,c_0_37]),c_0_30]),[96]c_0_38])).

 cnf(c_0_106,negated_conjecture,
     ( esk50_0 != nil
     | ~ segmentP(esk50_0,X1)
     | ~ segmentP(esk51_0,X1)
     | ~ ssList(X1)
     | ~ neq(X1,nil) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[97]c_0_100,c_0_65]),c_0_65]),[98]c_0_25])).

 cnf(c_0_107,plain,
     ( segmentP(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[99]c_0_101])).

 cnf(c_0_108,plain,
     ( ~ ssList(X1)
     | ~ ssList(X2)
     | X1 != X2
     | ~ neq(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[100]c_0_102])).

 cnf(c_0_109,negated_conjecture,
     ( neq(esk52_0,nil)
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[101]c_0_22])).

 cnf(c_0_110,negated_conjecture,
     ( esk52_0 = nil
     | esk50_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[102]c_0_103,c_0_104]),[103]c_0_105])).

 cnf(c_0_111,negated_conjecture,
     ( app(esk52_0,esk6_2(esk50_0,esk52_0)) = esk50_0
     | esk51_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[104]c_0_27,c_0_97]),c_0_74])]),[105]c_0_76])).

 cnf(c_0_112,negated_conjecture,
     ( esk51_0 = nil
     | ssList(esk6_2(esk50_0,esk52_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[106]c_0_31,c_0_97]),c_0_74])]),[107]c_0_76])).

 cnf(c_0_113,negated_conjecture,
     ( esk50_0 != nil
     | ~ segmentP(esk51_0,nil)
     | ~ neq(nil,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[108]c_0_106,c_0_107]),c_0_71]),c_0_74])])).

 cnf(c_0_114,plain,
     ( ~ ssList(X1)
     | ~ neq(X1,X1) ),
     inference(er,[status(thm)],[[109]c_0_108])).

 cnf(c_0_115,negated_conjecture,
     ( esk50_0 = nil
     | neq(nil,nil) ),
     inference(spm,[status(thm)],[[110]c_0_109,[111]c_0_110])).

 cnf(c_0_116,negated_conjecture,
     ( neq(esk49_0,nil)
     | nil != esk48_0 ),
     inference(split_conjunct,[status(thm)],[[112]c_0_22])).

 cnf(c_0_117,negated_conjecture,
     ( esk51_0 = nil
     | esk52_0 = nil
     | esk50_0 != nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[113]c_0_99,c_0_111]),c_0_76]),[114]c_0_112])).

 cnf(c_0_118,negated_conjecture,
     ( esk50_0 != nil
     | ~ neq(nil,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[115]c_0_113,c_0_107]),c_0_29])])).

 cnf(c_0_119,negated_conjecture,
     ( esk50_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[116]c_0_114,c_0_115]),c_0_71])])).

 cnf(c_0_120,negated_conjecture,
     ( neq(esk51_0,nil)
     | esk50_0 != nil ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[117]c_0_116,c_0_25]),[118]c_0_65])).

 cnf(c_0_121,negated_conjecture,
     ( esk52_0 = nil
     | esk51_0 = nil ),
     inference(spm,[status(thm)],[[119]c_0_117,[120]c_0_98])).

 cnf(c_0_122,negated_conjecture,
     ( ~ neq(nil,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[121]c_0_118,c_0_119])])).

 cnf(c_0_123,negated_conjecture,
     ( neq(esk51_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[122]c_0_120,c_0_119])])).

 cnf(c_0_124,negated_conjecture,
     ( esk51_0 = nil ),
     inference(sr,[status(thm)],[inference(spm,[status(thm)],[[123]c_0_92,c_0_121]),[124]c_0_122])).

 cnf(c_0_125,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[[125]c_0_123,c_0_124]),[126]c_0_122]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC053+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.22  % Computer   : n158.star.cs.uiowa.edu
 % 0.03/0.22  % Model      : x86_64 x86_64
 % 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.22  % Memory     : 32218.75MB
 % 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.22  % CPULimit   : 300
 % 0.03/0.22  % DateTime   : Tue May 10 20:34:53 CDT 2016
 % 0.03/0.22  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.06/0.47  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.06/0.47  # and selection function SelectCQPrecWNTNp.
 % 0.06/0.47  #
 % 0.06/0.47  # Preprocessing time       : 0.014 s
 % 0.06/0.47  # Presaturation interreduction done
 % 0.06/0.47
 % 0.06/0.47  # Proof found!
 % 0.06/0.47  # SZS status Theorem
 % 0.06/0.47  # SZS output start CNFRefutation
 % 0.06/0.47  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|(![X5]:(ssList(X5)=>((~(neq(X5,nil))|~(frontsegP(X4,X5)))|~(frontsegP(X3,X5))))&(nil!=X4|nil!=X3)))|((nil!=X2|nil=X1)&(~(neq(X2,nil))|?[X6]:(((ssList(X6)&neq(X6,nil))&segmentP(X2,X6))&segmentP(X1,X6)))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
 % 0.06/0.47  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)<=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax5)).
 % 0.06/0.47  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>hd(app(X1,X2))=hd(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax85)).
 % 0.06/0.47  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax78)).
 % 0.06/0.47  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax25)).
 % 0.06/0.47  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.06/0.47  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>nil!=cons(X2,X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax21)).
 % 0.06/0.47  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax22)).
 % 0.06/0.47  fof(c_0_8, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(frontsegP(cons(X1,X3),cons(X2,X4))<=>(X1=X2&frontsegP(X3,X4)))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax44)).
 % 0.06/0.47  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax42)).
 % 0.06/0.47  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax24)).
 % 0.06/0.47  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)<=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax7)).
 % 0.06/0.47  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 0.06/0.47  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 0.06/0.47  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(frontsegP(X1,X2)=>frontsegP(app(X1,X3),X2)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax43)).
 % 0.06/0.47  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax45)).
 % 0.06/0.47  fof(c_0_16, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.06/0.47  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>((frontsegP(X1,X2)&frontsegP(X2,X1))=>X1=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax41)).
 % 0.06/0.47  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 0.06/0.47  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>segmentP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax57)).
 % 0.06/0.47  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 0.06/0.47  fof(c_0_21, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|(![X5]:(ssList(X5)=>((~neq(X5,nil)|~frontsegP(X4,X5))|~frontsegP(X3,X5)))&(nil!=X4|nil!=X3)))|((nil!=X2|nil=X1)&(~neq(X2,nil)|?[X6]:(((ssList(X6)&neq(X6,nil))&segmentP(X2,X6))&segmentP(X1,X6))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.06/0.47  fof(c_0_22, negated_conjecture, (![X12]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((esk49_0=esk51_0&esk48_0=esk50_0)&(((nil=esk51_0|ssList(esk52_0))&(nil=esk50_0|ssList(esk52_0)))&((((nil=esk51_0|neq(esk52_0,nil))&(nil=esk50_0|neq(esk52_0,nil)))&((nil=esk51_0|frontsegP(esk51_0,esk52_0))&(nil=esk50_0|frontsegP(esk51_0,esk52_0))))&((nil=esk51_0|frontsegP(esk50_0,esk52_0))&(nil=esk50_0|frontsegP(esk50_0,esk52_0))))))&(((neq(esk49_0,nil)|nil=esk49_0)&((((~ssList(X12)|~neq(X12,nil))|~segmentP(esk49_0,X12))|~segmentP(esk48_0,X12))|nil=esk49_0))&((neq(esk49_0,nil)|nil!=esk48_0)&((((~ssList(X12)|~neq(X12,nil))|~segmentP(esk49_0,X12))|~segmentP(esk48_0,X12))|nil!=esk48_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).
 % 0.06/0.47  fof(c_0_23, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
 % 0.06/0.47  cnf(c_0_24, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_25, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  fof(c_0_26, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|hd(app(X3,X4))=hd(X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
 % 0.06/0.47  cnf(c_0_27, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
 % 0.06/0.47  cnf(c_0_28, negated_conjecture, (frontsegP(esk51_0,esk52_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_29, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_24, c_0_25])).
 % 0.06/0.47  cnf(c_0_30, negated_conjecture, (ssList(esk52_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_31, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
 % 0.06/0.47  fof(c_0_32, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
 % 0.06/0.47  fof(c_0_33, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 0.06/0.47  fof(c_0_34, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
 % 0.06/0.47  fof(c_0_35, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|nil!=cons(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
 % 0.06/0.47  cnf(c_0_36, plain, (hd(app(X1,X2))=hd(X1)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
 % 0.06/0.47  cnf(c_0_37, negated_conjecture, (app(esk52_0,esk6_2(esk51_0,esk52_0))=esk51_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])]), c_0_30])).
 % 0.06/0.47  cnf(c_0_38, negated_conjecture, (esk50_0=nil|ssList(esk6_2(esk51_0,esk52_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_28]), c_0_29])]), c_0_30])).
 % 0.06/0.47  fof(c_0_39, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
 % 0.06/0.47  fof(c_0_40, plain, (![X5]:![X6]:![X7]:![X8]:(((((((X5=X6|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5))&(((((frontsegP(X7,X8)|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))&((((((X5!=X6|~frontsegP(X7,X8))|frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
 % 0.06/0.47  cnf(c_0_41, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
 % 0.06/0.47  cnf(c_0_42, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
 % 0.06/0.47  cnf(c_0_43, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
 % 0.06/0.47  cnf(c_0_44, plain, (nil!=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
 % 0.06/0.47  cnf(c_0_45, negated_conjecture, (hd(esk52_0)=hd(esk51_0)|esk50_0=nil|esk52_0=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_30]), c_0_38])).
 % 0.06/0.47  cnf(c_0_46, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 0.06/0.47  cnf(c_0_47, plain, (frontsegP(cons(X1,X3),cons(X2,X4))|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~ssList(X4)|~frontsegP(X3,X4)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_40])).
 % 0.06/0.47  fof(c_0_48, plain, (![X2]:(~ssList(X2)|frontsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).
 % 0.06/0.47  cnf(c_0_49, plain, (cons(hd(cons(X1,X2)),X2)=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44])).
 % 0.06/0.47  cnf(c_0_50, negated_conjecture, (cons(hd(esk51_0),tl(esk52_0))=esk52_0|esk50_0=nil|esk52_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_45]), c_0_30])).
 % 0.06/0.47  cnf(c_0_51, negated_conjecture, (esk50_0=nil|esk52_0=nil|ssItem(hd(esk51_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_45]), c_0_30])).
 % 0.06/0.47  fof(c_0_52, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
 % 0.06/0.47  cnf(c_0_53, plain, (frontsegP(cons(X1,X2),cons(X1,X3))|~frontsegP(X2,X3)|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_47])).
 % 0.06/0.47  cnf(c_0_54, plain, (frontsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
 % 0.06/0.47  cnf(c_0_55, negated_conjecture, (cons(hd(esk52_0),tl(esk52_0))=esk52_0|esk52_0=nil|esk50_0=nil|~ssList(tl(esk52_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])).
 % 0.06/0.47  cnf(c_0_56, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
 % 0.06/0.47  cnf(c_0_57, plain, (frontsegP(cons(X1,X2),cons(X1,X2))|~ssList(X2)|~ssItem(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
 % 0.06/0.47  cnf(c_0_58, negated_conjecture, (cons(hd(esk52_0),tl(esk52_0))=esk52_0|esk50_0=nil|esk52_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_30])).
 % 0.06/0.47  fof(c_0_59, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
 % 0.06/0.47  cnf(c_0_60, negated_conjecture, (esk52_0=nil|esk50_0=nil|frontsegP(esk52_0,esk52_0)|~ssList(tl(esk52_0))|~ssItem(hd(esk52_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
 % 0.06/0.47  cnf(c_0_61, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_59])).
 % 0.06/0.47  fof(c_0_62, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).
 % 0.06/0.47  fof(c_0_63, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
 % 0.06/0.47  cnf(c_0_64, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_65, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  fof(c_0_66, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~frontsegP(X4,X5)|frontsegP(app(X4,X6),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
 % 0.06/0.47  fof(c_0_67, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])).
 % 0.06/0.47  cnf(c_0_68, negated_conjecture, (esk50_0=nil|esk52_0=nil|frontsegP(esk52_0,esk52_0)|~ssItem(hd(esk52_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_56]), c_0_30])).
 % 0.06/0.47  cnf(c_0_69, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_61])).
 % 0.06/0.47  cnf(c_0_70, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
 % 0.06/0.47  cnf(c_0_71, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_16])).
 % 0.06/0.47  cnf(c_0_72, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
 % 0.06/0.47  cnf(c_0_73, negated_conjecture, (frontsegP(esk50_0,esk52_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_74, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_64, c_0_65])).
 % 0.06/0.47  cnf(c_0_75, negated_conjecture, (frontsegP(esk51_0,esk52_0)|nil=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_76, negated_conjecture, (ssList(esk52_0)|nil=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_77, plain, (frontsegP(app(X1,X2),X3)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
 % 0.06/0.47  cnf(c_0_78, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
 % 0.06/0.47  cnf(c_0_79, negated_conjecture, (esk52_0=nil|esk50_0=nil|frontsegP(esk52_0,esk52_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_46]), c_0_30])).
 % 0.06/0.47  cnf(c_0_80, negated_conjecture, (nil=esk49_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_81, plain, (segmentP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])]), c_0_72])).
 % 0.06/0.47  cnf(c_0_82, negated_conjecture, (app(esk52_0,esk6_2(esk50_0,esk52_0))=esk50_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_73]), c_0_74])]), c_0_30])).
 % 0.06/0.47  cnf(c_0_83, negated_conjecture, (esk50_0=nil|ssList(esk6_2(esk50_0,esk52_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_73]), c_0_74])]), c_0_30])).
 % 0.06/0.47  cnf(c_0_84, negated_conjecture, (app(esk52_0,esk6_2(esk51_0,esk52_0))=esk51_0|esk51_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_75]), c_0_29])]), c_0_76])).
 % 0.06/0.47  cnf(c_0_85, negated_conjecture, (esk51_0=nil|ssList(esk6_2(esk51_0,esk52_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_75]), c_0_29])]), c_0_76])).
 % 0.06/0.47  fof(c_0_86, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|((~frontsegP(X3,X4)|~frontsegP(X4,X3))|X3=X4)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
 % 0.06/0.47  cnf(c_0_87, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_71])])).
 % 0.06/0.47  cnf(c_0_88, negated_conjecture, (app(esk52_0,esk6_2(esk52_0,esk52_0))=esk52_0|esk50_0=nil|esk52_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_79]), c_0_30])).
 % 0.06/0.47  cnf(c_0_89, negated_conjecture, (esk50_0=nil|esk52_0=nil|ssList(esk6_2(esk52_0,esk52_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_79]), c_0_30])).
 % 0.06/0.47  cnf(c_0_90, negated_conjecture, (esk51_0=nil|~segmentP(esk50_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_25]), c_0_65]), c_0_25])).
 % 0.06/0.47  cnf(c_0_91, negated_conjecture, (esk50_0=nil|segmentP(esk50_0,esk52_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_30]), c_0_83])).
 % 0.06/0.47  cnf(c_0_92, negated_conjecture, (neq(esk52_0,nil)|nil=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_93, negated_conjecture, (esk51_0=nil|segmentP(esk51_0,esk52_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_84]), c_0_76]), c_0_85])).
 % 0.06/0.47  fof(c_0_94, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])).
 % 0.06/0.47  cnf(c_0_95, plain, (X1=X2|~frontsegP(X2,X1)|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
 % 0.06/0.47  cnf(c_0_96, negated_conjecture, (esk52_0=nil|esk50_0=nil|frontsegP(esk52_0,nil)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_30]), c_0_89])).
 % 0.06/0.47  cnf(c_0_97, negated_conjecture, (frontsegP(esk50_0,esk52_0)|nil=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_98, negated_conjecture, (esk50_0=nil|esk51_0=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_92]), c_0_76]), c_0_93])).
 % 0.06/0.47  cnf(c_0_99, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_94])).
 % 0.06/0.47  cnf(c_0_100, negated_conjecture, (nil!=esk48_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  fof(c_0_101, plain, (![X2]:(~ssList(X2)|segmentP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).
 % 0.06/0.47  fof(c_0_102, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])])).
 % 0.06/0.47  cnf(c_0_103, negated_conjecture, (esk50_0=nil|esk52_0=nil|~frontsegP(nil,esk52_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_71])]), c_0_30])).
 % 0.06/0.47  cnf(c_0_104, negated_conjecture, (esk51_0=nil|frontsegP(nil,esk52_0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
 % 0.06/0.47  cnf(c_0_105, negated_conjecture, (esk50_0=nil|esk52_0=nil|esk51_0!=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_37]), c_0_30]), c_0_38])).
 % 0.06/0.47  cnf(c_0_106, negated_conjecture, (esk50_0!=nil|~segmentP(esk50_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_65]), c_0_65]), c_0_25])).
 % 0.06/0.47  cnf(c_0_107, plain, (segmentP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
 % 0.06/0.47  cnf(c_0_108, plain, (~ssList(X1)|~ssList(X2)|X1!=X2|~neq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
 % 0.06/0.47  cnf(c_0_109, negated_conjecture, (neq(esk52_0,nil)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_110, negated_conjecture, (esk52_0=nil|esk50_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_105])).
 % 0.06/0.47  cnf(c_0_111, negated_conjecture, (app(esk52_0,esk6_2(esk50_0,esk52_0))=esk50_0|esk51_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_97]), c_0_74])]), c_0_76])).
 % 0.06/0.47  cnf(c_0_112, negated_conjecture, (esk51_0=nil|ssList(esk6_2(esk50_0,esk52_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_97]), c_0_74])]), c_0_76])).
 % 0.06/0.47  cnf(c_0_113, negated_conjecture, (esk50_0!=nil|~segmentP(esk51_0,nil)|~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_71]), c_0_74])])).
 % 0.06/0.47  cnf(c_0_114, plain, (~ssList(X1)|~neq(X1,X1)), inference(er,[status(thm)],[c_0_108])).
 % 0.06/0.47  cnf(c_0_115, negated_conjecture, (esk50_0=nil|neq(nil,nil)), inference(spm,[status(thm)],[c_0_109, c_0_110])).
 % 0.06/0.47  cnf(c_0_116, negated_conjecture, (neq(esk49_0,nil)|nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.47  cnf(c_0_117, negated_conjecture, (esk51_0=nil|esk52_0=nil|esk50_0!=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_111]), c_0_76]), c_0_112])).
 % 0.06/0.47  cnf(c_0_118, negated_conjecture, (esk50_0!=nil|~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_107]), c_0_29])])).
 % 0.06/0.47  cnf(c_0_119, negated_conjecture, (esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_71])])).
 % 0.06/0.47  cnf(c_0_120, negated_conjecture, (neq(esk51_0,nil)|esk50_0!=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_25]), c_0_65])).
 % 0.06/0.47  cnf(c_0_121, negated_conjecture, (esk52_0=nil|esk51_0=nil), inference(spm,[status(thm)],[c_0_117, c_0_98])).
 % 0.06/0.47  cnf(c_0_122, negated_conjecture, (~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_118, c_0_119])])).
 % 0.06/0.47  cnf(c_0_123, negated_conjecture, (neq(esk51_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_119])])).
 % 0.06/0.47  cnf(c_0_124, negated_conjecture, (esk51_0=nil), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_121]), c_0_122])).
 % 0.06/0.47  cnf(c_0_125, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_123, c_0_124]), c_0_122]), ['proof']).
 % 0.06/0.47  # SZS output end CNFRefutation
 % 0.06/0.47  # Training examples: 0 positive, 0 negative
 % 0.06/0.47
 % 0.06/0.47  # -------------------------------------------------
 % 0.06/0.47  # User time                : 0.234 s
 % 0.06/0.47  # System time              : 0.010 s
 % 0.06/0.47  # Total time               : 0.245 s
 % 0.06/0.47  # Maximum resident set size: 1876 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
