   --------------------------------------------------------------------------

                   TSTP Solution File: SWC170+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC170+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n111.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:55 EDT 2016

 % Result     : Theorem 0.02s
 % Output     : CNFRefutation 0.02s
 % Verified   :
 % Statistics : Number of formulae       :   82 ( 437 expanded)
 %              Number of clauses        :   51 ( 163 expanded)
 %              Number of leaves         :   16 ( 110 expanded)
 %              Depth                    :   15
 %              Number of atoms          :  335 (3637 expanded)
 %              Number of equality atoms :   59 ( 718 expanded)
 %              Maximal formula depth    :   24 (   5 average)
 %              Maximal clause size      :   34 (   2 average)
 %              Maximal term depth       :    5 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ~ segmentP(X4,X3)
                     | ! [X5] :
                         ( ssItem(X5)
                        => ! [X6] :
                             ( ssItem(X6)
                            => ! [X7] :
                                 ( ssList(X7)
                                => ! [X8] :
                                     ( ssList(X8)
                                    => ( app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8) != X1
                                       | neq(X5,X6) ) ) ) ) )
                     | ( ~ singletonP(X3)
                       & neq(X4,nil) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax1)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_4,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => duplicatefreeP(cons(X1,nil)) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax71)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( singletonP(X1)
       <=> ? [X2] :
             ( ssItem(X2)
             & cons(X2,nil) = X1 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax4)).

 fof(c_0_8,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_9,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( duplicatefreeP(X1)
       <=> ! [X2] :
             ( ssItem(X2)
            => ! [X3] :
                 ( ssItem(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ! [X5] :
                         ( ssList(X5)
                        => ! [X6] :
                             ( ssList(X6)
                            => ( app(app(X4,cons(X2,X5)),cons(X3,X6)) = X1
                              => X2 != X3 ) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax13)).

 fof(c_0_10,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

 fof(c_0_11,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( cyclefreeP(X1)
       <=> ! [X2] :
             ( ssItem(X2)
            => ! [X3] :
                 ( ssItem(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ! [X5] :
                         ( ssList(X5)
                        => ! [X6] :
                             ( ssList(X6)
                            => ( app(app(X4,cons(X2,X5)),cons(X3,X6)) = X1
                              => ~ ( leq(X2,X3)
                                   & leq(X3,X2) ) ) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax8)).

 fof(c_0_12,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( ( segmentP(X1,X2)
               & segmentP(X2,X1) )
            => X1 = X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax54)).

 fof(c_0_13,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => leq(X1,X1) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax31)).

 fof(c_0_14,axiom,(
     ! [X1] :
       ( ssList(X1)
      => segmentP(X1,nil) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax57)).

 fof(c_0_15,axiom,(
     cyclefreeP(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax60)).

 fof(c_0_16,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ~ segmentP(X4,X3)
                       | ! [X5] :
                           ( ssItem(X5)
                          => ! [X6] :
                               ( ssItem(X6)
                              => ! [X7] :
                                   ( ssList(X7)
                                  => ! [X8] :
                                       ( ssList(X8)
                                      => ( app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8) != X1
                                         | neq(X5,X6) ) ) ) ) )
                       | ( ~ singletonP(X3)
                         & neq(X4,nil) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_17,negated_conjecture,
     ( ssList(esk48_0)
     & ssList(esk49_0)
     & ssList(esk50_0)
     & ssList(esk51_0)
     & esk49_0 = esk51_0
     & esk48_0 = esk50_0
     & segmentP(esk51_0,esk50_0)
     & ssItem(esk52_0)
     & ssItem(esk53_0)
     & ssList(esk54_0)
     & ssList(esk55_0)
     & app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0) = esk48_0
     & ~ neq(esk52_0,esk53_0)
     & ( singletonP(esk50_0)
       | ~ neq(esk51_0,nil) ) ),
     inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).

 fof(c_0_18,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssItem(X4)
         | ~ ssItem(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssItem(X4)
         | ~ ssItem(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).

 cnf(c_0_19,negated_conjecture,
     ( app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0) = esk48_0 ),
     inference(split_conjunct,[status(thm)],[[4]c_0_17])).

 cnf(c_0_20,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_17])).

 cnf(c_0_21,negated_conjecture,
     ( ~ neq(esk52_0,esk53_0) ),
     inference(split_conjunct,[status(thm)],[[6]c_0_17])).

 cnf(c_0_22,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssItem(X1)
     | ~ ssItem(X2) ),
     inference(split_conjunct,[status(thm)],[[7]c_0_18])).

 cnf(c_0_23,negated_conjecture,
     ( ssItem(esk53_0) ),
     inference(split_conjunct,[status(thm)],[[8]c_0_17])).

 cnf(c_0_24,negated_conjecture,
     ( ssItem(esk52_0) ),
     inference(split_conjunct,[status(thm)],[[9]c_0_17])).

 cnf(c_0_25,negated_conjecture,
     ( app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0) = esk50_0 ),
     inference(rw,[status(thm)],[[10]c_0_19,[11]c_0_20])).

 cnf(c_0_26,negated_conjecture,
     ( esk53_0 = esk52_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[12]c_0_21,c_0_22]),c_0_23]),c_0_24])])).

 fof(c_0_27,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

 cnf(c_0_28,negated_conjecture,
     ( app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk52_0,nil)),esk55_0) = esk50_0 ),
     inference(rw,[status(thm)],[[13]c_0_25,[14]c_0_26])).

 cnf(c_0_29,plain,
     ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_27])).

 cnf(c_0_30,negated_conjecture,
     ( ssList(esk55_0) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_17])).

 fof(c_0_31,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

 cnf(c_0_32,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0)) = esk50_0
     | ~ ssList(app(esk54_0,cons(esk52_0,nil)))
     | ~ ssList(cons(esk52_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_28,c_0_29]),c_0_30])])).

 cnf(c_0_33,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[18]c_0_31])).

 cnf(c_0_34,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_4])).

 fof(c_0_35,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

 cnf(c_0_36,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0)) = esk50_0
     | ~ ssList(app(esk54_0,cons(esk52_0,nil))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[20]c_0_32,c_0_33]),c_0_34]),c_0_24])])).

 cnf(c_0_37,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[21]c_0_35])).

 cnf(c_0_38,negated_conjecture,
     ( ssList(esk54_0) ),
     inference(split_conjunct,[status(thm)],[[22]c_0_17])).

 fof(c_0_39,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | duplicatefreeP(cons(X2,nil)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).

 fof(c_0_40,plain,(
     ! [X3,X5] :
       ( ( ssItem(esk5_1(X3))
         | ~ singletonP(X3)
         | ~ ssList(X3) )
       & ( cons(esk5_1(X3),nil) = X3
         | ~ singletonP(X3)
         | ~ ssList(X3) )
       & ( ~ ssItem(X5)
         | cons(X5,nil) != X3
         | singletonP(X3)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).

 fof(c_0_41,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).

 cnf(c_0_42,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[23]c_0_17])).

 cnf(c_0_43,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[24]c_0_17])).

 fof(c_0_44,plain,(
     ! [X7,X8,X9,X10,X11,X12] :
       ( ( ~ duplicatefreeP(X7)
         | ~ ssItem(X8)
         | ~ ssItem(X9)
         | ~ ssList(X10)
         | ~ ssList(X11)
         | ~ ssList(X12)
         | app(app(X10,cons(X8,X11)),cons(X9,X12)) != X7
         | X8 != X9
         | ~ ssList(X7) )
       & ( ssItem(esk35_1(X7))
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( ssItem(esk36_1(X7))
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk37_1(X7))
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk38_1(X7))
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk39_1(X7))
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( app(app(esk37_1(X7),cons(esk35_1(X7),esk38_1(X7))),cons(esk36_1(X7),esk39_1(X7))) = X7
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( esk35_1(X7) = esk36_1(X7)
         | duplicatefreeP(X7)
         | ~ ssList(X7) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).

 cnf(c_0_45,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0)) = esk50_0
     | ~ ssList(cons(esk52_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_36,c_0_37]),c_0_38])])).

 fof(c_0_46,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).

 cnf(c_0_47,plain,
     ( duplicatefreeP(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[26]c_0_39])).

 cnf(c_0_48,plain,
     ( cons(esk5_1(X1),nil) = X1
     | ~ ssList(X1)
     | ~ singletonP(X1) ),
     inference(split_conjunct,[status(thm)],[[27]c_0_40])).

 cnf(c_0_49,plain,
     ( ssItem(esk5_1(X1))
     | ~ ssList(X1)
     | ~ singletonP(X1) ),
     inference(split_conjunct,[status(thm)],[[28]c_0_40])).

 cnf(c_0_50,negated_conjecture,
     ( singletonP(esk50_0)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[29]c_0_17])).

 cnf(c_0_51,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[30]c_0_41])).

 cnf(c_0_52,negated_conjecture,
     ( ssList(esk51_0) ),
     inference(rw,[status(thm)],[[31]c_0_42,[32]c_0_43])).

 cnf(c_0_53,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[33]c_0_17])).

 cnf(c_0_54,plain,
     ( ~ ssList(X1)
     | X2 != X3
     | app(app(X4,cons(X2,X5)),cons(X3,X6)) != X1
     | ~ ssList(X6)
     | ~ ssList(X5)
     | ~ ssList(X4)
     | ~ ssItem(X3)
     | ~ ssItem(X2)
     | ~ duplicatefreeP(X1) ),
     inference(split_conjunct,[status(thm)],[[34]c_0_44])).

 cnf(c_0_55,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0)) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[35]c_0_45,c_0_33]),c_0_34]),c_0_24])])).

 cnf(c_0_56,plain,
     ( cons(X1,X2) = app(cons(X1,nil),X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[36]c_0_46])).

 fof(c_0_57,plain,(
     ! [X7,X8,X9,X10,X11,X12] :
       ( ( ~ cyclefreeP(X7)
         | ~ ssItem(X8)
         | ~ ssItem(X9)
         | ~ ssList(X10)
         | ~ ssList(X11)
         | ~ ssList(X12)
         | app(app(X10,cons(X8,X11)),cons(X9,X12)) != X7
         | ~ leq(X8,X9)
         | ~ leq(X9,X8)
         | ~ ssList(X7) )
       & ( ssItem(esk10_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( ssItem(esk11_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk12_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk13_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk14_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( app(app(esk12_1(X7),cons(esk10_1(X7),esk13_1(X7))),cons(esk11_1(X7),esk14_1(X7))) = X7
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( leq(esk10_1(X7),esk11_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( leq(esk11_1(X7),esk10_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).

 fof(c_0_58,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ~ segmentP(X3,X4)
       | ~ segmentP(X4,X3)
       | X3 = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).

 cnf(c_0_59,plain,
     ( duplicatefreeP(X1)
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[37]c_0_47,c_0_48]),[38]c_0_49])).

 cnf(c_0_60,negated_conjecture,
     ( esk51_0 = nil
     | singletonP(esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[39]c_0_50,c_0_51]),c_0_34]),c_0_52])])).

 cnf(c_0_61,negated_conjecture,
     ( ssList(esk50_0) ),
     inference(rw,[status(thm)],[[40]c_0_53,[41]c_0_20])).

 cnf(c_0_62,plain,
     ( ~ duplicatefreeP(app(app(X1,cons(X2,X3)),cons(X2,X4)))
     | ~ ssList(app(app(X1,cons(X2,X3)),cons(X2,X4)))
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_54])])).

 cnf(c_0_63,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,nil)),cons(esk52_0,esk55_0)) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[42]c_0_55,c_0_56]),c_0_30]),c_0_24])])).

 cnf(c_0_64,plain,
     ( ~ ssList(X1)
     | ~ leq(X2,X3)
     | ~ leq(X3,X2)
     | app(app(X4,cons(X3,X5)),cons(X2,X6)) != X1
     | ~ ssList(X6)
     | ~ ssList(X5)
     | ~ ssList(X4)
     | ~ ssItem(X2)
     | ~ ssItem(X3)
     | ~ cyclefreeP(X1) ),
     inference(split_conjunct,[status(thm)],[[43]c_0_57])).

 cnf(c_0_65,plain,
     ( X1 = X2
     | ~ segmentP(X2,X1)
     | ~ segmentP(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[44]c_0_58])).

 cnf(c_0_66,negated_conjecture,
     ( segmentP(esk51_0,esk50_0) ),
     inference(split_conjunct,[status(thm)],[[45]c_0_17])).

 cnf(c_0_67,negated_conjecture,
     ( esk51_0 = nil
     | duplicatefreeP(esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[46]c_0_59,c_0_60]),c_0_61])])).

 cnf(c_0_68,negated_conjecture,
     ( ~ duplicatefreeP(esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[47]c_0_62,c_0_63]),c_0_61]),c_0_30]),c_0_34]),c_0_38]),c_0_24])])).

 cnf(c_0_69,plain,
     ( ~ leq(X1,X2)
     | ~ leq(X2,X1)
     | ~ cyclefreeP(app(app(X3,cons(X1,X4)),cons(X2,X5)))
     | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
     | ~ ssList(X5)
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssItem(X1)
     | ~ ssItem(X2) ),
     inference(er,[status(thm)],[[48]c_0_64])).

 fof(c_0_70,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | leq(X2,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).

 cnf(c_0_71,negated_conjecture,
     ( esk50_0 = esk51_0
     | ~ segmentP(esk50_0,esk51_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[49]c_0_65,c_0_66]),c_0_52]),c_0_61])])).

 cnf(c_0_72,negated_conjecture,
     ( esk51_0 = nil ),
     inference(sr,[status(thm)],[[50]c_0_67,[51]c_0_68])).

 fof(c_0_73,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | segmentP(X2,nil) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).

 cnf(c_0_74,negated_conjecture,
     ( ~ leq(esk52_0,esk52_0)
     | ~ cyclefreeP(esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[52]c_0_69,c_0_63]),c_0_61]),c_0_30]),c_0_34]),c_0_38]),c_0_24])])).

 cnf(c_0_75,plain,
     ( leq(X1,X1)
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[53]c_0_70])).

 cnf(c_0_76,negated_conjecture,
     ( esk50_0 = nil
     | ~ segmentP(esk50_0,nil) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[54]c_0_71,c_0_72]),[55]c_0_72])).

 cnf(c_0_77,plain,
     ( segmentP(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[56]c_0_73])).

 cnf(c_0_78,negated_conjecture,
     ( ~ cyclefreeP(esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[57]c_0_74,c_0_75]),c_0_24])])).

 cnf(c_0_79,negated_conjecture,
     ( esk50_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[58]c_0_76,c_0_77]),c_0_61])])).

 cnf(c_0_80,plain,
     ( cyclefreeP(nil) ),
     inference(split_conjunct,[status(thm)],[[59]c_0_15])).

 cnf(c_0_81,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[60]c_0_78,c_0_79]),c_0_80])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC170+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.22  % Computer   : n111.star.cs.uiowa.edu
 % 0.02/0.22  % Model      : x86_64 x86_64
 % 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.22  % Memory     : 32218.75MB
 % 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.22  % CPULimit   : 300
 % 0.02/0.22  % DateTime   : Tue May 10 20:36:38 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.02/0.29  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.02/0.29  # and selection function SelectCQPrecWNTNp.
 % 0.02/0.29  #
 % 0.02/0.29  # Preprocessing time       : 0.014 s
 % 0.02/0.29  # Presaturation interreduction done
 % 0.02/0.29
 % 0.02/0.29  # Proof found!
 % 0.02/0.29  # SZS status Theorem
 % 0.02/0.29  # SZS output start CNFRefutation
 % 0.02/0.29  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|~(segmentP(X4,X3)))|![X5]:(ssItem(X5)=>![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>(app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8)!=X1|neq(X5,X6)))))))|(~(singletonP(X3))&neq(X4,nil)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.02/0.29  fof(c_0_1, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax1)).
 % 0.02/0.29  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
 % 0.02/0.29  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.02/0.29  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.02/0.29  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 0.02/0.29  fof(c_0_6, axiom, (![X1]:(ssItem(X1)=>duplicatefreeP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax71)).
 % 0.02/0.29  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)<=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax4)).
 % 0.02/0.29  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 0.02/0.29  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>(duplicatefreeP(X1)<=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(app(X4,cons(X2,X5)),cons(X3,X6))=X1=>X2!=X3))))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax13)).
 % 0.02/0.29  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
 % 0.02/0.29  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>(cyclefreeP(X1)<=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(app(X4,cons(X2,X5)),cons(X3,X6))=X1=>~((leq(X2,X3)&leq(X3,X2)))))))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax8)).
 % 0.02/0.29  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>((segmentP(X1,X2)&segmentP(X2,X1))=>X1=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax54)).
 % 0.02/0.29  fof(c_0_13, axiom, (![X1]:(ssItem(X1)=>leq(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax31)).
 % 0.02/0.29  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>segmentP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax57)).
 % 0.02/0.29  fof(c_0_15, axiom, (cyclefreeP(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax60)).
 % 0.02/0.29  fof(c_0_16, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|~segmentP(X4,X3))|![X5]:(ssItem(X5)=>![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>(app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8)!=X1|neq(X5,X6)))))))|(~singletonP(X3)&neq(X4,nil))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.02/0.29  fof(c_0_17, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&segmentP(esk51_0,esk50_0))&(ssItem(esk52_0)&(ssItem(esk53_0)&(ssList(esk54_0)&(ssList(esk55_0)&(app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0)=esk48_0&~neq(esk52_0,esk53_0)))))))&(singletonP(esk50_0)|~neq(esk51_0,nil)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
 % 0.02/0.29  fof(c_0_18, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssItem(X4))|~ssItem(X3))&(((X3=X4|neq(X3,X4))|~ssItem(X4))|~ssItem(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
 % 0.02/0.29  cnf(c_0_19, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.29  cnf(c_0_20, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.29  cnf(c_0_21, negated_conjecture, (~neq(esk52_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.29  cnf(c_0_22, plain, (neq(X1,X2)|X1=X2|~ssItem(X1)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
 % 0.02/0.29  cnf(c_0_23, negated_conjecture, (ssItem(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.29  cnf(c_0_24, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.29  cnf(c_0_25, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0)=esk50_0), inference(rw,[status(thm)],[c_0_19, c_0_20])).
 % 0.02/0.29  cnf(c_0_26, negated_conjecture, (esk53_0=esk52_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24])])).
 % 0.02/0.29  fof(c_0_27, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
 % 0.02/0.29  cnf(c_0_28, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk52_0,nil)),esk55_0)=esk50_0), inference(rw,[status(thm)],[c_0_25, c_0_26])).
 % 0.02/0.29  cnf(c_0_29, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
 % 0.02/0.29  cnf(c_0_30, negated_conjecture, (ssList(esk55_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.29  fof(c_0_31, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
 % 0.02/0.29  cnf(c_0_32, negated_conjecture, (app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0))=esk50_0|~ssList(app(esk54_0,cons(esk52_0,nil)))|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
 % 0.02/0.29  cnf(c_0_33, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
 % 0.02/0.29  cnf(c_0_34, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
 % 0.02/0.29  fof(c_0_35, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
 % 0.02/0.29  cnf(c_0_36, negated_conjecture, (app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0))=esk50_0|~ssList(app(esk54_0,cons(esk52_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_24])])).
 % 0.02/0.30  cnf(c_0_37, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
 % 0.02/0.30  cnf(c_0_38, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.30  fof(c_0_39, plain, (![X2]:(~ssItem(X2)|duplicatefreeP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
 % 0.02/0.30  fof(c_0_40, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
 % 0.02/0.30  fof(c_0_41, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
 % 0.02/0.30  cnf(c_0_42, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.30  cnf(c_0_43, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.30  fof(c_0_44, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((~duplicatefreeP(X7)|(~ssItem(X8)|(~ssItem(X9)|(~ssList(X10)|(~ssList(X11)|(~ssList(X12)|(app(app(X10,cons(X8,X11)),cons(X9,X12))!=X7|X8!=X9)))))))|~ssList(X7))&(((ssItem(esk35_1(X7))|duplicatefreeP(X7))|~ssList(X7))&(((ssItem(esk36_1(X7))|duplicatefreeP(X7))|~ssList(X7))&(((ssList(esk37_1(X7))|duplicatefreeP(X7))|~ssList(X7))&(((ssList(esk38_1(X7))|duplicatefreeP(X7))|~ssList(X7))&(((ssList(esk39_1(X7))|duplicatefreeP(X7))|~ssList(X7))&(((app(app(esk37_1(X7),cons(esk35_1(X7),esk38_1(X7))),cons(esk36_1(X7),esk39_1(X7)))=X7|duplicatefreeP(X7))|~ssList(X7))&((esk35_1(X7)=esk36_1(X7)|duplicatefreeP(X7))|~ssList(X7)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).
 % 0.02/0.30  cnf(c_0_45, negated_conjecture, (app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0))=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])])).
 % 0.02/0.30  fof(c_0_46, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
 % 0.02/0.30  cnf(c_0_47, plain, (duplicatefreeP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 0.02/0.30  cnf(c_0_48, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
 % 0.02/0.30  cnf(c_0_49, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
 % 0.02/0.30  cnf(c_0_50, negated_conjecture, (singletonP(esk50_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.30  cnf(c_0_51, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 0.02/0.30  cnf(c_0_52, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_42, c_0_43])).
 % 0.02/0.30  cnf(c_0_53, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.30  cnf(c_0_54, plain, (~ssList(X1)|X2!=X3|app(app(X4,cons(X2,X5)),cons(X3,X6))!=X1|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssItem(X3)|~ssItem(X2)|~duplicatefreeP(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
 % 0.02/0.30  cnf(c_0_55, negated_conjecture, (app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_33]), c_0_34]), c_0_24])])).
 % 0.02/0.30  cnf(c_0_56, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
 % 0.02/0.30  fof(c_0_57, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((~cyclefreeP(X7)|(~ssItem(X8)|(~ssItem(X9)|(~ssList(X10)|(~ssList(X11)|(~ssList(X12)|(app(app(X10,cons(X8,X11)),cons(X9,X12))!=X7|(~leq(X8,X9)|~leq(X9,X8)))))))))|~ssList(X7))&(((ssItem(esk10_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssItem(esk11_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssList(esk12_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssList(esk13_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssList(esk14_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((app(app(esk12_1(X7),cons(esk10_1(X7),esk13_1(X7))),cons(esk11_1(X7),esk14_1(X7)))=X7|cyclefreeP(X7))|~ssList(X7))&(((leq(esk10_1(X7),esk11_1(X7))|cyclefreeP(X7))|~ssList(X7))&((leq(esk11_1(X7),esk10_1(X7))|cyclefreeP(X7))|~ssList(X7))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
 % 0.02/0.30  fof(c_0_58, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|((~segmentP(X3,X4)|~segmentP(X4,X3))|X3=X4)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
 % 0.02/0.30  cnf(c_0_59, plain, (duplicatefreeP(X1)|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
 % 0.02/0.30  cnf(c_0_60, negated_conjecture, (esk51_0=nil|singletonP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_34]), c_0_52])])).
 % 0.02/0.30  cnf(c_0_61, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_53, c_0_20])).
 % 0.02/0.30  cnf(c_0_62, plain, (~duplicatefreeP(app(app(X1,cons(X2,X3)),cons(X2,X4)))|~ssList(app(app(X1,cons(X2,X3)),cons(X2,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_54])])).
 % 0.02/0.30  cnf(c_0_63, negated_conjecture, (app(app(esk54_0,cons(esk52_0,nil)),cons(esk52_0,esk55_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_30]), c_0_24])])).
 % 0.02/0.30  cnf(c_0_64, plain, (~ssList(X1)|~leq(X2,X3)|~leq(X3,X2)|app(app(X4,cons(X3,X5)),cons(X2,X6))!=X1|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssItem(X2)|~ssItem(X3)|~cyclefreeP(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
 % 0.02/0.30  cnf(c_0_65, plain, (X1=X2|~segmentP(X2,X1)|~segmentP(X1,X2)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
 % 0.02/0.30  cnf(c_0_66, negated_conjecture, (segmentP(esk51_0,esk50_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.02/0.30  cnf(c_0_67, negated_conjecture, (esk51_0=nil|duplicatefreeP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])])).
 % 0.02/0.30  cnf(c_0_68, negated_conjecture, (~duplicatefreeP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_61]), c_0_30]), c_0_34]), c_0_38]), c_0_24])])).
 % 0.02/0.30  cnf(c_0_69, plain, (~leq(X1,X2)|~leq(X2,X1)|~cyclefreeP(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssItem(X1)|~ssItem(X2)), inference(er,[status(thm)],[c_0_64])).
 % 0.02/0.30  fof(c_0_70, plain, (![X2]:(~ssItem(X2)|leq(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).
 % 0.02/0.30  cnf(c_0_71, negated_conjecture, (esk50_0=esk51_0|~segmentP(esk50_0,esk51_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_52]), c_0_61])])).
 % 0.02/0.30  cnf(c_0_72, negated_conjecture, (esk51_0=nil), inference(sr,[status(thm)],[c_0_67, c_0_68])).
 % 0.02/0.30  fof(c_0_73, plain, (![X2]:(~ssList(X2)|segmentP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).
 % 0.02/0.30  cnf(c_0_74, negated_conjecture, (~leq(esk52_0,esk52_0)|~cyclefreeP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_63]), c_0_61]), c_0_30]), c_0_34]), c_0_38]), c_0_24])])).
 % 0.02/0.30  cnf(c_0_75, plain, (leq(X1,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
 % 0.02/0.30  cnf(c_0_76, negated_conjecture, (esk50_0=nil|~segmentP(esk50_0,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_72]), c_0_72])).
 % 0.02/0.30  cnf(c_0_77, plain, (segmentP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
 % 0.02/0.30  cnf(c_0_78, negated_conjecture, (~cyclefreeP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_24])])).
 % 0.02/0.30  cnf(c_0_79, negated_conjecture, (esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_61])])).
 % 0.02/0.30  cnf(c_0_80, plain, (cyclefreeP(nil)), inference(split_conjunct,[status(thm)],[c_0_15])).
 % 0.02/0.30  cnf(c_0_81, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_79]), c_0_80])]), ['proof']).
 % 0.02/0.30  # SZS output end CNFRefutation
 % 0.02/0.30  # Training examples: 0 positive, 0 negative
 % 0.06/0.30
 % 0.06/0.30  # -------------------------------------------------
 % 0.06/0.30  # User time                : 0.066 s
 % 0.06/0.30  # System time              : 0.004 s
 % 0.06/0.30  # Total time               : 0.070 s
 % 0.06/0.30  # Maximum resident set size: 1876 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
