   --------------------------------------------------------------------------

                   TSTP Solution File: SWC152+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC152+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n142.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:51 EDT 2016

 % Result     : Theorem 0.02s
 % Output     : CNFRefutation 0.02s
 % Verified   :
 % Statistics : Number of formulae       :   77 ( 766 expanded)
 %              Number of clauses        :   53 ( 272 expanded)
 %              Number of leaves         :   12 ( 193 expanded)
 %              Depth                    :   18
 %              Number of atoms          :  371 (12678 expanded)
 %              Number of equality atoms :   98 (3716 expanded)
 %              Maximal formula depth    :   30 (   5 average)
 %              Maximal clause size      :   42 (   3 average)
 %              Maximal term depth       :    6 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ~ ssList(X4)
                   | X2 != X4
                   | X1 != X3
                   | ! [X5] :
                       ( ssItem(X5)
                      => ! [X6] :
                           ( ssItem(X6)
                          => ! [X7] :
                               ( ssList(X7)
                              => ! [X8] :
                                   ( ssList(X8)
                                  => ! [X9] :
                                       ( ~ ssList(X9)
                                       | app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9) != X1
                                       | ~ leq(X6,X5)
                                       | ( ! [X10] :
                                             ( ~ ssItem(X10)
                                             | ~ memberP(X8,X10)
                                             | ( leq(X5,X10)
                                               & leq(X10,X6) ) )
                                         & leq(X5,X6) ) ) ) ) ) )
                   | ( ! [X11] :
                         ( ~ ssItem(X11)
                         | cons(X11,nil) != X3
                         | ~ memberP(X4,X11)
                         | ? [X12] :
                             ( ssItem(X12)
                             & X11 != X12
                             & memberP(X4,X12)
                             & leq(X11,X12) ) )
                     & ( nil != X4
                       | nil != X3 ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_3,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( totalorderedP(X1)
       <=> ! [X2] :
             ( ssItem(X2)
            => ! [X3] :
                 ( ssItem(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ! [X5] :
                         ( ssList(X5)
                        => ! [X6] :
                             ( ssList(X6)
                            => ( app(app(X4,cons(X2,X5)),cons(X3,X6)) = X1
                              => leq(X2,X3) ) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax11)).

 fof(c_0_8,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( cyclefreeP(X1)
       <=> ! [X2] :
             ( ssItem(X2)
            => ! [X3] :
                 ( ssItem(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ! [X5] :
                         ( ssList(X5)
                        => ! [X6] :
                             ( ssList(X6)
                            => ( app(app(X4,cons(X2,X5)),cons(X3,X6)) = X1
                              => ~ ( leq(X2,X3)
                                   & leq(X3,X2) ) ) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax8)).

 fof(c_0_9,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => totalorderedP(cons(X1,nil)) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax65)).

 fof(c_0_10,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => cyclefreeP(cons(X1,nil)) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax59)).

 fof(c_0_11,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) != X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax18)).

 fof(c_0_12,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ~ ssList(X4)
                     | X2 != X4
                     | X1 != X3
                     | ! [X5] :
                         ( ssItem(X5)
                        => ! [X6] :
                             ( ssItem(X6)
                            => ! [X7] :
                                 ( ssList(X7)
                                => ! [X8] :
                                     ( ssList(X8)
                                    => ! [X9] :
                                         ( ~ ssList(X9)
                                         | app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9) != X1
                                         | ~ leq(X6,X5)
                                         | ( ! [X10] :
                                               ( ~ ssItem(X10)
                                               | ~ memberP(X8,X10)
                                               | ( leq(X5,X10)
                                                 & leq(X10,X6) ) )
                                           & leq(X5,X6) ) ) ) ) ) )
                     | ( ! [X11] :
                           ( ~ ssItem(X11)
                           | cons(X11,nil) != X3
                           | ~ memberP(X4,X11)
                           | ? [X12] :
                               ( ssItem(X12)
                               & X11 != X12
                               & memberP(X4,X12)
                               & leq(X11,X12) ) )
                       & ( nil != X4
                         | nil != X3 ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_13,negated_conjecture,(
     ! [X24] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ssItem(esk52_0)
       & ssItem(esk53_0)
       & ssList(esk54_0)
       & ssList(esk55_0)
       & ssList(esk56_0)
       & app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk48_0
       & leq(esk53_0,esk52_0)
       & ( ssItem(esk57_0)
         | ~ leq(esk52_0,esk53_0) )
       & ( memberP(esk55_0,esk57_0)
         | ~ leq(esk52_0,esk53_0) )
       & ( ~ leq(esk52_0,esk57_0)
         | ~ leq(esk57_0,esk53_0)
         | ~ leq(esk52_0,esk53_0) )
       & ( nil = esk51_0
         | ssItem(esk58_0) )
       & ( nil = esk50_0
         | ssItem(esk58_0) )
       & ( nil = esk51_0
         | cons(esk58_0,nil) = esk50_0 )
       & ( nil = esk50_0
         | cons(esk58_0,nil) = esk50_0 )
       & ( nil = esk51_0
         | memberP(esk51_0,esk58_0) )
       & ( nil = esk50_0
         | memberP(esk51_0,esk58_0) )
       & ( nil = esk51_0
         | ~ ssItem(X24)
         | esk58_0 = X24
         | ~ memberP(esk51_0,X24)
         | ~ leq(esk58_0,X24) )
       & ( nil = esk50_0
         | ~ ssItem(X24)
         | esk58_0 = X24
         | ~ memberP(esk51_0,X24)
         | ~ leq(esk58_0,X24) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])])).

 cnf(c_0_14,negated_conjecture,
     ( app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk48_0 ),
     inference(split_conjunct,[status(thm)],[[4]c_0_13])).

 cnf(c_0_15,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_13])).

 fof(c_0_16,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

 cnf(c_0_17,negated_conjecture,
     ( app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk50_0 ),
     inference(rw,[status(thm)],[[6]c_0_14,[7]c_0_15])).

 cnf(c_0_18,plain,
     ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[8]c_0_16])).

 cnf(c_0_19,negated_conjecture,
     ( ssList(esk55_0) ),
     inference(split_conjunct,[status(thm)],[[9]c_0_13])).

 cnf(c_0_20,negated_conjecture,
     ( ssList(esk54_0) ),
     inference(split_conjunct,[status(thm)],[[10]c_0_13])).

 fof(c_0_21,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

 cnf(c_0_22,negated_conjecture,
     ( app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0) = esk50_0
     | ~ ssList(cons(esk52_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[11]c_0_17,c_0_18]),c_0_19]),c_0_20])])).

 cnf(c_0_23,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_21])).

 cnf(c_0_24,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[13]c_0_3])).

 cnf(c_0_25,negated_conjecture,
     ( ssItem(esk52_0) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_13])).

 fof(c_0_26,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

 cnf(c_0_27,negated_conjecture,
     ( app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_22,c_0_23]),c_0_24]),c_0_25])])).

 cnf(c_0_28,plain,
     ( cons(X1,X2) = app(cons(X1,nil),X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_26])).

 cnf(c_0_29,negated_conjecture,
     ( app(app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,nil)),esk56_0) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_27,c_0_28]),c_0_19]),c_0_25])])).

 cnf(c_0_30,negated_conjecture,
     ( ssList(esk56_0) ),
     inference(split_conjunct,[status(thm)],[[18]c_0_13])).

 cnf(c_0_31,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0
     | ~ ssList(app(esk54_0,cons(esk52_0,esk55_0)))
     | ~ ssList(cons(esk53_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[19]c_0_18,c_0_29]),c_0_30])])).

 cnf(c_0_32,negated_conjecture,
     ( ssItem(esk53_0) ),
     inference(split_conjunct,[status(thm)],[[20]c_0_13])).

 fof(c_0_33,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

 fof(c_0_34,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).

 cnf(c_0_35,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0
     | ~ ssList(app(esk54_0,cons(esk52_0,esk55_0))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[21]c_0_31,c_0_23]),c_0_24]),c_0_32])])).

 cnf(c_0_36,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[22]c_0_33])).

 cnf(c_0_37,plain,
     ( nil = X2
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[23]c_0_34])).

 cnf(c_0_38,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0
     | ~ ssList(cons(esk52_0,esk55_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[24]c_0_35,c_0_36]),c_0_20])])).

 cnf(c_0_39,negated_conjecture,
     ( esk56_0 = nil
     | esk50_0 != nil
     | ~ ssList(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_37,c_0_17]),c_0_30])])).

 fof(c_0_40,plain,(
     ! [X7,X8,X9,X10,X11,X12] :
       ( ( ~ totalorderedP(X7)
         | ~ ssItem(X8)
         | ~ ssItem(X9)
         | ~ ssList(X10)
         | ~ ssList(X11)
         | ~ ssList(X12)
         | app(app(X10,cons(X8,X11)),cons(X9,X12)) != X7
         | leq(X8,X9)
         | ~ ssList(X7) )
       & ( ssItem(esk25_1(X7))
         | totalorderedP(X7)
         | ~ ssList(X7) )
       & ( ssItem(esk26_1(X7))
         | totalorderedP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk27_1(X7))
         | totalorderedP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk28_1(X7))
         | totalorderedP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk29_1(X7))
         | totalorderedP(X7)
         | ~ ssList(X7) )
       & ( app(app(esk27_1(X7),cons(esk25_1(X7),esk28_1(X7))),cons(esk26_1(X7),esk29_1(X7))) = X7
         | totalorderedP(X7)
         | ~ ssList(X7) )
       & ( ~ leq(esk25_1(X7),esk26_1(X7))
         | totalorderedP(X7)
         | ~ ssList(X7) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).

 cnf(c_0_41,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_38,c_0_23]),c_0_19]),c_0_25])])).

 cnf(c_0_42,negated_conjecture,
     ( esk56_0 = nil
     | esk50_0 != nil
     | ~ ssList(app(app(esk54_0,cons(esk52_0,nil)),esk55_0))
     | ~ ssList(cons(esk53_0,nil)) ),
     inference(spm,[status(thm)],[[27]c_0_39,[28]c_0_36])).

 fof(c_0_43,plain,(
     ! [X7,X8,X9,X10,X11,X12] :
       ( ( ~ cyclefreeP(X7)
         | ~ ssItem(X8)
         | ~ ssItem(X9)
         | ~ ssList(X10)
         | ~ ssList(X11)
         | ~ ssList(X12)
         | app(app(X10,cons(X8,X11)),cons(X9,X12)) != X7
         | ~ leq(X8,X9)
         | ~ leq(X9,X8)
         | ~ ssList(X7) )
       & ( ssItem(esk10_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( ssItem(esk11_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk12_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk13_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk14_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( app(app(esk12_1(X7),cons(esk10_1(X7),esk13_1(X7))),cons(esk11_1(X7),esk14_1(X7))) = X7
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( leq(esk10_1(X7),esk11_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) )
       & ( leq(esk11_1(X7),esk10_1(X7))
         | cyclefreeP(X7)
         | ~ ssList(X7) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).

 cnf(c_0_44,plain,
     ( leq(X2,X3)
     | ~ ssList(X1)
     | app(app(X4,cons(X2,X5)),cons(X3,X6)) != X1
     | ~ ssList(X6)
     | ~ ssList(X5)
     | ~ ssList(X4)
     | ~ ssItem(X3)
     | ~ ssItem(X2)
     | ~ totalorderedP(X1) ),
     inference(split_conjunct,[status(thm)],[[29]c_0_40])).

 cnf(c_0_45,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[30]c_0_13])).

 fof(c_0_46,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | totalorderedP(cons(X2,nil)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).

 cnf(c_0_47,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,esk56_0)) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[31]c_0_41,c_0_28]),c_0_30]),c_0_32])])).

 cnf(c_0_48,negated_conjecture,
     ( esk56_0 = nil
     | esk50_0 != nil
     | ~ ssList(app(app(esk54_0,cons(esk52_0,nil)),esk55_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[32]c_0_42,c_0_23]),c_0_24]),c_0_32])])).

 cnf(c_0_49,plain,
     ( ~ ssList(X1)
     | ~ leq(X2,X3)
     | ~ leq(X3,X2)
     | app(app(X4,cons(X3,X5)),cons(X2,X6)) != X1
     | ~ ssList(X6)
     | ~ ssList(X5)
     | ~ ssList(X4)
     | ~ ssItem(X2)
     | ~ ssItem(X3)
     | ~ cyclefreeP(X1) ),
     inference(split_conjunct,[status(thm)],[[33]c_0_43])).

 cnf(c_0_50,plain,
     ( leq(X1,X2)
     | ~ totalorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))
     | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
     | ~ ssList(X5)
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssItem(X2)
     | ~ ssItem(X1) ),
     inference(er,[status(thm)],[[34]c_0_44])).

 cnf(c_0_51,negated_conjecture,
     ( ssList(esk50_0) ),
     inference(rw,[status(thm)],[[35]c_0_45,[36]c_0_15])).

 cnf(c_0_52,plain,
     ( totalorderedP(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[37]c_0_46])).

 cnf(c_0_53,negated_conjecture,
     ( cons(esk58_0,nil) = esk50_0
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[38]c_0_13])).

 cnf(c_0_54,negated_conjecture,
     ( ssItem(esk58_0)
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[39]c_0_13])).

 fof(c_0_55,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | cyclefreeP(cons(X2,nil)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).

 cnf(c_0_56,negated_conjecture,
     ( cons(esk53_0,esk56_0) = nil
     | esk50_0 != nil
     | ~ ssList(app(esk54_0,cons(esk52_0,esk55_0)))
     | ~ ssList(cons(esk53_0,esk56_0)) ),
     inference(spm,[status(thm)],[[40]c_0_37,[41]c_0_47])).

 cnf(c_0_57,negated_conjecture,
     ( esk56_0 = nil
     | esk50_0 != nil
     | ~ ssList(app(esk54_0,cons(esk52_0,nil))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[42]c_0_48,c_0_36]),c_0_19])])).

 cnf(c_0_58,plain,
     ( ~ leq(X1,X2)
     | ~ leq(X2,X1)
     | ~ cyclefreeP(app(app(X3,cons(X1,X4)),cons(X2,X5)))
     | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
     | ~ ssList(X5)
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssItem(X1)
     | ~ ssItem(X2) ),
     inference(er,[status(thm)],[[43]c_0_49])).

 cnf(c_0_59,negated_conjecture,
     ( leq(esk53_0,esk52_0) ),
     inference(split_conjunct,[status(thm)],[[44]c_0_13])).

 cnf(c_0_60,negated_conjecture,
     ( leq(esk52_0,esk53_0)
     | ~ totalorderedP(esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[45]c_0_50,c_0_47]),c_0_51]),c_0_30]),c_0_19]),c_0_20]),c_0_32]),c_0_25])])).

 cnf(c_0_61,negated_conjecture,
     ( esk50_0 = nil
     | totalorderedP(esk50_0) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[46]c_0_52,c_0_53]),[47]c_0_54])).

 cnf(c_0_62,plain,
     ( cyclefreeP(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[48]c_0_55])).

 cnf(c_0_63,negated_conjecture,
     ( cons(esk53_0,esk56_0) = nil
     | esk50_0 != nil
     | ~ ssList(app(esk54_0,cons(esk52_0,esk55_0))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[49]c_0_56,c_0_23]),c_0_30]),c_0_32])])).

 cnf(c_0_64,negated_conjecture,
     ( esk56_0 = nil
     | esk50_0 != nil
     | ~ ssList(cons(esk52_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[50]c_0_57,c_0_36]),c_0_20])])).

 cnf(c_0_65,negated_conjecture,
     ( ~ leq(esk52_0,esk53_0)
     | ~ cyclefreeP(esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[51]c_0_58,c_0_47]),c_0_59]),c_0_51]),c_0_30]),c_0_19]),c_0_20]),c_0_25]),c_0_32])])).

 cnf(c_0_66,negated_conjecture,
     ( esk50_0 = nil
     | leq(esk52_0,esk53_0) ),
     inference(spm,[status(thm)],[[52]c_0_60,[53]c_0_61])).

 cnf(c_0_67,negated_conjecture,
     ( esk50_0 = nil
     | cyclefreeP(esk50_0) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[54]c_0_62,c_0_53]),[55]c_0_54])).

 cnf(c_0_68,negated_conjecture,
     ( cons(esk53_0,esk56_0) = nil
     | esk50_0 != nil
     | ~ ssList(cons(esk52_0,esk55_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[56]c_0_63,c_0_36]),c_0_20])])).

 cnf(c_0_69,negated_conjecture,
     ( esk56_0 = nil
     | esk50_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[57]c_0_64,c_0_23]),c_0_24]),c_0_25])])).

 cnf(c_0_70,negated_conjecture,
     ( esk50_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[58]c_0_65,c_0_66]),[59]c_0_67])).

 fof(c_0_71,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) != X3 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).

 cnf(c_0_72,negated_conjecture,
     ( cons(esk53_0,esk56_0) = nil
     | esk50_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[60]c_0_68,c_0_23]),c_0_19]),c_0_25])])).

 cnf(c_0_73,negated_conjecture,
     ( esk56_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[61]c_0_69,c_0_70])])).

 cnf(c_0_74,plain,
     ( cons(X1,X2) != X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[62]c_0_71])).

 cnf(c_0_75,negated_conjecture,
     ( cons(esk53_0,nil) = nil ),
     inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[[63]c_0_72,c_0_70])]),[64]c_0_73])).

 cnf(c_0_76,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[65]c_0_74,c_0_75]),c_0_24]),c_0_32])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC152+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.22  % Computer   : n142.star.cs.uiowa.edu
 % 0.02/0.22  % Model      : x86_64 x86_64
 % 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.22  % Memory     : 32218.75MB
 % 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.23  % CPULimit   : 300
 % 0.02/0.23  % DateTime   : Tue May 10 20:36:08 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.02/0.29  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.02/0.29  # and selection function SelectCQPrecWNTNp.
 % 0.02/0.29  #
 % 0.02/0.29  # Preprocessing time       : 0.014 s
 % 0.02/0.29  # Presaturation interreduction done
 % 0.02/0.29
 % 0.02/0.29  # Proof found!
 % 0.02/0.29  # SZS status Theorem
 % 0.02/0.29  # SZS output start CNFRefutation
 % 0.02/0.29  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~(ssList(X4))|X2!=X4)|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>![X9]:(((~(ssList(X9))|app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9)!=X1)|~(leq(X6,X5)))|(![X10]:((~(ssItem(X10))|~(memberP(X8,X10)))|(leq(X5,X10)&leq(X10,X6)))&leq(X5,X6))))))))|(![X11]:(((~(ssItem(X11))|cons(X11,nil)!=X3)|~(memberP(X4,X11)))|?[X12]:(((ssItem(X12)&X11!=X12)&memberP(X4,X12))&leq(X11,X12)))&(nil!=X4|nil!=X3))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.02/0.29  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
 % 0.02/0.29  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.02/0.29  fof(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.02/0.29  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
 % 0.02/0.29  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 0.02/0.29  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 0.02/0.29  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(totalorderedP(X1)<=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(app(X4,cons(X2,X5)),cons(X3,X6))=X1=>leq(X2,X3)))))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax11)).
 % 0.02/0.29  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>(cyclefreeP(X1)<=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(app(X4,cons(X2,X5)),cons(X3,X6))=X1=>~((leq(X2,X3)&leq(X3,X2)))))))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax8)).
 % 0.02/0.29  fof(c_0_9, axiom, (![X1]:(ssItem(X1)=>totalorderedP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax65)).
 % 0.02/0.29  fof(c_0_10, axiom, (![X1]:(ssItem(X1)=>cyclefreeP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax59)).
 % 0.02/0.29  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax18)).
 % 0.02/0.29  fof(c_0_12, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~ssList(X4)|X2!=X4)|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>![X9]:(((~ssList(X9)|app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9)!=X1)|~leq(X6,X5))|(![X10]:((~ssItem(X10)|~memberP(X8,X10))|(leq(X5,X10)&leq(X10,X6)))&leq(X5,X6))))))))|(![X11]:(((~ssItem(X11)|cons(X11,nil)!=X3)|~memberP(X4,X11))|?[X12]:(((ssItem(X12)&X11!=X12)&memberP(X4,X12))&leq(X11,X12)))&(nil!=X4|nil!=X3)))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.02/0.29  fof(c_0_13, negated_conjecture, (![X24]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(ssItem(esk52_0)&(ssItem(esk53_0)&(ssList(esk54_0)&(ssList(esk55_0)&(((ssList(esk56_0)&app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk48_0)&leq(esk53_0,esk52_0))&(((ssItem(esk57_0)|~leq(esk52_0,esk53_0))&(memberP(esk55_0,esk57_0)|~leq(esk52_0,esk53_0)))&((~leq(esk52_0,esk57_0)|~leq(esk57_0,esk53_0))|~leq(esk52_0,esk53_0)))))))))&(((((nil=esk51_0|ssItem(esk58_0))&(nil=esk50_0|ssItem(esk58_0)))&((nil=esk51_0|cons(esk58_0,nil)=esk50_0)&(nil=esk50_0|cons(esk58_0,nil)=esk50_0)))&((nil=esk51_0|memberP(esk51_0,esk58_0))&(nil=esk50_0|memberP(esk51_0,esk58_0))))&((nil=esk51_0|(((~ssItem(X24)|esk58_0=X24)|~memberP(esk51_0,X24))|~leq(esk58_0,X24)))&(nil=esk50_0|(((~ssItem(X24)|esk58_0=X24)|~memberP(esk51_0,X24))|~leq(esk58_0,X24)))))))))),
 % inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])])).
 % 0.02/0.29  cnf(c_0_14, negated_conjecture, (app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.02/0.29  cnf(c_0_15, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.02/0.29  fof(c_0_16, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
 % 0.02/0.29  cnf(c_0_17, negated_conjecture, (app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk50_0), inference(rw,[status(thm)],[c_0_14, c_0_15])).
 % 0.02/0.29  cnf(c_0_18, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
 % 0.02/0.29  cnf(c_0_19, negated_conjecture, (ssList(esk55_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.02/0.29  cnf(c_0_20, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.02/0.29  fof(c_0_21, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
 % 0.02/0.29  cnf(c_0_22, negated_conjecture, (app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0)=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])])).
 % 0.02/0.29  cnf(c_0_23, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
 % 0.02/0.29  cnf(c_0_24, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
 % 0.02/0.29  cnf(c_0_25, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.02/0.29  fof(c_0_26, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 0.02/0.29  cnf(c_0_27, negated_conjecture, (app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25])])).
 % 0.02/0.29  cnf(c_0_28, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
 % 0.02/0.29  cnf(c_0_29, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,nil)),esk56_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_19]), c_0_25])])).
 % 0.02/0.29  cnf(c_0_30, negated_conjecture, (ssList(esk56_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.02/0.29  cnf(c_0_31, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0|~ssList(app(esk54_0,cons(esk52_0,esk55_0)))|~ssList(cons(esk53_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_29]), c_0_30])])).
 % 0.02/0.29  cnf(c_0_32, negated_conjecture, (ssItem(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.02/0.29  fof(c_0_33, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
 % 0.02/0.29  fof(c_0_34, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
 % 0.02/0.29  cnf(c_0_35, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0|~ssList(app(esk54_0,cons(esk52_0,esk55_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_23]), c_0_24]), c_0_32])])).
 % 0.02/0.29  cnf(c_0_36, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
 % 0.02/0.29  cnf(c_0_37, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
 % 0.02/0.29  cnf(c_0_38, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0|~ssList(cons(esk52_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_20])])).
 % 0.02/0.29  cnf(c_0_39, negated_conjecture, (esk56_0=nil|esk50_0!=nil|~ssList(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_17]), c_0_30])])).
 % 0.02/0.29  fof(c_0_40, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((~totalorderedP(X7)|(~ssItem(X8)|(~ssItem(X9)|(~ssList(X10)|(~ssList(X11)|(~ssList(X12)|(app(app(X10,cons(X8,X11)),cons(X9,X12))!=X7|leq(X8,X9))))))))|~ssList(X7))&(((ssItem(esk25_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssItem(esk26_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk27_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk28_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk29_1(X7))|totalorderedP(X7))|~ssList(X7))&(((app(app(esk27_1(X7),cons(esk25_1(X7),esk28_1(X7))),cons(esk26_1(X7),esk29_1(X7)))=X7|totalorderedP(X7))|~ssList(X7))&((~leq(esk25_1(X7),esk26_1(X7))|totalorderedP(X7))|~ssList(X7)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
 % 0.02/0.29  cnf(c_0_41, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_23]), c_0_19]), c_0_25])])).
 % 0.02/0.29  cnf(c_0_42, negated_conjecture, (esk56_0=nil|esk50_0!=nil|~ssList(app(app(esk54_0,cons(esk52_0,nil)),esk55_0))|~ssList(cons(esk53_0,nil))), inference(spm,[status(thm)],[c_0_39, c_0_36])).
 % 0.02/0.29  fof(c_0_43, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((~cyclefreeP(X7)|(~ssItem(X8)|(~ssItem(X9)|(~ssList(X10)|(~ssList(X11)|(~ssList(X12)|(app(app(X10,cons(X8,X11)),cons(X9,X12))!=X7|(~leq(X8,X9)|~leq(X9,X8)))))))))|~ssList(X7))&(((ssItem(esk10_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssItem(esk11_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssList(esk12_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssList(esk13_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssList(esk14_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((app(app(esk12_1(X7),cons(esk10_1(X7),esk13_1(X7))),cons(esk11_1(X7),esk14_1(X7)))=X7|cyclefreeP(X7))|~ssList(X7))&(((leq(esk10_1(X7),esk11_1(X7))|cyclefreeP(X7))|~ssList(X7))&((leq(esk11_1(X7),esk10_1(X7))|cyclefreeP(X7))|~ssList(X7))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
 % 0.02/0.29  cnf(c_0_44, plain, (leq(X2,X3)|~ssList(X1)|app(app(X4,cons(X2,X5)),cons(X3,X6))!=X1|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssItem(X3)|~ssItem(X2)|~totalorderedP(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
 % 0.02/0.29  cnf(c_0_45, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.02/0.29  fof(c_0_46, plain, (![X2]:(~ssItem(X2)|totalorderedP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).
 % 0.02/0.29  cnf(c_0_47, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,esk56_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_28]), c_0_30]), c_0_32])])).
 % 0.02/0.29  cnf(c_0_48, negated_conjecture, (esk56_0=nil|esk50_0!=nil|~ssList(app(app(esk54_0,cons(esk52_0,nil)),esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_23]), c_0_24]), c_0_32])])).
 % 0.02/0.29  cnf(c_0_49, plain, (~ssList(X1)|~leq(X2,X3)|~leq(X3,X2)|app(app(X4,cons(X3,X5)),cons(X2,X6))!=X1|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssItem(X2)|~ssItem(X3)|~cyclefreeP(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
 % 0.02/0.29  cnf(c_0_50, plain, (leq(X1,X2)|~totalorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_44])).
 % 0.02/0.29  cnf(c_0_51, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_45, c_0_15])).
 % 0.02/0.29  cnf(c_0_52, plain, (totalorderedP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
 % 0.02/0.29  cnf(c_0_53, negated_conjecture, (cons(esk58_0,nil)=esk50_0|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.02/0.29  cnf(c_0_54, negated_conjecture, (ssItem(esk58_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.02/0.29  fof(c_0_55, plain, (![X2]:(~ssItem(X2)|cyclefreeP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
 % 0.02/0.29  cnf(c_0_56, negated_conjecture, (cons(esk53_0,esk56_0)=nil|esk50_0!=nil|~ssList(app(esk54_0,cons(esk52_0,esk55_0)))|~ssList(cons(esk53_0,esk56_0))), inference(spm,[status(thm)],[c_0_37, c_0_47])).
 % 0.02/0.29  cnf(c_0_57, negated_conjecture, (esk56_0=nil|esk50_0!=nil|~ssList(app(esk54_0,cons(esk52_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_36]), c_0_19])])).
 % 0.02/0.29  cnf(c_0_58, plain, (~leq(X1,X2)|~leq(X2,X1)|~cyclefreeP(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssItem(X1)|~ssItem(X2)), inference(er,[status(thm)],[c_0_49])).
 % 0.02/0.29  cnf(c_0_59, negated_conjecture, (leq(esk53_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.02/0.29  cnf(c_0_60, negated_conjecture, (leq(esk52_0,esk53_0)|~totalorderedP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_47]), c_0_51]), c_0_30]), c_0_19]), c_0_20]), c_0_32]), c_0_25])])).
 % 0.02/0.29  cnf(c_0_61, negated_conjecture, (esk50_0=nil|totalorderedP(esk50_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
 % 0.02/0.29  cnf(c_0_62, plain, (cyclefreeP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
 % 0.02/0.29  cnf(c_0_63, negated_conjecture, (cons(esk53_0,esk56_0)=nil|esk50_0!=nil|~ssList(app(esk54_0,cons(esk52_0,esk55_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_23]), c_0_30]), c_0_32])])).
 % 0.02/0.29  cnf(c_0_64, negated_conjecture, (esk56_0=nil|esk50_0!=nil|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_36]), c_0_20])])).
 % 0.02/0.29  cnf(c_0_65, negated_conjecture, (~leq(esk52_0,esk53_0)|~cyclefreeP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_47]), c_0_59]), c_0_51]), c_0_30]), c_0_19]), c_0_20]), c_0_25]), c_0_32])])).
 % 0.02/0.29  cnf(c_0_66, negated_conjecture, (esk50_0=nil|leq(esk52_0,esk53_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
 % 0.02/0.29  cnf(c_0_67, negated_conjecture, (esk50_0=nil|cyclefreeP(esk50_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_53]), c_0_54])).
 % 0.02/0.29  cnf(c_0_68, negated_conjecture, (cons(esk53_0,esk56_0)=nil|esk50_0!=nil|~ssList(cons(esk52_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_36]), c_0_20])])).
 % 0.02/0.29  cnf(c_0_69, negated_conjecture, (esk56_0=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_23]), c_0_24]), c_0_25])])).
 % 0.02/0.29  cnf(c_0_70, negated_conjecture, (esk50_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
 % 0.02/0.29  fof(c_0_71, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
 % 0.02/0.29  cnf(c_0_72, negated_conjecture, (cons(esk53_0,esk56_0)=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_23]), c_0_19]), c_0_25])])).
 % 0.02/0.29  cnf(c_0_73, negated_conjecture, (esk56_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_70])])).
 % 0.02/0.29  cnf(c_0_74, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
 % 0.02/0.29  cnf(c_0_75, negated_conjecture, (cons(esk53_0,nil)=nil), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_70])]), c_0_73])).
 % 0.02/0.29  cnf(c_0_76, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_24]), c_0_32])]), ['proof']).
 % 0.02/0.29  # SZS output end CNFRefutation
 % 0.02/0.29  # Training examples: 0 positive, 0 negative
 % 0.02/0.29
 % 0.02/0.29  # -------------------------------------------------
 % 0.02/0.29  # User time                : 0.062 s
 % 0.02/0.29  # System time              : 0.004 s
 % 0.02/0.29  # Total time               : 0.066 s
 % 0.02/0.29  # Maximum resident set size: 1884 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
