   --------------------------------------------------------------------------

                   TSTP Solution File: SWC112+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC112+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n130.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:42 EDT 2016

 % Result     : Theorem 0.03s
 % Output     : CNFRefutation 0.03s
 % Verified   :
 % Statistics : Number of formulae       :   59 (1079 expanded)
 %              Number of clauses        :   41 ( 373 expanded)
 %              Number of leaves         :    9 ( 262 expanded)
 %              Depth                    :   12
 %              Number of atoms          :  271 (20747 expanded)
 %              Number of equality atoms :  102 (7388 expanded)
 %              Maximal formula depth    :   30 (   5 average)
 %              Maximal clause size      :   36 (   2 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ! [X5] :
                         ( ssList(X5)
                        => ! [X6] :
                             ( ssList(X6)
                            => ( app(app(X5,X3),X6) != X4
                               | ~ strictorderedP(X3)
                               | ? [X7] :
                                   ( ssItem(X7)
                                   & ? [X8] :
                                       ( ssList(X8)
                                       & app(X8,cons(X7,nil)) = X5
                                       & ? [X9] :
                                           ( ssItem(X9)
                                           & ? [X10] :
                                               ( ssList(X10)
                                               & app(cons(X9,nil),X10) = X3
                                               & lt(X7,X9) ) ) ) )
                               | ? [X11] :
                                   ( ssItem(X11)
                                   & ? [X12] :
                                       ( ssList(X12)
                                       & app(cons(X11,nil),X12) = X6
                                       & ? [X13] :
                                           ( ssItem(X13)
                                           & ? [X14] :
                                               ( ssList(X14)
                                               & app(X14,cons(X13,nil)) = X3
                                               & lt(X13,X11) ) ) ) ) ) ) )
                     | ( nil != X4
                       & nil = X3 )
                     | ( ( nil != X2
                         | nil = X1 )
                       & ( ~ neq(X2,nil)
                         | ( neq(X1,nil)
                           & segmentP(X2,X1) ) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( segmentP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & ? [X4] :
                     ( ssList(X4)
                     & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax7)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( app(X3,X2) = app(X1,X2)
                => X3 = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax79)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(nil,X1) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

 fof(c_0_8,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_9,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ! [X5] :
                           ( ssList(X5)
                          => ! [X6] :
                               ( ssList(X6)
                              => ( app(app(X5,X3),X6) != X4
                                 | ~ strictorderedP(X3)
                                 | ? [X7] :
                                     ( ssItem(X7)
                                     & ? [X8] :
                                         ( ssList(X8)
                                         & app(X8,cons(X7,nil)) = X5
                                         & ? [X9] :
                                             ( ssItem(X9)
                                             & ? [X10] :
                                                 ( ssList(X10)
                                                 & app(cons(X9,nil),X10) = X3
                                                 & lt(X7,X9) ) ) ) )
                                 | ? [X11] :
                                     ( ssItem(X11)
                                     & ? [X12] :
                                         ( ssList(X12)
                                         & app(cons(X11,nil),X12) = X6
                                         & ? [X13] :
                                             ( ssItem(X13)
                                             & ? [X14] :
                                                 ( ssList(X14)
                                                 & app(X14,cons(X13,nil)) = X3
                                                 & lt(X13,X11) ) ) ) ) ) ) )
                       | ( nil != X4
                         & nil = X3 )
                       | ( ( nil != X2
                           | nil = X1 )
                         & ( ~ neq(X2,nil)
                           | ( neq(X1,nil)
                             & segmentP(X2,X1) ) ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_10,negated_conjecture,(
     ! [X21,X22,X23,X24,X25,X26,X27,X28] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ssList(esk52_0)
       & ssList(esk53_0)
       & app(app(esk52_0,esk50_0),esk53_0) = esk51_0
       & strictorderedP(esk50_0)
       & ( ~ ssItem(X21)
         | ~ ssList(X22)
         | app(X22,cons(X21,nil)) != esk52_0
         | ~ ssItem(X23)
         | ~ ssList(X24)
         | app(cons(X23,nil),X24) != esk50_0
         | ~ lt(X21,X23) )
       & ( ~ ssItem(X25)
         | ~ ssList(X26)
         | app(cons(X25,nil),X26) != esk53_0
         | ~ ssItem(X27)
         | ~ ssList(X28)
         | app(X28,cons(X27,nil)) != esk50_0
         | ~ lt(X27,X25) )
       & ( nil = esk51_0
         | nil != esk50_0 )
       & ( neq(esk49_0,nil)
         | nil = esk49_0 )
       & ( ~ neq(esk48_0,nil)
         | ~ segmentP(esk49_0,esk48_0)
         | nil = esk49_0 )
       & ( neq(esk49_0,nil)
         | nil != esk48_0 )
       & ( ~ neq(esk48_0,nil)
         | ~ segmentP(esk49_0,esk48_0)
         | nil != esk48_0 ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])])).

 fof(c_0_11,plain,(
     ! [X5,X6,X9,X10] :
       ( ( ssList(esk8_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ssList(esk9_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ~ ssList(X9)
         | ~ ssList(X10)
         | app(app(X9,X6),X10) != X5
         | segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

 cnf(c_0_12,negated_conjecture,
     ( app(app(esk52_0,esk50_0),esk53_0) = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[4]c_0_10])).

 cnf(c_0_13,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_10])).

 cnf(c_0_14,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_10])).

 fof(c_0_15,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

 cnf(c_0_16,plain,
     ( segmentP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(app(X3,X2),X4) != X1
     | ~ ssList(X4)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[7]c_0_11])).

 fof(c_0_17,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(X6,X5) != app(X4,X5)
       | X6 = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

 fof(c_0_18,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).

 cnf(c_0_19,negated_conjecture,
     ( app(app(esk52_0,esk48_0),esk53_0) = esk49_0 ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[8]c_0_12,c_0_13]),[9]c_0_14])).

 cnf(c_0_20,plain,
     ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[10]c_0_15])).

 cnf(c_0_21,negated_conjecture,
     ( ssList(esk53_0) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_10])).

 cnf(c_0_22,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_10])).

 cnf(c_0_23,negated_conjecture,
     ( ssList(esk52_0) ),
     inference(split_conjunct,[status(thm)],[[13]c_0_10])).

 cnf(c_0_24,plain,
     ( segmentP(app(app(X1,X2),X3),X2)
     | ~ ssList(app(app(X1,X2),X3))
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(er,[status(thm)],[[14]c_0_16])).

 cnf(c_0_25,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_10])).

 cnf(c_0_26,plain,
     ( X1 = X2
     | app(X1,X3) != app(X2,X3)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_17])).

 fof(c_0_27,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

 cnf(c_0_28,plain,
     ( nil = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[17]c_0_18])).

 cnf(c_0_29,negated_conjecture,
     ( app(esk52_0,app(esk48_0,esk53_0)) = esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_19,c_0_20]),c_0_21]),c_0_22]),c_0_23])])).

 cnf(c_0_30,negated_conjecture,
     ( nil = esk49_0
     | ~ segmentP(esk49_0,esk48_0)
     | ~ neq(esk48_0,nil) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_10])).

 cnf(c_0_31,negated_conjecture,
     ( segmentP(esk49_0,esk48_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[20]c_0_24,c_0_19]),c_0_25]),c_0_21]),c_0_23]),c_0_22])])).

 fof(c_0_32,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).

 cnf(c_0_33,negated_conjecture,
     ( nil = esk51_0
     | nil != esk50_0 ),
     inference(split_conjunct,[status(thm)],[[21]c_0_10])).

 cnf(c_0_34,plain,
     ( nil = X2
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[22]c_0_18])).

 cnf(c_0_35,negated_conjecture,
     ( X1 = app(esk52_0,esk48_0)
     | app(X1,esk53_0) != esk49_0
     | ~ ssList(app(esk52_0,esk48_0))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[23]c_0_26,c_0_19]),c_0_21])])).

 cnf(c_0_36,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[24]c_0_27])).

 fof(c_0_37,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(nil,X2) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

 cnf(c_0_38,negated_conjecture,
     ( esk52_0 = nil
     | esk49_0 != nil
     | ~ ssList(app(esk48_0,esk53_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_28,c_0_29]),c_0_23])])).

 cnf(c_0_39,negated_conjecture,
     ( esk49_0 = nil
     | ~ neq(esk48_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[26]c_0_30,c_0_31])])).

 cnf(c_0_40,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[27]c_0_32])).

 cnf(c_0_41,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[28]c_0_8])).

 cnf(c_0_42,negated_conjecture,
     ( esk49_0 = nil
     | esk48_0 != nil ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[29]c_0_33,c_0_14]),[30]c_0_13])).

 cnf(c_0_43,negated_conjecture,
     ( esk53_0 = nil
     | esk49_0 != nil
     | ~ ssList(app(esk52_0,esk48_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[31]c_0_34,c_0_19]),c_0_21])])).

 cnf(c_0_44,negated_conjecture,
     ( X1 = app(esk52_0,esk48_0)
     | app(X1,esk53_0) != esk49_0
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[32]c_0_35,c_0_36]),c_0_22]),c_0_23])])).

 cnf(c_0_45,plain,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[33]c_0_37])).

 cnf(c_0_46,negated_conjecture,
     ( esk52_0 = nil
     | esk49_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[34]c_0_38,c_0_36]),c_0_21]),c_0_22])])).

 cnf(c_0_47,negated_conjecture,
     ( esk49_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[35]c_0_39,c_0_40]),c_0_41]),c_0_22])]),[36]c_0_42])).

 cnf(c_0_48,negated_conjecture,
     ( esk53_0 = nil
     | esk49_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[37]c_0_43,c_0_36]),c_0_22]),c_0_23])])).

 cnf(c_0_49,negated_conjecture,
     ( app(esk52_0,esk48_0) = nil
     | esk53_0 != esk49_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[38]c_0_44,c_0_45]),c_0_41]),c_0_21])])).

 cnf(c_0_50,negated_conjecture,
     ( esk52_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[39]c_0_46,c_0_47])])).

 cnf(c_0_51,negated_conjecture,
     ( esk53_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[40]c_0_48,c_0_47])])).

 cnf(c_0_52,negated_conjecture,
     ( nil != esk48_0
     | ~ segmentP(esk49_0,esk48_0)
     | ~ neq(esk48_0,nil) ),
     inference(split_conjunct,[status(thm)],[[41]c_0_10])).

 cnf(c_0_53,negated_conjecture,
     ( app(nil,esk48_0) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[42]c_0_49,c_0_47]),c_0_50]),c_0_51])])).

 cnf(c_0_54,negated_conjecture,
     ( esk48_0 != nil
     | ~ neq(esk48_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[43]c_0_52,c_0_31])])).

 cnf(c_0_55,negated_conjecture,
     ( esk48_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[44]c_0_34,c_0_53]),c_0_22]),c_0_41])])).

 cnf(c_0_56,negated_conjecture,
     ( neq(esk49_0,nil)
     | nil != esk48_0 ),
     inference(split_conjunct,[status(thm)],[[45]c_0_10])).

 cnf(c_0_57,negated_conjecture,
     ( ~ neq(nil,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[46]c_0_54,c_0_55]),c_0_55])])).

 cnf(c_0_58,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[47]c_0_56,c_0_47]),c_0_55])]),[48]c_0_57]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC112+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.23  % Computer   : n130.star.cs.uiowa.edu
 % 0.03/0.23  % Model      : x86_64 x86_64
 % 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.23  % Memory     : 32218.75MB
 % 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.23  % CPULimit   : 300
 % 0.03/0.23  % DateTime   : Tue May 10 20:35:53 CDT 2016
 % 0.03/0.23  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.03/0.27  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.03/0.27  # and selection function SelectCQPrecWNTNp.
 % 0.03/0.27  #
 % 0.03/0.27  # Preprocessing time       : 0.014 s
 % 0.03/0.27  # Presaturation interreduction done
 % 0.03/0.27
 % 0.03/0.27  # Proof found!
 % 0.03/0.27  # SZS status Theorem
 % 0.03/0.27  # SZS output start CNFRefutation
 % 0.03/0.27  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(((app(app(X5,X3),X6)!=X4|~(strictorderedP(X3)))|?[X7]:(ssItem(X7)&?[X8]:((ssList(X8)&app(X8,cons(X7,nil))=X5)&?[X9]:(ssItem(X9)&?[X10]:((ssList(X10)&app(cons(X9,nil),X10)=X3)<(X7,X9))))))|?[X11]:(ssItem(X11)&?[X12]:((ssList(X12)&app(cons(X11,nil),X12)=X6)&?[X13]:(ssItem(X13)&?[X14]:((ssList(X14)&app(X14,cons(X13,nil))=X3)<(X13,X11)))))))))|(nil!=X4&nil=X3))|((nil!=X2|nil=X1)&(~(neq(X2,nil))|(neq(X1,nil)&segmentP(X2,X1)))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.03/0.27  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)<=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax7)).
 % 0.03/0.27  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
 % 0.03/0.27  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax79)).
 % 0.03/0.27  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 0.03/0.27  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 0.03/0.27  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 0.03/0.27  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 0.03/0.27  fof(c_0_8, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.03/0.27  fof(c_0_9, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(((app(app(X5,X3),X6)!=X4|~strictorderedP(X3))|?[X7]:(ssItem(X7)&?[X8]:((ssList(X8)&app(X8,cons(X7,nil))=X5)&?[X9]:(ssItem(X9)&?[X10]:((ssList(X10)&app(cons(X9,nil),X10)=X3)<(X7,X9))))))|?[X11]:(ssItem(X11)&?[X12]:((ssList(X12)&app(cons(X11,nil),X12)=X6)&?[X13]:(ssItem(X13)&?[X14]:((ssList(X14)&app(X14,cons(X13,nil))=X3)<(X13,X11)))))))))|(nil!=X4&nil=X3))|((nil!=X2|nil=X1)&(~neq(X2,nil)|(neq(X1,nil)&segmentP(X2,X1))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.03/0.27  fof(c_0_10, negated_conjecture, (![X21]:![X22]:![X23]:![X24]:![X25]:![X26]:![X27]:![X28]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(ssList(esk52_0)&(ssList(esk53_0)&(((app(app(esk52_0,esk50_0),esk53_0)=esk51_0&strictorderedP(esk50_0))&(~ssItem(X21)|((~ssList(X22)|app(X22,cons(X21,nil))!=esk52_0)|(~ssItem(X23)|((~ssList(X24)|app(cons(X23,nil),X24)!=esk50_0)|~lt(X21,X23))))))&(~ssItem(X25)|((~ssList(X26)|app(cons(X25,nil),X26)!=esk53_0)|(~ssItem(X27)|((~ssList(X28)|app(X28,cons(X27,nil))!=esk50_0)|~lt(X27,X25)))))))))&(nil=esk51_0|nil!=esk50_0))&(((neq(esk49_0,nil)|nil=esk49_0)&((~neq(esk48_0,nil)|~segmentP(esk49_0,esk48_0))|nil=esk49_0))&((neq(esk49_0,nil)|nil!=esk48_0)&((~neq(esk48_0,nil)|~segmentP(esk49_0,esk48_0))|nil!=esk48_0))))))))),
 % inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])])).
 % 0.03/0.27  fof(c_0_11, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
 % 0.03/0.27  cnf(c_0_12, negated_conjecture, (app(app(esk52_0,esk50_0),esk53_0)=esk51_0), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.03/0.27  cnf(c_0_13, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.03/0.27  cnf(c_0_14, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.03/0.27  fof(c_0_15, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
 % 0.03/0.27  cnf(c_0_16, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.03/0.27  fof(c_0_17, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
 % 0.03/0.27  fof(c_0_18, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
 % 0.03/0.27  cnf(c_0_19, negated_conjecture, (app(app(esk52_0,esk48_0),esk53_0)=esk49_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13]), c_0_14])).
 % 0.03/0.27  cnf(c_0_20, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
 % 0.03/0.27  cnf(c_0_21, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.03/0.27  cnf(c_0_22, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.03/0.27  cnf(c_0_23, negated_conjecture, (ssList(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.03/0.27  cnf(c_0_24, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_16])).
 % 0.03/0.27  cnf(c_0_25, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.03/0.27  cnf(c_0_26, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.03/0.27  fof(c_0_27, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
 % 0.03/0.27  cnf(c_0_28, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
 % 0.03/0.27  cnf(c_0_29, negated_conjecture, (app(esk52_0,app(esk48_0,esk53_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22]), c_0_23])])).
 % 0.03/0.27  cnf(c_0_30, negated_conjecture, (nil=esk49_0|~segmentP(esk49_0,esk48_0)|~neq(esk48_0,nil)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.03/0.27  cnf(c_0_31, negated_conjecture, (segmentP(esk49_0,esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_19]), c_0_25]), c_0_21]), c_0_23]), c_0_22])])).
 % 0.03/0.27  fof(c_0_32, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
 % 0.03/0.27  cnf(c_0_33, negated_conjecture, (nil=esk51_0|nil!=esk50_0), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.03/0.27  cnf(c_0_34, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
 % 0.03/0.27  cnf(c_0_35, negated_conjecture, (X1=app(esk52_0,esk48_0)|app(X1,esk53_0)!=esk49_0|~ssList(app(esk52_0,esk48_0))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_19]), c_0_21])])).
 % 0.03/0.27  cnf(c_0_36, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
 % 0.03/0.27  fof(c_0_37, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
 % 0.03/0.27  cnf(c_0_38, negated_conjecture, (esk52_0=nil|esk49_0!=nil|~ssList(app(esk48_0,esk53_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_23])])).
 % 0.03/0.27  cnf(c_0_39, negated_conjecture, (esk49_0=nil|~neq(esk48_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31])])).
 % 0.03/0.27  cnf(c_0_40, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
 % 0.03/0.27  cnf(c_0_41, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_8])).
 % 0.03/0.27  cnf(c_0_42, negated_conjecture, (esk49_0=nil|esk48_0!=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_14]), c_0_13])).
 % 0.03/0.27  cnf(c_0_43, negated_conjecture, (esk53_0=nil|esk49_0!=nil|~ssList(app(esk52_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_19]), c_0_21])])).
 % 0.03/0.27  cnf(c_0_44, negated_conjecture, (X1=app(esk52_0,esk48_0)|app(X1,esk53_0)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_22]), c_0_23])])).
 % 0.03/0.27  cnf(c_0_45, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
 % 0.03/0.27  cnf(c_0_46, negated_conjecture, (esk52_0=nil|esk49_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_36]), c_0_21]), c_0_22])])).
 % 0.03/0.27  cnf(c_0_47, negated_conjecture, (esk49_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_22])]), c_0_42])).
 % 0.03/0.27  cnf(c_0_48, negated_conjecture, (esk53_0=nil|esk49_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_36]), c_0_22]), c_0_23])])).
 % 0.03/0.27  cnf(c_0_49, negated_conjecture, (app(esk52_0,esk48_0)=nil|esk53_0!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_41]), c_0_21])])).
 % 0.03/0.27  cnf(c_0_50, negated_conjecture, (esk52_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47])])).
 % 0.03/0.27  cnf(c_0_51, negated_conjecture, (esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_47])])).
 % 0.03/0.27  cnf(c_0_52, negated_conjecture, (nil!=esk48_0|~segmentP(esk49_0,esk48_0)|~neq(esk48_0,nil)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.03/0.27  cnf(c_0_53, negated_conjecture, (app(nil,esk48_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_47]), c_0_50]), c_0_51])])).
 % 0.03/0.27  cnf(c_0_54, negated_conjecture, (esk48_0!=nil|~neq(esk48_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_31])])).
 % 0.03/0.27  cnf(c_0_55, negated_conjecture, (esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_53]), c_0_22]), c_0_41])])).
 % 0.03/0.27  cnf(c_0_56, negated_conjecture, (neq(esk49_0,nil)|nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.03/0.27  cnf(c_0_57, negated_conjecture, (~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55]), c_0_55])])).
 % 0.03/0.27  cnf(c_0_58, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_47]), c_0_55])]), c_0_57]), ['proof']).
 % 0.03/0.27  # SZS output end CNFRefutation
 % 0.03/0.27  # Training examples: 0 positive, 0 negative
 % 0.03/0.27
 % 0.03/0.27  # -------------------------------------------------
 % 0.03/0.27  # User time                : 0.039 s
 % 0.03/0.27  # System time              : 0.002 s
 % 0.03/0.27  # Total time               : 0.042 s
 % 0.03/0.27  # Maximum resident set size: 1880 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
