   --------------------------------------------------------------------------

                   TSTP Solution File: SWC183+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC183+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n112.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:58 EDT 2016

 % Result     : Theorem 0.06s
 % Output     : CNFRefutation 0.06s
 % Verified   :
 % Statistics : Number of formulae       :  248 (65439 expanded)
 %              Number of clauses        :  171 (24013 expanded)
 %              Number of leaves         :   38 (16497 expanded)
 %              Depth                    :   33
 %              Number of atoms          :  918 (828432 expanded)
 %              Number of equality atoms :  301 (305031 expanded)
 %              Maximal formula depth    :   25 (   5 average)
 %              Maximal clause size      :   34 (   3 average)
 %              Maximal term depth       :    5 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ~ ssList(X4)
                   | X2 != X4
                   | X1 != X3
                   | ! [X5] :
                       ( ssItem(X5)
                      => ! [X6] :
                           ( ssList(X6)
                          => ! [X7] :
                               ( ~ ssList(X7)
                               | app(app(X6,cons(X5,nil)),X7) != X1
                               | ( ~ memberP(X6,X5)
                                 & ~ memberP(X7,X5) ) ) ) )
                   | ( ! [X8] :
                         ( ~ ssItem(X8)
                         | cons(X8,nil) != X3
                         | ~ memberP(X4,X8)
                         | ? [X9] :
                             ( ssItem(X9)
                             & X8 != X9
                             & memberP(X4,X9)
                             & leq(X8,X9) ) )
                     & ( nil != X4
                       | nil != X3 ) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax82)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( frontsegP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & app(X2,X3) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax5)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_5,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax81)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => cons(hd(X1),tl(X1)) = X1 ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax78)).

 fof(c_0_8,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => hd(cons(X2,X1)) = X2 ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax23)).

 fof(c_0_9,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => nil != cons(X2,X1) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax21)).

 fof(c_0_10,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( singletonP(X1)
       <=> ? [X2] :
             ( ssItem(X2)
             & cons(X2,nil) = X1 ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax4)).

 fof(c_0_11,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => tl(cons(X2,X1)) = X1 ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax25)).

 fof(c_0_12,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil != X1
            => hd(app(X1,X2)) = hd(X1) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax85)).

 fof(c_0_13,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_14,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssItem(hd(X1)) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax22)).

 fof(c_0_15,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ~ memberP(nil,X1) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax38)).

 fof(c_0_16,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ( memberP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & ? [X4] :
                     ( ssList(X4)
                     & app(X3,cons(X2,X4)) = X1 ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax3)).

 fof(c_0_17,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssList(tl(X1)) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax24)).

 fof(c_0_18,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil != X1
            => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax86)).

 fof(c_0_19,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( memberP(cons(X2,X3),X1)
               <=> ( X1 = X2
                   | memberP(X3,X1) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax37)).

 fof(c_0_20,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( app(X2,X3) = app(X2,X1)
                => X3 = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax80)).

 fof(c_0_21,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(X1,nil) = X1 ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax84)).

 fof(c_0_22,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( equalelemsP(X1)
       <=> ! [X2] :
             ( ssItem(X2)
            => ! [X3] :
                 ( ssItem(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ! [X5] :
                         ( ssList(X5)
                        => ( app(X4,cons(X2,cons(X3,X5))) = X1
                          => X2 = X3 ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax14)).

 fof(c_0_23,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil = X1
         | ? [X2] :
             ( ssList(X2)
             & ? [X3] :
                 ( ssItem(X3)
                 & cons(X3,X2) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax20)).

 fof(c_0_24,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => equalelemsP(cons(X1,nil)) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax73)).

 fof(c_0_25,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( frontsegP(cons(X1,X3),cons(X2,X4))
                   <=> ( X1 = X2
                       & frontsegP(X3,X4) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax44)).

 fof(c_0_26,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( app(X3,X2) = app(X1,X2)
                => X3 = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax79)).

 fof(c_0_27,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(nil,X1) = X1 ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax28)).

 fof(c_0_28,axiom,(
     ! [X1] :
       ( ssList(X1)
      => frontsegP(X1,X1) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax42)).

 fof(c_0_29,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( segmentP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & ? [X4] :
                     ( ssList(X4)
                     & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax7)).

 fof(c_0_30,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( rearsegP(X1,X2)
                => rearsegP(app(X3,X1),X2) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax50)).

 fof(c_0_31,axiom,(
     ! [X1] :
       ( ssList(X1)
      => rearsegP(X1,nil) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax51)).

 fof(c_0_32,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( frontsegP(X1,X2)
                => frontsegP(app(X1,X3),X2) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax43)).

 fof(c_0_33,axiom,(
     ! [X1] :
       ( ssList(X1)
      => frontsegP(X1,nil) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax45)).

 fof(c_0_34,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( rearsegP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & app(X3,X2) = X1 ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax6)).

 fof(c_0_35,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssItem(X3)
              => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax27)).

 fof(c_0_36,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( ( rearsegP(X1,X2)
               & rearsegP(X2,X1) )
            => X1 = X2 ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax48)).

 fof(c_0_37,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( ( segmentP(X1,X2)
               & segmentP(X2,X1) )
            => X1 = X2 ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax',ax54)).

 fof(c_0_38,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ~ ssList(X4)
                     | X2 != X4
                     | X1 != X3
                     | ! [X5] :
                         ( ssItem(X5)
                        => ! [X6] :
                             ( ssList(X6)
                            => ! [X7] :
                                 ( ~ ssList(X7)
                                 | app(app(X6,cons(X5,nil)),X7) != X1
                                 | ( ~ memberP(X6,X5)
                                   & ~ memberP(X7,X5) ) ) ) )
                     | ( ! [X8] :
                           ( ~ ssItem(X8)
                           | cons(X8,nil) != X3
                           | ~ memberP(X4,X8)
                           | ? [X9] :
                               ( ssItem(X9)
                               & X8 != X9
                               & memberP(X4,X9)
                               & leq(X8,X9) ) )
                       & ( nil != X4
                         | nil != X3 ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_39,negated_conjecture,(
     ! [X18] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ssItem(esk52_0)
       & ssList(esk53_0)
       & ssList(esk54_0)
       & app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk48_0
       & ( memberP(esk53_0,esk52_0)
         | memberP(esk54_0,esk52_0) )
       & ( nil = esk51_0
         | ssItem(esk55_0) )
       & ( nil = esk50_0
         | ssItem(esk55_0) )
       & ( nil = esk51_0
         | cons(esk55_0,nil) = esk50_0 )
       & ( nil = esk50_0
         | cons(esk55_0,nil) = esk50_0 )
       & ( nil = esk51_0
         | memberP(esk51_0,esk55_0) )
       & ( nil = esk50_0
         | memberP(esk51_0,esk55_0) )
       & ( nil = esk51_0
         | ~ ssItem(X18)
         | esk55_0 = X18
         | ~ memberP(esk51_0,X18)
         | ~ leq(esk55_0,X18) )
       & ( nil = esk50_0
         | ~ ssItem(X18)
         | esk55_0 = X18
         | ~ memberP(esk51_0,X18)
         | ~ leq(esk55_0,X18) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])])])])).

 cnf(c_0_40,negated_conjecture,
     ( app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk48_0 ),
     inference(split_conjunct,[status(thm)],[[4]c_0_39])).

 cnf(c_0_41,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_39])).

 fof(c_0_42,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

 cnf(c_0_43,negated_conjecture,
     ( app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk50_0 ),
     inference(rw,[status(thm)],[[6]c_0_40,[7]c_0_41])).

 cnf(c_0_44,plain,
     ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[8]c_0_42])).

 cnf(c_0_45,negated_conjecture,
     ( ssList(esk54_0) ),
     inference(split_conjunct,[status(thm)],[[9]c_0_39])).

 cnf(c_0_46,negated_conjecture,
     ( ssList(esk53_0) ),
     inference(split_conjunct,[status(thm)],[[10]c_0_39])).

 fof(c_0_47,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

 fof(c_0_48,plain,(
     ! [X4,X5,X7] :
       ( ( ssList(esk6_2(X4,X5))
         | ~ frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( app(X5,esk6_2(X4,X5)) = X4
         | ~ frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( ~ ssList(X7)
         | app(X5,X7) != X4
         | frontsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).

 fof(c_0_49,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

 cnf(c_0_50,negated_conjecture,
     ( app(esk53_0,app(cons(esk52_0,nil),esk54_0)) = esk50_0
     | ~ ssList(cons(esk52_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[11]c_0_43,c_0_44]),c_0_45]),c_0_46])])).

 cnf(c_0_51,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_47])).

 cnf(c_0_52,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[13]c_0_5])).

 cnf(c_0_53,negated_conjecture,
     ( ssItem(esk52_0) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_39])).

 fof(c_0_54,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).

 cnf(c_0_55,plain,
     ( frontsegP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(X2,X3) != X1
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_48])).

 cnf(c_0_56,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_49])).

 cnf(c_0_57,negated_conjecture,
     ( app(esk53_0,app(cons(esk52_0,nil),esk54_0)) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_50,c_0_51]),c_0_52]),c_0_53])])).

 cnf(c_0_58,plain,
     ( cons(X1,X2) = app(cons(X1,nil),X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[18]c_0_54])).

 fof(c_0_59,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | cons(hd(X2),tl(X2)) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

 fof(c_0_60,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | hd(cons(X4,X3)) = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

 fof(c_0_61,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | nil != cons(X4,X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

 fof(c_0_62,plain,(
     ! [X3,X5] :
       ( ( ssItem(esk5_1(X3))
         | ~ singletonP(X3)
         | ~ ssList(X3) )
       & ( cons(esk5_1(X3),nil) = X3
         | ~ singletonP(X3)
         | ~ ssList(X3) )
       & ( ~ ssItem(X5)
         | cons(X5,nil) != X3
         | singletonP(X3)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).

 cnf(c_0_63,plain,
     ( frontsegP(app(X1,X2),X1)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_55]),[19]c_0_56])).

 cnf(c_0_64,negated_conjecture,
     ( app(esk53_0,cons(esk52_0,esk54_0)) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[20]c_0_57,c_0_58]),c_0_45]),c_0_53])])).

 cnf(c_0_65,plain,
     ( cons(hd(X1),tl(X1)) = X1
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[21]c_0_59])).

 cnf(c_0_66,plain,
     ( hd(cons(X1,X2)) = X1
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[22]c_0_60])).

 cnf(c_0_67,plain,
     ( nil != cons(X1,X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[23]c_0_61])).

 fof(c_0_68,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | tl(cons(X4,X3)) = X3 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).

 cnf(c_0_69,plain,
     ( singletonP(X1)
     | ~ ssList(X1)
     | cons(X2,nil) != X1
     | ~ ssItem(X2) ),
     inference(split_conjunct,[status(thm)],[[24]c_0_62])).

 cnf(c_0_70,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[25]c_0_39])).

 cnf(c_0_71,negated_conjecture,
     ( frontsegP(esk50_0,esk53_0)
     | ~ ssList(cons(esk52_0,esk54_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_63,c_0_64]),c_0_46])])).

 cnf(c_0_72,plain,
     ( cons(X1,tl(cons(X1,X2))) = cons(X1,X2)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_65,c_0_66]),c_0_51]),[28]c_0_67])).

 cnf(c_0_73,plain,
     ( cons(esk5_1(X1),nil) = X1
     | ~ ssList(X1)
     | ~ singletonP(X1) ),
     inference(split_conjunct,[status(thm)],[[29]c_0_62])).

 cnf(c_0_74,plain,
     ( ssItem(esk5_1(X1))
     | ~ ssList(X1)
     | ~ singletonP(X1) ),
     inference(split_conjunct,[status(thm)],[[30]c_0_62])).

 cnf(c_0_75,plain,
     ( tl(cons(X1,X2)) = X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[31]c_0_68])).

 cnf(c_0_76,negated_conjecture,
     ( cons(esk55_0,nil) = esk50_0
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[32]c_0_39])).

 cnf(c_0_77,negated_conjecture,
     ( ssItem(esk55_0)
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[33]c_0_39])).

 cnf(c_0_78,plain,
     ( singletonP(cons(X1,nil))
     | ~ ssList(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(er,[status(thm)],[[34]c_0_69])).

 cnf(c_0_79,negated_conjecture,
     ( ssList(esk50_0) ),
     inference(rw,[status(thm)],[[35]c_0_70,[36]c_0_41])).

 fof(c_0_80,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | nil = X3
       | hd(app(X3,X4)) = hd(X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).

 cnf(c_0_81,plain,
     ( app(X2,esk6_2(X1,X2)) = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ frontsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[37]c_0_48])).

 cnf(c_0_82,negated_conjecture,
     ( frontsegP(esk50_0,esk53_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[38]c_0_71,c_0_51]),c_0_45]),c_0_53])])).

 cnf(c_0_83,plain,
     ( ssList(esk6_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ frontsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[39]c_0_48])).

 fof(c_0_84,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).

 cnf(c_0_85,plain,
     ( cons(esk5_1(X1),tl(X1)) = X1
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[40]c_0_72,c_0_73]),c_0_52])]),[41]c_0_74])).

 cnf(c_0_86,negated_conjecture,
     ( tl(esk50_0) = nil
     | esk50_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[42]c_0_75,c_0_76]),c_0_52])]),[43]c_0_77])).

 cnf(c_0_87,negated_conjecture,
     ( esk50_0 = nil
     | singletonP(esk50_0) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[44]c_0_78,c_0_76]),c_0_79])]),[45]c_0_77])).

 cnf(c_0_88,plain,
     ( hd(app(X1,X2)) = hd(X1)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[46]c_0_80])).

 cnf(c_0_89,negated_conjecture,
     ( app(esk53_0,esk6_2(esk50_0,esk53_0)) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[47]c_0_81,c_0_82]),c_0_46]),c_0_79])])).

 cnf(c_0_90,negated_conjecture,
     ( ssList(esk6_2(esk50_0,esk53_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[48]c_0_83,c_0_82]),c_0_46]),c_0_79])])).

 cnf(c_0_91,plain,
     ( nil = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[49]c_0_84])).

 fof(c_0_92,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssItem(hd(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).

 fof(c_0_93,plain,(
     ! [X1] :
       ( ssItem(X1)
      => ~ memberP(nil,X1) ) ),
     inference(fof_simplification,[status(thm)],[[50]c_0_15])).

 fof(c_0_94,plain,(
     ! [X5,X6,X9,X10] :
       ( ( ssList(esk3_2(X5,X6))
         | ~ memberP(X5,X6)
         | ~ ssItem(X6)
         | ~ ssList(X5) )
       & ( ssList(esk4_2(X5,X6))
         | ~ memberP(X5,X6)
         | ~ ssItem(X6)
         | ~ ssList(X5) )
       & ( app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6))) = X5
         | ~ memberP(X5,X6)
         | ~ ssItem(X6)
         | ~ ssList(X5) )
       & ( ~ ssList(X9)
         | ~ ssList(X10)
         | app(X9,cons(X6,X10)) != X5
         | memberP(X5,X6)
         | ~ ssItem(X6)
         | ~ ssList(X5) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])])).

 fof(c_0_95,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssList(tl(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).

 fof(c_0_96,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | nil = X3
       | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).

 cnf(c_0_97,negated_conjecture,
     ( cons(esk5_1(esk50_0),nil) = esk50_0
     | esk50_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[51]c_0_85,c_0_86]),c_0_79])]),[52]c_0_87])).

 cnf(c_0_98,plain,
     ( cons(esk5_1(X1),X2) = app(X1,X2)
     | ~ singletonP(X1)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[53]c_0_58,c_0_73]),[54]c_0_74])).

 cnf(c_0_99,negated_conjecture,
     ( hd(esk53_0) = hd(esk50_0)
     | esk53_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[55]c_0_88,c_0_89]),c_0_90]),c_0_46])])).

 cnf(c_0_100,negated_conjecture,
     ( esk53_0 = nil
     | esk50_0 != nil
     | ~ ssList(cons(esk52_0,esk54_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[56]c_0_91,c_0_64]),c_0_46])])).

 cnf(c_0_101,plain,
     ( ssItem(hd(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[57]c_0_92])).

 fof(c_0_102,plain,(
     ! [X4,X5,X6] :
       ( ( ~ memberP(cons(X5,X6),X4)
         | X4 = X5
         | memberP(X6,X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) )
       & ( X4 != X5
         | memberP(cons(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) )
       & ( ~ memberP(X6,X4)
         | memberP(cons(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).

 fof(c_0_103,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | ~ memberP(nil,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_93])])).

 cnf(c_0_104,plain,
     ( memberP(X1,X2)
     | ~ ssList(X1)
     | ~ ssItem(X2)
     | app(X3,cons(X2,X4)) != X1
     | ~ ssList(X4)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[58]c_0_94])).

 cnf(c_0_105,plain,
     ( ssList(tl(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[59]c_0_95])).

 cnf(c_0_106,plain,
     ( tl(app(X1,X2)) = app(tl(X1),X2)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[60]c_0_96])).

 cnf(c_0_107,negated_conjecture,
     ( app(esk50_0,nil) = esk50_0
     | esk50_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[61]c_0_97,c_0_98]),c_0_52]),c_0_79])]),[62]c_0_87])).

 cnf(c_0_108,negated_conjecture,
     ( cons(hd(esk50_0),tl(esk53_0)) = esk53_0
     | esk53_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[63]c_0_65,c_0_99]),c_0_46])])).

 cnf(c_0_109,negated_conjecture,
     ( hd(esk50_0) = esk55_0
     | esk50_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[64]c_0_66,c_0_76]),c_0_52])]),[65]c_0_77])).

 cnf(c_0_110,negated_conjecture,
     ( esk53_0 = nil
     | esk50_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[66]c_0_100,c_0_51]),c_0_45]),c_0_53])])).

 cnf(c_0_111,negated_conjecture,
     ( esk53_0 = nil
     | ssItem(hd(esk50_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[67]c_0_101,c_0_99]),c_0_46])])).

 cnf(c_0_112,plain,
     ( memberP(X3,X1)
     | X1 = X2
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | ~ memberP(cons(X2,X3),X1) ),
     inference(split_conjunct,[status(thm)],[[68]c_0_102])).

 cnf(c_0_113,plain,
     ( ~ memberP(nil,X1)
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[69]c_0_103])).

 cnf(c_0_114,plain,
     ( memberP(app(X1,cons(X2,X3)),X2)
     | ~ ssList(app(X1,cons(X2,X3)))
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(er,[status(thm)],[[70]c_0_104])).

 cnf(c_0_115,plain,
     ( nil = X1
     | ssList(app(tl(X1),X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[71]c_0_105,c_0_106]),c_0_56]),[72]c_0_91])).

 cnf(c_0_116,negated_conjecture,
     ( app(tl(esk50_0),nil) = tl(esk50_0)
     | esk50_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[73]c_0_106,c_0_107]),c_0_52]),c_0_79])])).

 fof(c_0_117,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(X5,X6) != app(X5,X4)
       | X6 = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).

 fof(c_0_118,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(X2,nil) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).

 cnf(c_0_119,negated_conjecture,
     ( cons(esk55_0,tl(esk53_0)) = esk53_0
     | esk53_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[74]c_0_108,c_0_109]),[75]c_0_110])).

 cnf(c_0_120,negated_conjecture,
     ( esk53_0 = nil
     | ssItem(esk55_0) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[76]c_0_111,c_0_109]),[77]c_0_110])).

 cnf(c_0_121,plain,
     ( X1 = esk5_1(X2)
     | ~ singletonP(X2)
     | ~ memberP(X2,X1)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[78]c_0_112,c_0_73]),c_0_52])]),c_0_74]),[79]c_0_113])).

 cnf(c_0_122,negated_conjecture,
     ( memberP(esk50_0,esk52_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[80]c_0_114,c_0_64]),c_0_79]),c_0_45]),c_0_46]),c_0_53])])).

 cnf(c_0_123,negated_conjecture,
     ( cons(esk55_0,tl(esk50_0)) = esk50_0
     | esk50_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[81]c_0_65,c_0_109]),c_0_79])])).

 cnf(c_0_124,negated_conjecture,
     ( esk50_0 = nil
     | ssList(tl(esk50_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[82]c_0_115,c_0_116]),c_0_52]),c_0_79])])).

 cnf(c_0_125,plain,
     ( X1 = X2
     | app(X3,X1) != app(X3,X2)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[83]c_0_117])).

 cnf(c_0_126,plain,
     ( app(X1,nil) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[84]c_0_118])).

 cnf(c_0_127,negated_conjecture,
     ( hd(esk53_0) = esk55_0
     | esk53_0 = nil
     | ~ ssList(tl(esk53_0)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[85]c_0_66,c_0_119]),[86]c_0_120])).

 cnf(c_0_128,negated_conjecture,
     ( esk5_1(esk50_0) = hd(esk50_0)
     | esk50_0 = nil
     | ~ ssItem(esk5_1(esk50_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[87]c_0_66,c_0_97]),c_0_52])])).

 cnf(c_0_129,negated_conjecture,
     ( esk5_1(esk50_0) = esk52_0
     | ~ singletonP(esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[88]c_0_121,c_0_122]),c_0_79]),c_0_53])])).

 cnf(c_0_130,negated_conjecture,
     ( esk50_0 = nil
     | X1 = esk55_0
     | memberP(tl(esk50_0),X1)
     | ~ memberP(esk50_0,X1)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[89]c_0_112,c_0_123]),c_0_77]),[90]c_0_124])).

 cnf(c_0_131,plain,
     ( X1 = nil
     | app(X2,X1) != X2
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[91]c_0_125,c_0_126]),c_0_52])])).

 cnf(c_0_132,negated_conjecture,
     ( hd(esk53_0) = esk55_0
     | esk53_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[92]c_0_127,c_0_105]),c_0_46])])).

 cnf(c_0_133,negated_conjecture,
     ( esk5_1(esk50_0) = hd(esk50_0)
     | esk50_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[93]c_0_128,c_0_74]),c_0_79])]),[94]c_0_87])).

 cnf(c_0_134,negated_conjecture,
     ( esk5_1(esk50_0) = esk52_0
     | esk50_0 = nil ),
     inference(spm,[status(thm)],[[95]c_0_129,[96]c_0_87])).

 cnf(c_0_135,negated_conjecture,
     ( esk52_0 = esk55_0
     | esk50_0 = nil
     | memberP(tl(esk50_0),esk52_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[97]c_0_130,c_0_122]),c_0_53])])).

 cnf(c_0_136,negated_conjecture,
     ( cons(esk52_0,esk54_0) = nil
     | esk53_0 != esk50_0
     | ~ ssList(cons(esk52_0,esk54_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[98]c_0_131,c_0_64]),c_0_46])])).

 cnf(c_0_137,negated_conjecture,
     ( hd(esk50_0) = esk55_0
     | esk53_0 = nil ),
     inference(spm,[status(thm)],[[99]c_0_99,[100]c_0_132])).

 cnf(c_0_138,negated_conjecture,
     ( hd(esk50_0) = esk52_0
     | esk50_0 = nil ),
     inference(spm,[status(thm)],[[101]c_0_133,[102]c_0_134])).

 cnf(c_0_139,negated_conjecture,
     ( esk50_0 = nil
     | esk52_0 = esk55_0
     | memberP(nil,esk52_0) ),
     inference(spm,[status(thm)],[[103]c_0_135,[104]c_0_86])).

 fof(c_0_140,plain,(
     ! [X6,X7,X8,X9,X10] :
       ( ( ~ equalelemsP(X6)
         | ~ ssItem(X7)
         | ~ ssItem(X8)
         | ~ ssList(X9)
         | ~ ssList(X10)
         | app(X9,cons(X7,cons(X8,X10))) != X6
         | X7 = X8
         | ~ ssList(X6) )
       & ( ssItem(esk40_1(X6))
         | equalelemsP(X6)
         | ~ ssList(X6) )
       & ( ssItem(esk41_1(X6))
         | equalelemsP(X6)
         | ~ ssList(X6) )
       & ( ssList(esk42_1(X6))
         | equalelemsP(X6)
         | ~ ssList(X6) )
       & ( ssList(esk43_1(X6))
         | equalelemsP(X6)
         | ~ ssList(X6) )
       & ( app(esk42_1(X6),cons(esk40_1(X6),cons(esk41_1(X6),esk43_1(X6)))) = X6
         | equalelemsP(X6)
         | ~ ssList(X6) )
       & ( esk40_1(X6) != esk41_1(X6)
         | equalelemsP(X6)
         | ~ ssList(X6) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])])).

 cnf(c_0_141,negated_conjecture,
     ( cons(esk52_0,esk54_0) = nil
     | esk53_0 != esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[105]c_0_136,c_0_51]),c_0_45]),c_0_53])])).

 cnf(c_0_142,negated_conjecture,
     ( esk53_0 = nil
     | esk52_0 = esk55_0 ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[106]c_0_137,c_0_138]),[107]c_0_110])).

 cnf(c_0_143,negated_conjecture,
     ( esk52_0 = esk55_0
     | esk50_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[108]c_0_113,c_0_139]),c_0_53])])).

 cnf(c_0_144,plain,
     ( X2 = X3
     | ~ ssList(X1)
     | app(X4,cons(X2,cons(X3,X5))) != X1
     | ~ ssList(X5)
     | ~ ssList(X4)
     | ~ ssItem(X3)
     | ~ ssItem(X2)
     | ~ equalelemsP(X1) ),
     inference(split_conjunct,[status(thm)],[[109]c_0_140])).

 fof(c_0_145,plain,(
     ! [X4] :
       ( ( ssList(esk44_1(X4))
         | nil = X4
         | ~ ssList(X4) )
       & ( ssItem(esk45_1(X4))
         | nil = X4
         | ~ ssList(X4) )
       & ( cons(esk45_1(X4),esk44_1(X4)) = X4
         | nil = X4
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).

 cnf(c_0_146,negated_conjecture,
     ( cons(esk52_0,esk54_0) = nil
     | esk52_0 = esk55_0 ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[110]c_0_141,c_0_142]),[111]c_0_143])).

 cnf(c_0_147,plain,
     ( nil = X2
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[112]c_0_84])).

 cnf(c_0_148,plain,
     ( X1 = X2
     | ~ equalelemsP(app(X3,cons(X1,cons(X2,X4))))
     | ~ ssList(app(X3,cons(X1,cons(X2,X4))))
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssItem(X2)
     | ~ ssItem(X1) ),
     inference(er,[status(thm)],[[113]c_0_144])).

 cnf(c_0_149,plain,
     ( nil = X1
     | cons(esk45_1(X1),esk44_1(X1)) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[114]c_0_145])).

 cnf(c_0_150,plain,
     ( nil = X1
     | ssItem(esk45_1(X1))
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[115]c_0_145])).

 cnf(c_0_151,plain,
     ( nil = X1
     | ssList(esk44_1(X1))
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[116]c_0_145])).

 cnf(c_0_152,negated_conjecture,
     ( esk52_0 = esk55_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[117]c_0_67,c_0_146]),c_0_45]),c_0_53])])).

 fof(c_0_153,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | equalelemsP(cons(X2,nil)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])).

 cnf(c_0_154,negated_conjecture,
     ( esk54_0 = nil
     | esk50_0 != nil
     | ~ ssList(app(esk53_0,cons(esk52_0,nil))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[118]c_0_147,c_0_43]),c_0_45])])).

 fof(c_0_155,plain,(
     ! [X5,X6,X7,X8] :
       ( ( X5 = X6
         | ~ frontsegP(cons(X5,X7),cons(X6,X8))
         | ~ ssList(X8)
         | ~ ssList(X7)
         | ~ ssItem(X6)
         | ~ ssItem(X5) )
       & ( frontsegP(X7,X8)
         | ~ frontsegP(cons(X5,X7),cons(X6,X8))
         | ~ ssList(X8)
         | ~ ssList(X7)
         | ~ ssItem(X6)
         | ~ ssItem(X5) )
       & ( X5 != X6
         | ~ frontsegP(X7,X8)
         | frontsegP(cons(X5,X7),cons(X6,X8))
         | ~ ssList(X8)
         | ~ ssList(X7)
         | ~ ssItem(X6)
         | ~ ssItem(X5) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])).

 cnf(c_0_156,plain,
     ( X1 = esk45_1(X2)
     | nil = X2
     | ~ equalelemsP(app(X3,cons(X1,X2)))
     | ~ ssList(app(X3,cons(X1,X2)))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[119]c_0_148,c_0_149]),c_0_150]),[120]c_0_151])).

 cnf(c_0_157,negated_conjecture,
     ( app(esk53_0,cons(esk55_0,esk54_0)) = esk50_0 ),
     inference(rw,[status(thm)],[[121]c_0_64,[122]c_0_152])).

 cnf(c_0_158,negated_conjecture,
     ( ssItem(esk55_0) ),
     inference(rw,[status(thm)],[[123]c_0_53,[124]c_0_152])).

 cnf(c_0_159,plain,
     ( equalelemsP(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[125]c_0_153])).

 cnf(c_0_160,negated_conjecture,
     ( esk54_0 = nil
     | esk50_0 != nil
     | ~ ssList(cons(esk52_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[126]c_0_154,c_0_56]),c_0_46])])).

 fof(c_0_161,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(X6,X5) != app(X4,X5)
       | X6 = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).

 fof(c_0_162,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(nil,X2) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])).

 cnf(c_0_163,plain,
     ( frontsegP(cons(X1,X3),cons(X2,X4))
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | ~ ssList(X4)
     | ~ frontsegP(X3,X4)
     | X1 != X2 ),
     inference(split_conjunct,[status(thm)],[[127]c_0_155])).

 fof(c_0_164,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | frontsegP(X2,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])).

 cnf(c_0_165,negated_conjecture,
     ( esk45_1(esk54_0) = esk55_0
     | esk54_0 = nil
     | ~ equalelemsP(esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[128]c_0_156,c_0_157]),c_0_79]),c_0_46]),c_0_45]),c_0_158])])).

 cnf(c_0_166,negated_conjecture,
     ( esk50_0 = nil
     | equalelemsP(esk50_0) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[129]c_0_159,c_0_76]),[130]c_0_77])).

 cnf(c_0_167,negated_conjecture,
     ( esk54_0 = nil
     | esk50_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[131]c_0_160,c_0_51]),c_0_52]),c_0_53])])).

 cnf(c_0_168,plain,
     ( X1 = X2
     | app(X1,X3) != app(X2,X3)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[132]c_0_161])).

 cnf(c_0_169,plain,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[133]c_0_162])).

 fof(c_0_170,plain,(
     ! [X5,X6,X9,X10] :
       ( ( ssList(esk8_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ssList(esk9_2(X5,X6))
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
         | ~ segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) )
       & ( ~ ssList(X9)
         | ~ ssList(X10)
         | app(app(X9,X6),X10) != X5
         | segmentP(X5,X6)
         | ~ ssList(X6)
         | ~ ssList(X5) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])])])).

 cnf(c_0_171,plain,
     ( frontsegP(cons(X1,X2),cons(X1,X3))
     | ~ frontsegP(X2,X3)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(er,[status(thm)],[[134]c_0_163])).

 cnf(c_0_172,plain,
     ( frontsegP(X1,X1)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[135]c_0_164])).

 cnf(c_0_173,negated_conjecture,
     ( esk45_1(esk54_0) = esk55_0
     | esk54_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[136]c_0_165,c_0_166]),[137]c_0_167])).

 cnf(c_0_174,plain,
     ( X1 = nil
     | app(X1,X2) != X2
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[138]c_0_168,c_0_169]),c_0_52])])).

 cnf(c_0_175,plain,
     ( segmentP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(app(X3,X2),X4) != X1
     | ~ ssList(X4)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[139]c_0_170])).

 cnf(c_0_176,plain,
     ( frontsegP(cons(X1,X2),cons(X1,X2))
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(spm,[status(thm)],[[140]c_0_171,[141]c_0_172])).

 cnf(c_0_177,negated_conjecture,
     ( cons(esk55_0,esk44_1(esk54_0)) = esk54_0
     | esk54_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[142]c_0_149,c_0_173]),c_0_45])])).

 fof(c_0_178,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | ~ rearsegP(X4,X5)
       | rearsegP(app(X6,X4),X5) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).

 fof(c_0_179,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | rearsegP(X2,nil) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])).

 fof(c_0_180,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | ~ frontsegP(X4,X5)
       | frontsegP(app(X4,X6),X5) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])).

 fof(c_0_181,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | frontsegP(X2,nil) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])).

 cnf(c_0_182,plain,
     ( cons(hd(X1),tl(app(X1,X2))) = app(X1,X2)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[143]c_0_65,c_0_88]),c_0_56]),[144]c_0_91])).

 cnf(c_0_183,negated_conjecture,
     ( esk53_0 = nil
     | app(cons(esk52_0,nil),esk54_0) != esk50_0
     | ~ ssList(app(cons(esk52_0,nil),esk54_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[145]c_0_174,c_0_57]),c_0_46])])).

 cnf(c_0_184,plain,
     ( segmentP(app(app(X1,X2),X3),X2)
     | ~ ssList(app(app(X1,X2),X3))
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(er,[status(thm)],[[146]c_0_175])).

 cnf(c_0_185,negated_conjecture,
     ( esk54_0 = nil
     | frontsegP(esk54_0,esk54_0)
     | ~ ssList(esk44_1(esk54_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[147]c_0_176,c_0_177]),c_0_158])])).

 cnf(c_0_186,plain,
     ( rearsegP(app(X1,X2),X3)
     | ~ rearsegP(X2,X3)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[148]c_0_178])).

 cnf(c_0_187,plain,
     ( rearsegP(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[149]c_0_179])).

 cnf(c_0_188,plain,
     ( frontsegP(app(X1,X2),X3)
     | ~ frontsegP(X1,X3)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[150]c_0_180])).

 cnf(c_0_189,plain,
     ( frontsegP(X1,nil)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[151]c_0_181])).

 fof(c_0_190,plain,(
     ! [X4,X5,X7] :
       ( ( ssList(esk7_2(X4,X5))
         | ~ rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( app(esk7_2(X4,X5),X5) = X4
         | ~ rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) )
       & ( ~ ssList(X7)
         | app(X7,X5) != X4
         | rearsegP(X4,X5)
         | ~ ssList(X5)
         | ~ ssList(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])])])])).

 cnf(c_0_191,negated_conjecture,
     ( cons(hd(esk53_0),tl(esk50_0)) = esk50_0
     | esk53_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[152]c_0_182,c_0_89]),c_0_90]),c_0_46])])).

 cnf(c_0_192,negated_conjecture,
     ( esk53_0 = nil
     | app(cons(esk52_0,nil),esk54_0) != esk50_0
     | ~ ssList(cons(esk52_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[153]c_0_183,c_0_56]),c_0_45])])).

 cnf(c_0_193,plain,
     ( segmentP(app(X1,X2),X2)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[154]c_0_184,c_0_126]),c_0_52])]),[155]c_0_56])).

 fof(c_0_194,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssItem(X6)
       | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])).

 cnf(c_0_195,negated_conjecture,
     ( esk54_0 = nil
     | frontsegP(esk54_0,esk54_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[156]c_0_185,c_0_151]),c_0_45])])).

 cnf(c_0_196,plain,
     ( rearsegP(app(X1,X2),nil)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[157]c_0_186,c_0_187]),c_0_52])])).

 cnf(c_0_197,plain,
     ( frontsegP(app(X1,X2),nil)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[158]c_0_188,c_0_189]),c_0_52])])).

 fof(c_0_198,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ~ rearsegP(X3,X4)
       | ~ rearsegP(X4,X3)
       | X3 = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).

 cnf(c_0_199,plain,
     ( rearsegP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | app(X3,X2) != X1
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[159]c_0_190])).

 cnf(c_0_200,negated_conjecture,
     ( cons(esk55_0,tl(esk50_0)) = esk50_0
     | esk53_0 = nil ),
     inference(spm,[status(thm)],[[160]c_0_191,[161]c_0_132])).

 cnf(c_0_201,negated_conjecture,
     ( esk53_0 = nil
     | app(cons(esk52_0,nil),esk54_0) != esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[162]c_0_192,c_0_51]),c_0_52]),c_0_53])])).

 fof(c_0_202,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ~ segmentP(X3,X4)
       | ~ segmentP(X4,X3)
       | X3 = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])).

 cnf(c_0_203,negated_conjecture,
     ( segmentP(esk50_0,cons(esk55_0,esk54_0))
     | ~ ssList(cons(esk55_0,esk54_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[163]c_0_193,c_0_157]),c_0_46])])).

 cnf(c_0_204,plain,
     ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[164]c_0_194])).

 cnf(c_0_205,plain,
     ( app(nil,app(X1,X2)) = app(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[165]c_0_44,c_0_169]),c_0_52])])).

 cnf(c_0_206,negated_conjecture,
     ( app(esk54_0,esk6_2(esk54_0,esk54_0)) = esk54_0
     | esk54_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[166]c_0_81,c_0_195]),c_0_45])])).

 cnf(c_0_207,negated_conjecture,
     ( esk54_0 = nil
     | ssList(esk6_2(esk54_0,esk54_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[167]c_0_83,c_0_195]),c_0_45])])).

 cnf(c_0_208,negated_conjecture,
     ( rearsegP(esk50_0,nil)
     | ~ ssList(cons(esk55_0,esk54_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[168]c_0_196,c_0_157]),c_0_46])])).

 cnf(c_0_209,negated_conjecture,
     ( frontsegP(esk50_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[169]c_0_197,c_0_89]),c_0_90]),c_0_46])])).

 cnf(c_0_210,plain,
     ( X1 = X2
     | ~ rearsegP(X2,X1)
     | ~ rearsegP(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[170]c_0_198])).

 cnf(c_0_211,plain,
     ( rearsegP(app(X1,X2),X2)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_199]),[171]c_0_56])).

 cnf(c_0_212,negated_conjecture,
     ( app(esk53_0,app(cons(esk55_0,nil),esk54_0)) = esk50_0 ),
     inference(rw,[status(thm)],[[172]c_0_57,[173]c_0_152])).

 cnf(c_0_213,negated_conjecture,
     ( cons(esk55_0,nil) = esk50_0
     | esk53_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[174]c_0_200,c_0_86]),[175]c_0_110])).

 cnf(c_0_214,negated_conjecture,
     ( esk53_0 = nil
     | app(cons(esk55_0,nil),esk54_0) != esk50_0 ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[176]c_0_201,c_0_143]),[177]c_0_110])).

 cnf(c_0_215,plain,
     ( X1 = X2
     | ~ segmentP(X2,X1)
     | ~ segmentP(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[178]c_0_202])).

 cnf(c_0_216,negated_conjecture,
     ( segmentP(esk50_0,cons(esk55_0,esk54_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[179]c_0_203,c_0_51]),c_0_45]),c_0_158])])).

 cnf(c_0_217,negated_conjecture,
     ( cons(esk55_0,app(nil,X1)) = app(esk50_0,X1)
     | esk50_0 = nil
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[180]c_0_204,c_0_76]),c_0_52])]),c_0_158])])).

 cnf(c_0_218,negated_conjecture,
     ( app(nil,esk54_0) = esk54_0
     | esk54_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[181]c_0_205,c_0_206]),c_0_45])]),[182]c_0_207])).

 cnf(c_0_219,plain,
     ( app(esk7_2(X1,X2),X2) = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ rearsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[183]c_0_190])).

 cnf(c_0_220,negated_conjecture,
     ( rearsegP(esk50_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[184]c_0_208,c_0_51]),c_0_45]),c_0_158])])).

 cnf(c_0_221,plain,
     ( ssList(esk7_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ rearsegP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[185]c_0_190])).

 cnf(c_0_222,negated_conjecture,
     ( app(nil,esk6_2(esk50_0,nil)) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[186]c_0_81,c_0_209]),c_0_52]),c_0_79])])).

 cnf(c_0_223,negated_conjecture,
     ( ssList(esk6_2(esk50_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[187]c_0_83,c_0_209]),c_0_52]),c_0_79])])).

 cnf(c_0_224,plain,
     ( app(X1,X2) = X2
     | ~ rearsegP(X2,app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[188]c_0_210,c_0_211]),[189]c_0_56])).

 cnf(c_0_225,negated_conjecture,
     ( app(esk53_0,app(esk50_0,esk54_0)) = esk50_0
     | esk53_0 = nil ),
     inference(spm,[status(thm)],[[190]c_0_212,[191]c_0_213])).

 cnf(c_0_226,negated_conjecture,
     ( esk53_0 = nil
     | app(esk50_0,esk54_0) != esk50_0 ),
     inference(spm,[status(thm)],[[192]c_0_214,[193]c_0_213])).

 cnf(c_0_227,negated_conjecture,
     ( cons(esk55_0,esk54_0) = esk50_0
     | ~ segmentP(cons(esk55_0,esk54_0),esk50_0)
     | ~ ssList(cons(esk55_0,esk54_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[194]c_0_215,c_0_216]),c_0_79])])).

 cnf(c_0_228,plain,
     ( segmentP(app(X1,X2),X1)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[195]c_0_184,c_0_169]),c_0_52])]),[196]c_0_56])).

 cnf(c_0_229,negated_conjecture,
     ( app(esk50_0,esk54_0) = cons(esk55_0,esk54_0)
     | esk54_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[197]c_0_217,c_0_218]),c_0_45])]),[198]c_0_167])).

 cnf(c_0_230,negated_conjecture,
     ( app(esk7_2(esk50_0,nil),nil) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[199]c_0_219,c_0_220]),c_0_52]),c_0_79])])).

 cnf(c_0_231,negated_conjecture,
     ( ssList(esk7_2(esk50_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[200]c_0_221,c_0_220]),c_0_52]),c_0_79])])).

 cnf(c_0_232,negated_conjecture,
     ( esk6_2(esk50_0,nil) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[201]c_0_169,c_0_222]),c_0_223])])).

 cnf(c_0_233,negated_conjecture,
     ( esk53_0 = nil
     | ~ rearsegP(app(esk50_0,esk54_0),esk50_0)
     | ~ ssList(app(esk50_0,esk54_0)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[202]c_0_224,c_0_225]),c_0_46])]),[203]c_0_226])).

 cnf(c_0_234,negated_conjecture,
     ( cons(esk55_0,esk54_0) = esk50_0
     | ~ segmentP(cons(esk55_0,esk54_0),esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[204]c_0_227,c_0_51]),c_0_45]),c_0_158])])).

 cnf(c_0_235,negated_conjecture,
     ( esk54_0 = nil
     | segmentP(cons(esk55_0,esk54_0),esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[205]c_0_228,c_0_229]),c_0_45]),c_0_79])])).

 cnf(c_0_236,negated_conjecture,
     ( esk54_0 = nil
     | cons(esk55_0,esk54_0) != esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[206]c_0_131,c_0_229]),c_0_79]),c_0_45])])).

 cnf(c_0_237,negated_conjecture,
     ( esk7_2(esk50_0,nil) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[207]c_0_126,c_0_230]),c_0_231])])).

 cnf(c_0_238,negated_conjecture,
     ( app(nil,esk50_0) = esk50_0 ),
     inference(rw,[status(thm)],[[208]c_0_222,[209]c_0_232])).

 cnf(c_0_239,negated_conjecture,
     ( memberP(esk54_0,esk52_0)
     | memberP(esk53_0,esk52_0) ),
     inference(split_conjunct,[status(thm)],[[210]c_0_39])).

 cnf(c_0_240,negated_conjecture,
     ( esk53_0 = nil
     | ~ rearsegP(app(esk50_0,esk54_0),esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[211]c_0_233,c_0_56]),c_0_45]),c_0_79])])).

 cnf(c_0_241,negated_conjecture,
     ( esk54_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[212]c_0_234,c_0_235]),[213]c_0_236])).

 cnf(c_0_242,negated_conjecture,
     ( app(esk50_0,nil) = esk50_0 ),
     inference(rw,[status(thm)],[[214]c_0_230,[215]c_0_237])).

 cnf(c_0_243,negated_conjecture,
     ( rearsegP(esk50_0,esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[216]c_0_211,c_0_238]),c_0_52]),c_0_79])])).

 cnf(c_0_244,negated_conjecture,
     ( memberP(esk54_0,esk55_0)
     | memberP(esk53_0,esk55_0) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[217]c_0_239,c_0_152]),[218]c_0_152])).

 cnf(c_0_245,negated_conjecture,
     ( esk53_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[219]c_0_240,c_0_241]),c_0_242]),c_0_243])])).

 cnf(c_0_246,negated_conjecture,
     ( memberP(nil,esk55_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[220]c_0_244,c_0_241]),c_0_245])])).

 cnf(c_0_247,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[221]c_0_113,c_0_246]),c_0_158])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC183+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.23  % Computer   : n112.star.cs.uiowa.edu
 % 0.03/0.23  % Model      : x86_64 x86_64
 % 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.23  % Memory     : 32218.75MB
 % 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.23  % CPULimit   : 300
 % 0.03/0.23  % DateTime   : Tue May 10 20:36:38 CDT 2016
 % 0.03/0.23  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.06/0.70  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.06/0.70  # and selection function SelectCQPrecWNTNp.
 % 0.06/0.70  #
 % 0.06/0.70  # Preprocessing time       : 0.014 s
 % 0.06/0.70  # Presaturation interreduction done
 % 0.06/0.70
 % 0.06/0.70  # Proof found!
 % 0.06/0.70  # SZS status Theorem
 % 0.06/0.70  # SZS output start CNFRefutation
 % 0.06/0.70  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~(ssList(X4))|X2!=X4)|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssList(X6)=>![X7]:((~(ssList(X7))|app(app(X6,cons(X5,nil)),X7)!=X1)|(~(memberP(X6,X5))&~(memberP(X7,X5)))))))|(![X8]:(((~(ssItem(X8))|cons(X8,nil)!=X3)|~(memberP(X4,X8)))|?[X9]:(((ssItem(X9)&X8!=X9)&memberP(X4,X9))&leq(X8,X9)))&(nil!=X4|nil!=X3))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
 % 0.06/0.70  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax82)).
 % 0.06/0.70  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.06/0.70  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)<=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax5)).
 % 0.06/0.70  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 0.06/0.70  fof(c_0_5, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.06/0.70  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax81)).
 % 0.06/0.70  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax78)).
 % 0.06/0.70  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax23)).
 % 0.06/0.70  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>nil!=cons(X2,X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax21)).
 % 0.06/0.70  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)<=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax4)).
 % 0.06/0.70  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax25)).
 % 0.06/0.70  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>hd(app(X1,X2))=hd(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax85)).
 % 0.06/0.70  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 0.06/0.70  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax22)).
 % 0.06/0.70  fof(c_0_15, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax38)).
 % 0.06/0.70  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>(memberP(X1,X2)<=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(X3,cons(X2,X4))=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax3)).
 % 0.06/0.70  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax24)).
 % 0.06/0.70  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax86)).
 % 0.06/0.70  fof(c_0_19, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)<=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax37)).
 % 0.06/0.70  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X2,X3)=app(X2,X1)=>X3=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax80)).
 % 0.06/0.70  fof(c_0_21, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax84)).
 % 0.06/0.70  fof(c_0_22, axiom, (![X1]:(ssList(X1)=>(equalelemsP(X1)<=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>(app(X4,cons(X2,cons(X3,X5)))=X1=>X2=X3)))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax14)).
 % 0.06/0.70  fof(c_0_23, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax20)).
 % 0.06/0.70  fof(c_0_24, axiom, (![X1]:(ssItem(X1)=>equalelemsP(cons(X1,nil)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax73)).
 % 0.06/0.70  fof(c_0_25, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(frontsegP(cons(X1,X3),cons(X2,X4))<=>(X1=X2&frontsegP(X3,X4)))))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax44)).
 % 0.06/0.70  fof(c_0_26, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax79)).
 % 0.06/0.70  fof(c_0_27, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 0.06/0.70  fof(c_0_28, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax42)).
 % 0.06/0.70  fof(c_0_29, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)<=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax7)).
 % 0.06/0.70  fof(c_0_30, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(rearsegP(X1,X2)=>rearsegP(app(X3,X1),X2)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax50)).
 % 0.06/0.70  fof(c_0_31, axiom, (![X1]:(ssList(X1)=>rearsegP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax51)).
 % 0.06/0.70  fof(c_0_32, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(frontsegP(X1,X2)=>frontsegP(app(X1,X3),X2)))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax43)).
 % 0.06/0.70  fof(c_0_33, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax45)).
 % 0.06/0.70  fof(c_0_34, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)<=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax6)).
 % 0.06/0.70  fof(c_0_35, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax27)).
 % 0.06/0.70  fof(c_0_36, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>((rearsegP(X1,X2)&rearsegP(X2,X1))=>X1=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax48)).
 % 0.06/0.70  fof(c_0_37, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>((segmentP(X1,X2)&segmentP(X2,X1))=>X1=X2)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax54)).
 % 0.06/0.70  fof(c_0_38, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~ssList(X4)|X2!=X4)|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssList(X6)=>![X7]:((~ssList(X7)|app(app(X6,cons(X5,nil)),X7)!=X1)|(~memberP(X6,X5)&~memberP(X7,X5))))))|(![X8]:(((~ssItem(X8)|cons(X8,nil)!=X3)|~memberP(X4,X8))|?[X9]:(((ssItem(X9)&X8!=X9)&memberP(X4,X9))&leq(X8,X9)))&(nil!=X4|nil!=X3)))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.06/0.70  fof(c_0_39, negated_conjecture, (![X18]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(ssItem(esk52_0)&(ssList(esk53_0)&((ssList(esk54_0)&app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk48_0)&(memberP(esk53_0,esk52_0)|memberP(esk54_0,esk52_0))))))&(((((nil=esk51_0|ssItem(esk55_0))&(nil=esk50_0|ssItem(esk55_0)))&((nil=esk51_0|cons(esk55_0,nil)=esk50_0)&(nil=esk50_0|cons(esk55_0,nil)=esk50_0)))&((nil=esk51_0|memberP(esk51_0,esk55_0))&(nil=esk50_0|memberP(esk51_0,esk55_0))))&((nil=esk51_0|(((~ssItem(X18)|esk55_0=X18)|~memberP(esk51_0,X18))|~leq(esk55_0,X18)))&(nil=esk50_0|(((~ssItem(X18)|esk55_0=X18)|~memberP(esk51_0,X18))|~leq(esk55_0,X18)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])])])])).
 % 0.06/0.70  cnf(c_0_40, negated_conjecture, (app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 0.06/0.70  cnf(c_0_41, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 0.06/0.70  fof(c_0_42, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
 % 0.06/0.70  cnf(c_0_43, negated_conjecture, (app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk50_0), inference(rw,[status(thm)],[c_0_40, c_0_41])).
 % 0.06/0.70  cnf(c_0_44, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
 % 0.06/0.70  cnf(c_0_45, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 0.06/0.70  cnf(c_0_46, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 0.06/0.70  fof(c_0_47, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
 % 0.06/0.70  fof(c_0_48, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
 % 0.06/0.70  fof(c_0_49, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 0.06/0.70  cnf(c_0_50, negated_conjecture, (app(esk53_0,app(cons(esk52_0,nil),esk54_0))=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_51, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
 % 0.06/0.70  cnf(c_0_52, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.06/0.70  cnf(c_0_53, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 0.06/0.70  fof(c_0_54, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
 % 0.06/0.70  cnf(c_0_55, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
 % 0.06/0.70  cnf(c_0_56, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
 % 0.06/0.70  cnf(c_0_57, negated_conjecture, (app(esk53_0,app(cons(esk52_0,nil),esk54_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52]), c_0_53])])).
 % 0.06/0.70  cnf(c_0_58, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
 % 0.06/0.70  fof(c_0_59, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
 % 0.06/0.70  fof(c_0_60, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
 % 0.06/0.70  fof(c_0_61, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|nil!=cons(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
 % 0.06/0.70  fof(c_0_62, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).
 % 0.06/0.70  cnf(c_0_63, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_55]), c_0_56])).
 % 0.06/0.70  cnf(c_0_64, negated_conjecture, (app(esk53_0,cons(esk52_0,esk54_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_45]), c_0_53])])).
 % 0.06/0.70  cnf(c_0_65, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
 % 0.06/0.70  cnf(c_0_66, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
 % 0.06/0.70  cnf(c_0_67, plain, (nil!=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
 % 0.06/0.70  fof(c_0_68, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
 % 0.06/0.70  cnf(c_0_69, plain, (singletonP(X1)|~ssList(X1)|cons(X2,nil)!=X1|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
 % 0.06/0.70  cnf(c_0_70, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 0.06/0.70  cnf(c_0_71, negated_conjecture, (frontsegP(esk50_0,esk53_0)|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_72, plain, (cons(X1,tl(cons(X1,X2)))=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_51]), c_0_67])).
 % 0.06/0.70  cnf(c_0_73, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
 % 0.06/0.70  cnf(c_0_74, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
 % 0.06/0.70  cnf(c_0_75, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
 % 0.06/0.70  cnf(c_0_76, negated_conjecture, (cons(esk55_0,nil)=esk50_0|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 0.06/0.70  cnf(c_0_77, negated_conjecture, (ssItem(esk55_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 0.06/0.70  cnf(c_0_78, plain, (singletonP(cons(X1,nil))|~ssList(cons(X1,nil))|~ssItem(X1)), inference(er,[status(thm)],[c_0_69])).
 % 0.06/0.70  cnf(c_0_79, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_70, c_0_41])).
 % 0.06/0.70  fof(c_0_80, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|hd(app(X3,X4))=hd(X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
 % 0.06/0.70  cnf(c_0_81, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
 % 0.06/0.70  cnf(c_0_82, negated_conjecture, (frontsegP(esk50_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_51]), c_0_45]), c_0_53])])).
 % 0.06/0.70  cnf(c_0_83, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
 % 0.06/0.70  fof(c_0_84, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).
 % 0.06/0.70  cnf(c_0_85, plain, (cons(esk5_1(X1),tl(X1))=X1|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_52])]), c_0_74])).
 % 0.06/0.70  cnf(c_0_86, negated_conjecture, (tl(esk50_0)=nil|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_52])]), c_0_77])).
 % 0.06/0.70  cnf(c_0_87, negated_conjecture, (esk50_0=nil|singletonP(esk50_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_76]), c_0_79])]), c_0_77])).
 % 0.06/0.70  cnf(c_0_88, plain, (hd(app(X1,X2))=hd(X1)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
 % 0.06/0.70  cnf(c_0_89, negated_conjecture, (app(esk53_0,esk6_2(esk50_0,esk53_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_46]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_90, negated_conjecture, (ssList(esk6_2(esk50_0,esk53_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_82]), c_0_46]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_91, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
 % 0.06/0.70  fof(c_0_92, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).
 % 0.06/0.70  fof(c_0_93, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_15])).
 % 0.06/0.70  fof(c_0_94, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk3_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&((((ssList(esk4_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&(((app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6)))=X5|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(X9,cons(X6,X10))!=X5))|memberP(X5,X6))|~ssItem(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])])).
 % 0.06/0.70  fof(c_0_95, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).
 % 0.06/0.70  fof(c_0_96, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
 % 0.06/0.70  cnf(c_0_97, negated_conjecture, (cons(esk5_1(esk50_0),nil)=esk50_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_79])]), c_0_87])).
 % 0.06/0.70  cnf(c_0_98, plain, (cons(esk5_1(X1),X2)=app(X1,X2)|~singletonP(X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_73]), c_0_74])).
 % 0.06/0.70  cnf(c_0_99, negated_conjecture, (hd(esk53_0)=hd(esk50_0)|esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_100, negated_conjecture, (esk53_0=nil|esk50_0!=nil|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_64]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_101, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
 % 0.06/0.70  fof(c_0_102, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
 % 0.06/0.70  fof(c_0_103, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_93])])).
 % 0.06/0.70  cnf(c_0_104, plain, (memberP(X1,X2)|~ssList(X1)|~ssItem(X2)|app(X3,cons(X2,X4))!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_94])).
 % 0.06/0.70  cnf(c_0_105, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
 % 0.06/0.70  cnf(c_0_106, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
 % 0.06/0.70  cnf(c_0_107, negated_conjecture, (app(esk50_0,nil)=esk50_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_52]), c_0_79])]), c_0_87])).
 % 0.06/0.70  cnf(c_0_108, negated_conjecture, (cons(hd(esk50_0),tl(esk53_0))=esk53_0|esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_99]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_109, negated_conjecture, (hd(esk50_0)=esk55_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_76]), c_0_52])]), c_0_77])).
 % 0.06/0.70  cnf(c_0_110, negated_conjecture, (esk53_0=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_51]), c_0_45]), c_0_53])])).
 % 0.06/0.70  cnf(c_0_111, negated_conjecture, (esk53_0=nil|ssItem(hd(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_99]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_112, plain, (memberP(X3,X1)|X1=X2|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~memberP(cons(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
 % 0.06/0.70  cnf(c_0_113, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
 % 0.06/0.70  cnf(c_0_114, plain, (memberP(app(X1,cons(X2,X3)),X2)|~ssList(app(X1,cons(X2,X3)))|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(er,[status(thm)],[c_0_104])).
 % 0.06/0.70  cnf(c_0_115, plain, (nil=X1|ssList(app(tl(X1),X2))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_56]), c_0_91])).
 % 0.06/0.70  cnf(c_0_116, negated_conjecture, (app(tl(esk50_0),nil)=tl(esk50_0)|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_52]), c_0_79])])).
 % 0.06/0.70  fof(c_0_117, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X5,X6)!=app(X5,X4)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
 % 0.06/0.70  fof(c_0_118, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).
 % 0.06/0.70  cnf(c_0_119, negated_conjecture, (cons(esk55_0,tl(esk53_0))=esk53_0|esk53_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_110])).
 % 0.06/0.70  cnf(c_0_120, negated_conjecture, (esk53_0=nil|ssItem(esk55_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_109]), c_0_110])).
 % 0.06/0.70  cnf(c_0_121, plain, (X1=esk5_1(X2)|~singletonP(X2)|~memberP(X2,X1)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_73]), c_0_52])]), c_0_74]), c_0_113])).
 % 0.06/0.70  cnf(c_0_122, negated_conjecture, (memberP(esk50_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_64]), c_0_79]), c_0_45]), c_0_46]), c_0_53])])).
 % 0.06/0.70  cnf(c_0_123, negated_conjecture, (cons(esk55_0,tl(esk50_0))=esk50_0|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_109]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_124, negated_conjecture, (esk50_0=nil|ssList(tl(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_52]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_125, plain, (X1=X2|app(X3,X1)!=app(X3,X2)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_117])).
 % 0.06/0.70  cnf(c_0_126, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
 % 0.06/0.70  cnf(c_0_127, negated_conjecture, (hd(esk53_0)=esk55_0|esk53_0=nil|~ssList(tl(esk53_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_119]), c_0_120])).
 % 0.06/0.70  cnf(c_0_128, negated_conjecture, (esk5_1(esk50_0)=hd(esk50_0)|esk50_0=nil|~ssItem(esk5_1(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_97]), c_0_52])])).
 % 0.06/0.70  cnf(c_0_129, negated_conjecture, (esk5_1(esk50_0)=esk52_0|~singletonP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_79]), c_0_53])])).
 % 0.06/0.70  cnf(c_0_130, negated_conjecture, (esk50_0=nil|X1=esk55_0|memberP(tl(esk50_0),X1)|~memberP(esk50_0,X1)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_123]), c_0_77]), c_0_124])).
 % 0.06/0.70  cnf(c_0_131, plain, (X1=nil|app(X2,X1)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_52])])).
 % 0.06/0.70  cnf(c_0_132, negated_conjecture, (hd(esk53_0)=esk55_0|esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_105]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_133, negated_conjecture, (esk5_1(esk50_0)=hd(esk50_0)|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_74]), c_0_79])]), c_0_87])).
 % 0.06/0.70  cnf(c_0_134, negated_conjecture, (esk5_1(esk50_0)=esk52_0|esk50_0=nil), inference(spm,[status(thm)],[c_0_129, c_0_87])).
 % 0.06/0.70  cnf(c_0_135, negated_conjecture, (esk52_0=esk55_0|esk50_0=nil|memberP(tl(esk50_0),esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_122]), c_0_53])])).
 % 0.06/0.70  cnf(c_0_136, negated_conjecture, (cons(esk52_0,esk54_0)=nil|esk53_0!=esk50_0|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_64]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_137, negated_conjecture, (hd(esk50_0)=esk55_0|esk53_0=nil), inference(spm,[status(thm)],[c_0_99, c_0_132])).
 % 0.06/0.70  cnf(c_0_138, negated_conjecture, (hd(esk50_0)=esk52_0|esk50_0=nil), inference(spm,[status(thm)],[c_0_133, c_0_134])).
 % 0.06/0.70  cnf(c_0_139, negated_conjecture, (esk50_0=nil|esk52_0=esk55_0|memberP(nil,esk52_0)), inference(spm,[status(thm)],[c_0_135, c_0_86])).
 % 0.06/0.70  fof(c_0_140, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((~equalelemsP(X6)|(~ssItem(X7)|(~ssItem(X8)|(~ssList(X9)|(~ssList(X10)|(app(X9,cons(X7,cons(X8,X10)))!=X6|X7=X8))))))|~ssList(X6))&(((ssItem(esk40_1(X6))|equalelemsP(X6))|~ssList(X6))&(((ssItem(esk41_1(X6))|equalelemsP(X6))|~ssList(X6))&(((ssList(esk42_1(X6))|equalelemsP(X6))|~ssList(X6))&(((ssList(esk43_1(X6))|equalelemsP(X6))|~ssList(X6))&(((app(esk42_1(X6),cons(esk40_1(X6),cons(esk41_1(X6),esk43_1(X6))))=X6|equalelemsP(X6))|~ssList(X6))&((esk40_1(X6)!=esk41_1(X6)|equalelemsP(X6))|~ssList(X6))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])])).
 % 0.06/0.70  cnf(c_0_141, negated_conjecture, (cons(esk52_0,esk54_0)=nil|esk53_0!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_51]), c_0_45]), c_0_53])])).
 % 0.06/0.70  cnf(c_0_142, negated_conjecture, (esk53_0=nil|esk52_0=esk55_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_110])).
 % 0.06/0.70  cnf(c_0_143, negated_conjecture, (esk52_0=esk55_0|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_139]), c_0_53])])).
 % 0.06/0.70  cnf(c_0_144, plain, (X2=X3|~ssList(X1)|app(X4,cons(X2,cons(X3,X5)))!=X1|~ssList(X5)|~ssList(X4)|~ssItem(X3)|~ssItem(X2)|~equalelemsP(X1)), inference(split_conjunct,[status(thm)],[c_0_140])).
 % 0.06/0.70  fof(c_0_145, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).
 % 0.06/0.70  cnf(c_0_146, negated_conjecture, (cons(esk52_0,esk54_0)=nil|esk52_0=esk55_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142]), c_0_143])).
 % 0.06/0.70  cnf(c_0_147, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
 % 0.06/0.70  cnf(c_0_148, plain, (X1=X2|~equalelemsP(app(X3,cons(X1,cons(X2,X4))))|~ssList(app(X3,cons(X1,cons(X2,X4))))|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_144])).
 % 0.06/0.70  cnf(c_0_149, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_145])).
 % 0.06/0.70  cnf(c_0_150, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_145])).
 % 0.06/0.70  cnf(c_0_151, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_145])).
 % 0.06/0.70  cnf(c_0_152, negated_conjecture, (esk52_0=esk55_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_146]), c_0_45]), c_0_53])])).
 % 0.06/0.70  fof(c_0_153, plain, (![X2]:(~ssItem(X2)|equalelemsP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])).
 % 0.06/0.70  cnf(c_0_154, negated_conjecture, (esk54_0=nil|esk50_0!=nil|~ssList(app(esk53_0,cons(esk52_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_43]), c_0_45])])).
 % 0.06/0.70  fof(c_0_155, plain, (![X5]:![X6]:![X7]:![X8]:(((((((X5=X6|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5))&(((((frontsegP(X7,X8)|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))&((((((X5!=X6|~frontsegP(X7,X8))|frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])).
 % 0.06/0.70  cnf(c_0_156, plain, (X1=esk45_1(X2)|nil=X2|~equalelemsP(app(X3,cons(X1,X2)))|~ssList(app(X3,cons(X1,X2)))|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_150]), c_0_151])).
 % 0.06/0.70  cnf(c_0_157, negated_conjecture, (app(esk53_0,cons(esk55_0,esk54_0))=esk50_0), inference(rw,[status(thm)],[c_0_64, c_0_152])).
 % 0.06/0.70  cnf(c_0_158, negated_conjecture, (ssItem(esk55_0)), inference(rw,[status(thm)],[c_0_53, c_0_152])).
 % 0.06/0.70  cnf(c_0_159, plain, (equalelemsP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_153])).
 % 0.06/0.70  cnf(c_0_160, negated_conjecture, (esk54_0=nil|esk50_0!=nil|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_56]), c_0_46])])).
 % 0.06/0.70  fof(c_0_161, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
 % 0.06/0.70  fof(c_0_162, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])).
 % 0.06/0.70  cnf(c_0_163, plain, (frontsegP(cons(X1,X3),cons(X2,X4))|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~ssList(X4)|~frontsegP(X3,X4)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_155])).
 % 0.06/0.70  fof(c_0_164, plain, (![X2]:(~ssList(X2)|frontsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])).
 % 0.06/0.70  cnf(c_0_165, negated_conjecture, (esk45_1(esk54_0)=esk55_0|esk54_0=nil|~equalelemsP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_79]), c_0_46]), c_0_45]), c_0_158])])).
 % 0.06/0.70  cnf(c_0_166, negated_conjecture, (esk50_0=nil|equalelemsP(esk50_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_76]), c_0_77])).
 % 0.06/0.70  cnf(c_0_167, negated_conjecture, (esk54_0=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_51]), c_0_52]), c_0_53])])).
 % 0.06/0.70  cnf(c_0_168, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_161])).
 % 0.06/0.70  cnf(c_0_169, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_162])).
 % 0.06/0.70  fof(c_0_170, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])])])).
 % 0.06/0.70  cnf(c_0_171, plain, (frontsegP(cons(X1,X2),cons(X1,X3))|~frontsegP(X2,X3)|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_163])).
 % 0.06/0.70  cnf(c_0_172, plain, (frontsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_164])).
 % 0.06/0.70  cnf(c_0_173, negated_conjecture, (esk45_1(esk54_0)=esk55_0|esk54_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_167])).
 % 0.06/0.70  cnf(c_0_174, plain, (X1=nil|app(X1,X2)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169]), c_0_52])])).
 % 0.06/0.70  cnf(c_0_175, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_170])).
 % 0.06/0.70  cnf(c_0_176, plain, (frontsegP(cons(X1,X2),cons(X1,X2))|~ssList(X2)|~ssItem(X1)), inference(spm,[status(thm)],[c_0_171, c_0_172])).
 % 0.06/0.70  cnf(c_0_177, negated_conjecture, (cons(esk55_0,esk44_1(esk54_0))=esk54_0|esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_173]), c_0_45])])).
 % 0.06/0.70  fof(c_0_178, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~rearsegP(X4,X5)|rearsegP(app(X6,X4),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
 % 0.06/0.70  fof(c_0_179, plain, (![X2]:(~ssList(X2)|rearsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])).
 % 0.06/0.70  fof(c_0_180, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~frontsegP(X4,X5)|frontsegP(app(X4,X6),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])).
 % 0.06/0.70  fof(c_0_181, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])).
 % 0.06/0.70  cnf(c_0_182, plain, (cons(hd(X1),tl(app(X1,X2)))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_88]), c_0_56]), c_0_91])).
 % 0.06/0.70  cnf(c_0_183, negated_conjecture, (esk53_0=nil|app(cons(esk52_0,nil),esk54_0)!=esk50_0|~ssList(app(cons(esk52_0,nil),esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_57]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_184, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_175])).
 % 0.06/0.70  cnf(c_0_185, negated_conjecture, (esk54_0=nil|frontsegP(esk54_0,esk54_0)|~ssList(esk44_1(esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177]), c_0_158])])).
 % 0.06/0.70  cnf(c_0_186, plain, (rearsegP(app(X1,X2),X3)|~rearsegP(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_178])).
 % 0.06/0.70  cnf(c_0_187, plain, (rearsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_179])).
 % 0.06/0.70  cnf(c_0_188, plain, (frontsegP(app(X1,X2),X3)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_180])).
 % 0.06/0.70  cnf(c_0_189, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_181])).
 % 0.06/0.70  fof(c_0_190, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])])])])).
 % 0.06/0.70  cnf(c_0_191, negated_conjecture, (cons(hd(esk53_0),tl(esk50_0))=esk50_0|esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_89]), c_0_90]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_192, negated_conjecture, (esk53_0=nil|app(cons(esk52_0,nil),esk54_0)!=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_56]), c_0_45])])).
 % 0.06/0.70  cnf(c_0_193, plain, (segmentP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_126]), c_0_52])]), c_0_56])).
 % 0.06/0.70  fof(c_0_194, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])).
 % 0.06/0.70  cnf(c_0_195, negated_conjecture, (esk54_0=nil|frontsegP(esk54_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_151]), c_0_45])])).
 % 0.06/0.70  cnf(c_0_196, plain, (rearsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187]), c_0_52])])).
 % 0.06/0.70  cnf(c_0_197, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_189]), c_0_52])])).
 % 0.06/0.70  fof(c_0_198, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|((~rearsegP(X3,X4)|~rearsegP(X4,X3))|X3=X4)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).
 % 0.06/0.70  cnf(c_0_199, plain, (rearsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X3,X2)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_190])).
 % 0.06/0.70  cnf(c_0_200, negated_conjecture, (cons(esk55_0,tl(esk50_0))=esk50_0|esk53_0=nil), inference(spm,[status(thm)],[c_0_191, c_0_132])).
 % 0.06/0.70  cnf(c_0_201, negated_conjecture, (esk53_0=nil|app(cons(esk52_0,nil),esk54_0)!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_51]), c_0_52]), c_0_53])])).
 % 0.06/0.70  fof(c_0_202, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|((~segmentP(X3,X4)|~segmentP(X4,X3))|X3=X4)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])).
 % 0.06/0.70  cnf(c_0_203, negated_conjecture, (segmentP(esk50_0,cons(esk55_0,esk54_0))|~ssList(cons(esk55_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_157]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_204, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_194])).
 % 0.06/0.70  cnf(c_0_205, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_169]), c_0_52])])).
 % 0.06/0.70  cnf(c_0_206, negated_conjecture, (app(esk54_0,esk6_2(esk54_0,esk54_0))=esk54_0|esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_195]), c_0_45])])).
 % 0.06/0.70  cnf(c_0_207, negated_conjecture, (esk54_0=nil|ssList(esk6_2(esk54_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_195]), c_0_45])])).
 % 0.06/0.70  cnf(c_0_208, negated_conjecture, (rearsegP(esk50_0,nil)|~ssList(cons(esk55_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_157]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_209, negated_conjecture, (frontsegP(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_89]), c_0_90]), c_0_46])])).
 % 0.06/0.70  cnf(c_0_210, plain, (X1=X2|~rearsegP(X2,X1)|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_198])).
 % 0.06/0.70  cnf(c_0_211, plain, (rearsegP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_199]), c_0_56])).
 % 0.06/0.70  cnf(c_0_212, negated_conjecture, (app(esk53_0,app(cons(esk55_0,nil),esk54_0))=esk50_0), inference(rw,[status(thm)],[c_0_57, c_0_152])).
 % 0.06/0.70  cnf(c_0_213, negated_conjecture, (cons(esk55_0,nil)=esk50_0|esk53_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_86]), c_0_110])).
 % 0.06/0.70  cnf(c_0_214, negated_conjecture, (esk53_0=nil|app(cons(esk55_0,nil),esk54_0)!=esk50_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_143]), c_0_110])).
 % 0.06/0.70  cnf(c_0_215, plain, (X1=X2|~segmentP(X2,X1)|~segmentP(X1,X2)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_202])).
 % 0.06/0.70  cnf(c_0_216, negated_conjecture, (segmentP(esk50_0,cons(esk55_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_51]), c_0_45]), c_0_158])])).
 % 0.06/0.70  cnf(c_0_217, negated_conjecture, (cons(esk55_0,app(nil,X1))=app(esk50_0,X1)|esk50_0=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_76]), c_0_52])]), c_0_158])])).
 % 0.06/0.70  cnf(c_0_218, negated_conjecture, (app(nil,esk54_0)=esk54_0|esk54_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206]), c_0_45])]), c_0_207])).
 % 0.06/0.70  cnf(c_0_219, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_190])).
 % 0.06/0.70  cnf(c_0_220, negated_conjecture, (rearsegP(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_51]), c_0_45]), c_0_158])])).
 % 0.06/0.70  cnf(c_0_221, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_190])).
 % 0.06/0.70  cnf(c_0_222, negated_conjecture, (app(nil,esk6_2(esk50_0,nil))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_209]), c_0_52]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_223, negated_conjecture, (ssList(esk6_2(esk50_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_209]), c_0_52]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_224, plain, (app(X1,X2)=X2|~rearsegP(X2,app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_56])).
 % 0.06/0.70  cnf(c_0_225, negated_conjecture, (app(esk53_0,app(esk50_0,esk54_0))=esk50_0|esk53_0=nil), inference(spm,[status(thm)],[c_0_212, c_0_213])).
 % 0.06/0.70  cnf(c_0_226, negated_conjecture, (esk53_0=nil|app(esk50_0,esk54_0)!=esk50_0), inference(spm,[status(thm)],[c_0_214, c_0_213])).
 % 0.06/0.70  cnf(c_0_227, negated_conjecture, (cons(esk55_0,esk54_0)=esk50_0|~segmentP(cons(esk55_0,esk54_0),esk50_0)|~ssList(cons(esk55_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_228, plain, (segmentP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_169]), c_0_52])]), c_0_56])).
 % 0.06/0.70  cnf(c_0_229, negated_conjecture, (app(esk50_0,esk54_0)=cons(esk55_0,esk54_0)|esk54_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_218]), c_0_45])]), c_0_167])).
 % 0.06/0.70  cnf(c_0_230, negated_conjecture, (app(esk7_2(esk50_0,nil),nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_220]), c_0_52]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_231, negated_conjecture, (ssList(esk7_2(esk50_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_220]), c_0_52]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_232, negated_conjecture, (esk6_2(esk50_0,nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_222]), c_0_223])])).
 % 0.06/0.70  cnf(c_0_233, negated_conjecture, (esk53_0=nil|~rearsegP(app(esk50_0,esk54_0),esk50_0)|~ssList(app(esk50_0,esk54_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_225]), c_0_46])]), c_0_226])).
 % 0.06/0.70  cnf(c_0_234, negated_conjecture, (cons(esk55_0,esk54_0)=esk50_0|~segmentP(cons(esk55_0,esk54_0),esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_51]), c_0_45]), c_0_158])])).
 % 0.06/0.70  cnf(c_0_235, negated_conjecture, (esk54_0=nil|segmentP(cons(esk55_0,esk54_0),esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_229]), c_0_45]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_236, negated_conjecture, (esk54_0=nil|cons(esk55_0,esk54_0)!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_229]), c_0_79]), c_0_45])])).
 % 0.06/0.70  cnf(c_0_237, negated_conjecture, (esk7_2(esk50_0,nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_230]), c_0_231])])).
 % 0.06/0.70  cnf(c_0_238, negated_conjecture, (app(nil,esk50_0)=esk50_0), inference(rw,[status(thm)],[c_0_222, c_0_232])).
 % 0.06/0.70  cnf(c_0_239, negated_conjecture, (memberP(esk54_0,esk52_0)|memberP(esk53_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
 % 0.06/0.70  cnf(c_0_240, negated_conjecture, (esk53_0=nil|~rearsegP(app(esk50_0,esk54_0),esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_56]), c_0_45]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_241, negated_conjecture, (esk54_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_235]), c_0_236])).
 % 0.06/0.70  cnf(c_0_242, negated_conjecture, (app(esk50_0,nil)=esk50_0), inference(rw,[status(thm)],[c_0_230, c_0_237])).
 % 0.06/0.70  cnf(c_0_243, negated_conjecture, (rearsegP(esk50_0,esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_238]), c_0_52]), c_0_79])])).
 % 0.06/0.70  cnf(c_0_244, negated_conjecture, (memberP(esk54_0,esk55_0)|memberP(esk53_0,esk55_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_239, c_0_152]), c_0_152])).
 % 0.06/0.70  cnf(c_0_245, negated_conjecture, (esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_240, c_0_241]), c_0_242]), c_0_243])])).
 % 0.06/0.70  cnf(c_0_246, negated_conjecture, (memberP(nil,esk55_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_244, c_0_241]), c_0_245])])).
 % 0.06/0.70  cnf(c_0_247, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_246]), c_0_158])]), ['proof']).
 % 0.06/0.70  # SZS output end CNFRefutation
 % 0.06/0.70  # Training examples: 0 positive, 0 negative
 % 0.06/0.70
 % 0.06/0.70  # -------------------------------------------------
 % 0.06/0.70  # User time                : 0.467 s
 % 0.06/0.70  # System time              : 0.011 s
 % 0.06/0.70  # Total time               : 0.478 s
 % 0.06/0.70  # Maximum resident set size: 1876 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
