   --------------------------------------------------------------------------

                   TSTP Solution File: SWC171+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC171+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n150.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:55 EDT 2016

 % Result     : Theorem 0.09s
 % Output     : CNFRefutation 0.09s
 % Verified   :
 % Statistics : Number of formulae       :   52 ( 220 expanded)
 %              Number of clauses        :   33 (  81 expanded)
 %              Number of leaves         :   10 (  56 expanded)
 %              Depth                    :   13
 %              Number of atoms          :  225 (2908 expanded)
 %              Number of equality atoms :   62 (1049 expanded)
 %              Maximal formula depth    :   26 (   5 average)
 %              Maximal clause size      :   34 (   2 average)
 %              Maximal term depth       :    5 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ~ ssList(X4)
                   | X2 != X4
                   | X1 != X3
                   | ! [X5] :
                       ( ssItem(X5)
                      => ! [X6] :
                           ( ~ ssItem(X6)
                           | ! [X7] :
                               ( ssList(X7)
                              => ! [X8] :
                                   ( ssList(X8)
                                  => app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8) != X1 ) )
                           | neq(X5,X6) ) )
                   | ( ! [X9] :
                         ( ~ ssItem(X9)
                         | cons(X9,nil) != X3
                         | ~ memberP(X4,X9)
                         | ? [X10] :
                             ( ssItem(X10)
                             & X9 != X10
                             & memberP(X4,X10)
                             & leq(X9,X10) ) )
                     & ( nil != X4
                       | nil != X3 ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax1)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_4,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( duplicatefreeP(X1)
       <=> ! [X2] :
             ( ssItem(X2)
            => ! [X3] :
                 ( ssItem(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ! [X5] :
                         ( ssList(X5)
                        => ! [X6] :
                             ( ssList(X6)
                            => ( app(app(X4,cons(X2,X5)),cons(X3,X6)) = X1
                              => X2 != X3 ) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax13)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

 fof(c_0_8,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => duplicatefreeP(cons(X1,nil)) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax71)).

 fof(c_0_9,axiom,(
     duplicatefreeP(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax72)).

 fof(c_0_10,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ~ ssList(X4)
                     | X2 != X4
                     | X1 != X3
                     | ! [X5] :
                         ( ssItem(X5)
                        => ! [X6] :
                             ( ~ ssItem(X6)
                             | ! [X7] :
                                 ( ssList(X7)
                                => ! [X8] :
                                     ( ssList(X8)
                                    => app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8) != X1 ) )
                             | neq(X5,X6) ) )
                     | ( ! [X9] :
                           ( ~ ssItem(X9)
                           | cons(X9,nil) != X3
                           | ~ memberP(X4,X9)
                           | ? [X10] :
                               ( ssItem(X10)
                               & X9 != X10
                               & memberP(X4,X10)
                               & leq(X9,X10) ) )
                       & ( nil != X4
                         | nil != X3 ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_11,negated_conjecture,(
     ! [X20] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ssItem(esk52_0)
       & ssItem(esk53_0)
       & ssList(esk54_0)
       & ssList(esk55_0)
       & app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0) = esk48_0
       & ~ neq(esk52_0,esk53_0)
       & ( nil = esk51_0
         | ssItem(esk56_0) )
       & ( nil = esk50_0
         | ssItem(esk56_0) )
       & ( nil = esk51_0
         | cons(esk56_0,nil) = esk50_0 )
       & ( nil = esk50_0
         | cons(esk56_0,nil) = esk50_0 )
       & ( nil = esk51_0
         | memberP(esk51_0,esk56_0) )
       & ( nil = esk50_0
         | memberP(esk51_0,esk56_0) )
       & ( nil = esk51_0
         | ~ ssItem(X20)
         | esk56_0 = X20
         | ~ memberP(esk51_0,X20)
         | ~ leq(esk56_0,X20) )
       & ( nil = esk50_0
         | ~ ssItem(X20)
         | esk56_0 = X20
         | ~ memberP(esk51_0,X20)
         | ~ leq(esk56_0,X20) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).

 fof(c_0_12,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssItem(X4)
         | ~ ssItem(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssItem(X4)
         | ~ ssItem(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).

 cnf(c_0_13,negated_conjecture,
     ( app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0) = esk48_0 ),
     inference(split_conjunct,[status(thm)],[[4]c_0_11])).

 cnf(c_0_14,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_11])).

 cnf(c_0_15,negated_conjecture,
     ( ~ neq(esk52_0,esk53_0) ),
     inference(split_conjunct,[status(thm)],[[6]c_0_11])).

 cnf(c_0_16,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssItem(X1)
     | ~ ssItem(X2) ),
     inference(split_conjunct,[status(thm)],[[7]c_0_12])).

 cnf(c_0_17,negated_conjecture,
     ( ssItem(esk53_0) ),
     inference(split_conjunct,[status(thm)],[[8]c_0_11])).

 cnf(c_0_18,negated_conjecture,
     ( ssItem(esk52_0) ),
     inference(split_conjunct,[status(thm)],[[9]c_0_11])).

 cnf(c_0_19,negated_conjecture,
     ( app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0) = esk50_0 ),
     inference(rw,[status(thm)],[[10]c_0_13,[11]c_0_14])).

 cnf(c_0_20,negated_conjecture,
     ( esk53_0 = esk52_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[12]c_0_15,c_0_16]),c_0_17]),c_0_18])])).

 fof(c_0_21,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

 cnf(c_0_22,negated_conjecture,
     ( app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk52_0,nil)),esk55_0) = esk50_0 ),
     inference(rw,[status(thm)],[[13]c_0_19,[14]c_0_20])).

 cnf(c_0_23,plain,
     ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_21])).

 cnf(c_0_24,negated_conjecture,
     ( ssList(esk55_0) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_11])).

 fof(c_0_25,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

 cnf(c_0_26,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0)) = esk50_0
     | ~ ssList(app(esk54_0,cons(esk52_0,nil)))
     | ~ ssList(cons(esk52_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_22,c_0_23]),c_0_24])])).

 cnf(c_0_27,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[18]c_0_25])).

 cnf(c_0_28,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_4])).

 fof(c_0_29,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

 cnf(c_0_30,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0)) = esk50_0
     | ~ ssList(app(esk54_0,cons(esk52_0,nil))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[20]c_0_26,c_0_27]),c_0_28]),c_0_18])])).

 cnf(c_0_31,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[21]c_0_29])).

 cnf(c_0_32,negated_conjecture,
     ( ssList(esk54_0) ),
     inference(split_conjunct,[status(thm)],[[22]c_0_11])).

 fof(c_0_33,plain,(
     ! [X7,X8,X9,X10,X11,X12] :
       ( ( ~ duplicatefreeP(X7)
         | ~ ssItem(X8)
         | ~ ssItem(X9)
         | ~ ssList(X10)
         | ~ ssList(X11)
         | ~ ssList(X12)
         | app(app(X10,cons(X8,X11)),cons(X9,X12)) != X7
         | X8 != X9
         | ~ ssList(X7) )
       & ( ssItem(esk35_1(X7))
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( ssItem(esk36_1(X7))
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk37_1(X7))
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk38_1(X7))
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( ssList(esk39_1(X7))
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( app(app(esk37_1(X7),cons(esk35_1(X7),esk38_1(X7))),cons(esk36_1(X7),esk39_1(X7))) = X7
         | duplicatefreeP(X7)
         | ~ ssList(X7) )
       & ( esk35_1(X7) = esk36_1(X7)
         | duplicatefreeP(X7)
         | ~ ssList(X7) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).

 cnf(c_0_34,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0)) = esk50_0
     | ~ ssList(cons(esk52_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[23]c_0_30,c_0_31]),c_0_32])])).

 fof(c_0_35,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).

 fof(c_0_36,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | duplicatefreeP(cons(X2,nil)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).

 cnf(c_0_37,plain,
     ( ~ ssList(X1)
     | X2 != X3
     | app(app(X4,cons(X2,X5)),cons(X3,X6)) != X1
     | ~ ssList(X6)
     | ~ ssList(X5)
     | ~ ssList(X4)
     | ~ ssItem(X3)
     | ~ ssItem(X2)
     | ~ duplicatefreeP(X1) ),
     inference(split_conjunct,[status(thm)],[[24]c_0_33])).

 cnf(c_0_38,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0)) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_34,c_0_27]),c_0_28]),c_0_18])])).

 cnf(c_0_39,plain,
     ( cons(X1,X2) = app(cons(X1,nil),X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[26]c_0_35])).

 cnf(c_0_40,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[27]c_0_11])).

 cnf(c_0_41,plain,
     ( duplicatefreeP(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[28]c_0_36])).

 cnf(c_0_42,negated_conjecture,
     ( cons(esk56_0,nil) = esk50_0
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[29]c_0_11])).

 cnf(c_0_43,negated_conjecture,
     ( ssItem(esk56_0)
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[30]c_0_11])).

 cnf(c_0_44,plain,
     ( ~ duplicatefreeP(app(app(X1,cons(X2,X3)),cons(X2,X4)))
     | ~ ssList(app(app(X1,cons(X2,X3)),cons(X2,X4)))
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_37])])).

 cnf(c_0_45,negated_conjecture,
     ( app(app(esk54_0,cons(esk52_0,nil)),cons(esk52_0,esk55_0)) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[31]c_0_38,c_0_39]),c_0_24]),c_0_18])])).

 cnf(c_0_46,negated_conjecture,
     ( ssList(esk50_0) ),
     inference(rw,[status(thm)],[[32]c_0_40,[33]c_0_14])).

 cnf(c_0_47,negated_conjecture,
     ( esk50_0 = nil
     | duplicatefreeP(esk50_0) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[34]c_0_41,c_0_42]),[35]c_0_43])).

 cnf(c_0_48,negated_conjecture,
     ( ~ duplicatefreeP(esk50_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[36]c_0_44,c_0_45]),c_0_46]),c_0_24]),c_0_28]),c_0_32]),c_0_18])])).

 cnf(c_0_49,negated_conjecture,
     ( esk50_0 = nil ),
     inference(sr,[status(thm)],[[37]c_0_47,[38]c_0_48])).

 cnf(c_0_50,plain,
     ( duplicatefreeP(nil) ),
     inference(split_conjunct,[status(thm)],[[39]c_0_9])).

 cnf(c_0_51,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[40]c_0_48,c_0_49]),c_0_50])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.01/0.03  % Problem    : SWC171+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.01/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.26  % Computer   : n150.star.cs.uiowa.edu
 % 0.03/0.26  % Model      : x86_64 x86_64
 % 0.03/0.26  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.26  % Memory     : 32218.75MB
 % 0.03/0.26  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.26  % CPULimit   : 300
 % 0.03/0.26  % DateTime   : Tue May 10 20:36:38 CDT 2016
 % 0.03/0.26  % CPUTime    :
 % 0.03/0.26  # Version: 1.9.1pre011
 % 0.03/0.26  # No SInE strategy applied
 % 0.03/0.26  # Trying AutoSched0 for 151 seconds
 % 0.09/0.41  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.09/0.41  # and selection function SelectCQPrecWNTNp.
 % 0.09/0.41  #
 % 0.09/0.41  # Preprocessing time       : 0.027 s
 % 0.09/0.41  # Presaturation interreduction done
 % 0.09/0.41
 % 0.09/0.41  # Proof found!
 % 0.09/0.41  # SZS status Theorem
 % 0.09/0.41  # SZS output start CNFRefutation
 % 0.09/0.41  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~(ssList(X4))|X2!=X4)|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:((~(ssItem(X6))|![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8)!=X1)))|neq(X5,X6))))|(![X9]:(((~(ssItem(X9))|cons(X9,nil)!=X3)|~(memberP(X4,X9)))|?[X10]:(((ssItem(X10)&X9!=X10)&memberP(X4,X10))&leq(X9,X10)))&(nil!=X4|nil!=X3))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.09/0.41  fof(c_0_1, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax1)).
 % 0.09/0.41  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
 % 0.09/0.41  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.09/0.41  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.09/0.41  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 0.09/0.41  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>(duplicatefreeP(X1)<=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(app(X4,cons(X2,X5)),cons(X3,X6))=X1=>X2!=X3))))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax13)).
 % 0.09/0.41  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
 % 0.09/0.41  fof(c_0_8, axiom, (![X1]:(ssItem(X1)=>duplicatefreeP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax71)).
 % 0.09/0.41  fof(c_0_9, axiom, (duplicatefreeP(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax72)).
 % 0.09/0.41  fof(c_0_10, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~ssList(X4)|X2!=X4)|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:((~ssItem(X6)|![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>app(app(app(X7,cons(X5,nil)),cons(X6,nil)),X8)!=X1)))|neq(X5,X6))))|(![X9]:(((~ssItem(X9)|cons(X9,nil)!=X3)|~memberP(X4,X9))|?[X10]:(((ssItem(X10)&X9!=X10)&memberP(X4,X10))&leq(X9,X10)))&(nil!=X4|nil!=X3)))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.09/0.41  fof(c_0_11, negated_conjecture, (![X20]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(ssItem(esk52_0)&((ssItem(esk53_0)&(ssList(esk54_0)&(ssList(esk55_0)&app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0)=esk48_0)))&~neq(esk52_0,esk53_0))))&(((((nil=esk51_0|ssItem(esk56_0))&(nil=esk50_0|ssItem(esk56_0)))&((nil=esk51_0|cons(esk56_0,nil)=esk50_0)&(nil=esk50_0|cons(esk56_0,nil)=esk50_0)))&((nil=esk51_0|memberP(esk51_0,esk56_0))&(nil=esk50_0|memberP(esk51_0,esk56_0))))&((nil=esk51_0|(((~ssItem(X20)|esk56_0=X20)|~memberP(esk51_0,X20))|~leq(esk56_0,X20)))&(nil=esk50_0|(((~ssItem(X20)|esk56_0=X20)|~memberP(esk51_0,X20))|~leq(esk56_0,X20)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).
 % 0.09/0.41  fof(c_0_12, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssItem(X4))|~ssItem(X3))&(((X3=X4|neq(X3,X4))|~ssItem(X4))|~ssItem(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
 % 0.09/0.41  cnf(c_0_13, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.09/0.41  cnf(c_0_14, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.09/0.41  cnf(c_0_15, negated_conjecture, (~neq(esk52_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.09/0.41  cnf(c_0_16, plain, (neq(X1,X2)|X1=X2|~ssItem(X1)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
 % 0.09/0.41  cnf(c_0_17, negated_conjecture, (ssItem(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.09/0.41  cnf(c_0_18, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.09/0.41  cnf(c_0_19, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk53_0,nil)),esk55_0)=esk50_0), inference(rw,[status(thm)],[c_0_13, c_0_14])).
 % 0.09/0.41  cnf(c_0_20, negated_conjecture, (esk53_0=esk52_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])])).
 % 0.09/0.41  fof(c_0_21, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
 % 0.09/0.41  cnf(c_0_22, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,nil)),cons(esk52_0,nil)),esk55_0)=esk50_0), inference(rw,[status(thm)],[c_0_19, c_0_20])).
 % 0.09/0.41  cnf(c_0_23, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
 % 0.09/0.41  cnf(c_0_24, negated_conjecture, (ssList(esk55_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.09/0.41  fof(c_0_25, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
 % 0.09/0.41  cnf(c_0_26, negated_conjecture, (app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0))=esk50_0|~ssList(app(esk54_0,cons(esk52_0,nil)))|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])])).
 % 0.09/0.41  cnf(c_0_27, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
 % 0.09/0.41  cnf(c_0_28, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
 % 0.09/0.41  fof(c_0_29, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
 % 0.09/0.41  cnf(c_0_30, negated_conjecture, (app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0))=esk50_0|~ssList(app(esk54_0,cons(esk52_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_18])])).
 % 0.09/0.41  cnf(c_0_31, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
 % 0.09/0.41  cnf(c_0_32, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.09/0.41  fof(c_0_33, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((~duplicatefreeP(X7)|(~ssItem(X8)|(~ssItem(X9)|(~ssList(X10)|(~ssList(X11)|(~ssList(X12)|(app(app(X10,cons(X8,X11)),cons(X9,X12))!=X7|X8!=X9)))))))|~ssList(X7))&(((ssItem(esk35_1(X7))|duplicatefreeP(X7))|~ssList(X7))&(((ssItem(esk36_1(X7))|duplicatefreeP(X7))|~ssList(X7))&(((ssList(esk37_1(X7))|duplicatefreeP(X7))|~ssList(X7))&(((ssList(esk38_1(X7))|duplicatefreeP(X7))|~ssList(X7))&(((ssList(esk39_1(X7))|duplicatefreeP(X7))|~ssList(X7))&(((app(app(esk37_1(X7),cons(esk35_1(X7),esk38_1(X7))),cons(esk36_1(X7),esk39_1(X7)))=X7|duplicatefreeP(X7))|~ssList(X7))&((esk35_1(X7)=esk36_1(X7)|duplicatefreeP(X7))|~ssList(X7)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
 % 0.09/0.41  cnf(c_0_34, negated_conjecture, (app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0))=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])])).
 % 0.09/0.41  fof(c_0_35, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
 % 0.09/0.41  fof(c_0_36, plain, (![X2]:(~ssItem(X2)|duplicatefreeP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
 % 0.09/0.41  cnf(c_0_37, plain, (~ssList(X1)|X2!=X3|app(app(X4,cons(X2,X5)),cons(X3,X6))!=X1|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssItem(X3)|~ssItem(X2)|~duplicatefreeP(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
 % 0.09/0.41  cnf(c_0_38, negated_conjecture, (app(app(esk54_0,cons(esk52_0,nil)),app(cons(esk52_0,nil),esk55_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_27]), c_0_28]), c_0_18])])).
 % 0.09/0.41  cnf(c_0_39, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
 % 0.09/0.41  cnf(c_0_40, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.09/0.41  cnf(c_0_41, plain, (duplicatefreeP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
 % 0.09/0.41  cnf(c_0_42, negated_conjecture, (cons(esk56_0,nil)=esk50_0|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.09/0.41  cnf(c_0_43, negated_conjecture, (ssItem(esk56_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.09/0.41  cnf(c_0_44, plain, (~duplicatefreeP(app(app(X1,cons(X2,X3)),cons(X2,X4)))|~ssList(app(app(X1,cons(X2,X3)),cons(X2,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_37])])).
 % 0.09/0.41  cnf(c_0_45, negated_conjecture, (app(app(esk54_0,cons(esk52_0,nil)),cons(esk52_0,esk55_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_24]), c_0_18])])).
 % 0.09/0.41  cnf(c_0_46, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_40, c_0_14])).
 % 0.09/0.41  cnf(c_0_47, negated_conjecture, (esk50_0=nil|duplicatefreeP(esk50_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])).
 % 0.09/0.41  cnf(c_0_48, negated_conjecture, (~duplicatefreeP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46]), c_0_24]), c_0_28]), c_0_32]), c_0_18])])).
 % 0.09/0.41  cnf(c_0_49, negated_conjecture, (esk50_0=nil), inference(sr,[status(thm)],[c_0_47, c_0_48])).
 % 0.09/0.41  cnf(c_0_50, plain, (duplicatefreeP(nil)), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.09/0.41  cnf(c_0_51, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49]), c_0_50])]), ['proof']).
 % 0.09/0.41  # SZS output end CNFRefutation
 % 0.09/0.41  # Training examples: 0 positive, 0 negative
 % 0.09/0.41
 % 0.09/0.41  # -------------------------------------------------
 % 0.09/0.41  # User time                : 0.142 s
 % 0.09/0.41  # System time              : 0.008 s
 % 0.09/0.41  # Total time               : 0.150 s
 % 0.09/0.41  # Maximum resident set size: 1876 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
