   --------------------------------------------------------------------------

                   TSTP Solution File: SWC183-1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC183-1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n124.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:58 EDT 2016

 % Result     : Unsatisfiable 0.02s
 % Output     : CNFRefutation 0.02s
 % Verified   :
 % Statistics : Number of clauses        :   68 ( 926 expanded)
 %              Number of leaves         :   29 ( 609 expanded)
 %              Depth                    :   23
 %              Number of atoms          :  154 (1795 expanded)
 %              Number of equality atoms :   53 ( 376 expanded)
 %              Maximal clause size      :    7 (   2 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 cnf(c_0_0,negated_conjecture,
     ( app(app(sk6,cons(sk5,nil)),sk7) = sk1 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_10)).

 cnf(c_0_1,negated_conjecture,
     ( sk1 = sk3 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

 cnf(c_0_2,axiom,
     ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

 cnf(c_0_3,negated_conjecture,
     ( ssList(sk6) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

 cnf(c_0_4,negated_conjecture,
     ( ssList(sk7) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

 cnf(c_0_5,axiom,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

 cnf(c_0_6,negated_conjecture,
     ( ssItem(sk5) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

 cnf(c_0_7,axiom,
     ( ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

 cnf(c_0_8,axiom,
     ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2)
     | ~ duplicatefreeP(X5)
     | ~ ssList(X5) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

 cnf(c_0_9,axiom,
     ( app(cons(X1,nil),X2) = cons(X1,X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

 cnf(c_0_10,axiom,
     ( duplicatefreeP(X1)
     | ssItem(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

 cnf(c_0_11,negated_conjecture,
     ( ssList(sk1) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

 cnf(c_0_12,axiom,
     ( ssList(app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

 cnf(c_0_13,axiom,
     ( frontsegP(X3,X1)
     | app(X1,X2) != X3
     | ~ ssList(X2)
     | ~ ssList(X1)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause144)).

 cnf(c_0_14,axiom,
     ( ssList(skaf42(X1,X2)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause53)).

 cnf(c_0_15,axiom,
     ( nil = X2
     | tl(app(X2,X1)) = app(tl(X2),X1)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause133)).

 cnf(c_0_16,axiom,
     ( app(X2,skaf45(X1,X2)) = X1
     | ~ frontsegP(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause132)).

 cnf(c_0_17,axiom,
     ( nil = X2
     | app(X1,X2) != nil
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause119)).

 cnf(c_0_18,axiom,
     ( tl(cons(X1,X2)) = X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause96)).

 cnf(c_0_19,negated_conjecture,
     ( cons(sk8,nil) = sk3
     | nil = sk3 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_17)).

 cnf(c_0_20,negated_conjecture,
     ( ssItem(sk8)
     | nil = sk3 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

 cnf(c_0_21,axiom,
     ( nil = X1
     | app(X1,X2) != nil
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause118)).

 cnf(c_0_22,axiom,
     ( ssList(skaf45(X1,X2)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause51)).

 cnf(c_0_23,axiom,
     ( memberP(cons(X2,X3),X1)
     | X1 != X2
     | ~ ssList(X3)
     | ~ ssItem(X2)
     | ~ ssItem(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause138)).

 cnf(c_0_24,axiom,
     ( ssList(tl(X1))
     | nil = X1
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause75)).

 cnf(c_0_25,axiom,
     ( ~ memberP(nil,X1)
     | ~ ssItem(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause71)).

 cnf(c_0_26,axiom,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause74)).

 cnf(c_0_27,axiom,
     ( cons(X1,X2) != nil
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause98)).

 cnf(c_0_28,negated_conjecture,
     ( memberP(sk6,sk5)
     | memberP(sk7,sk5) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_11)).

 cnf(c_0_29,negated_conjecture,
     ( app(app(sk6,cons(sk5,nil)),sk7) = sk3 ),
     inference(rw,[status(thm)],[[4]c_0_0,[5]c_0_1])).

 cnf(c_0_30,negated_conjecture,
     ( app(sk6,app(cons(sk5,nil),sk7)) = sk3
     | ~ ssList(cons(sk5,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[6]c_0_29,c_0_2]),c_0_3]),c_0_4])])).

 cnf(c_0_31,negated_conjecture,
     ( app(sk6,app(cons(sk5,nil),sk7)) = sk3 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[7]c_0_30,c_0_5]),c_0_6]),c_0_7])])).

 cnf(c_0_32,plain,
     ( ~ ssItem(X1)
     | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
     inference(er,[status(thm)],[[8]c_0_8])).

 cnf(c_0_33,negated_conjecture,
     ( app(sk6,cons(sk5,sk7)) = sk3 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[9]c_0_31,c_0_9]),c_0_6]),c_0_4])])).

 cnf(c_0_34,negated_conjecture,
     ( ~ ssList(app(sk3,cons(sk5,X1)))
     | ~ ssList(X1)
     | ~ duplicatefreeP(app(sk3,cons(sk5,X1))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[10]c_0_32,c_0_33]),c_0_6]),c_0_4]),c_0_3])])).

 cnf(c_0_35,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(app(sk3,cons(sk5,X2)))
     | ~ ssList(X2) ),
     inference(spm,[status(thm)],[[11]c_0_34,[12]c_0_10])).

 cnf(c_0_36,negated_conjecture,
     ( ssList(sk3) ),
     inference(rw,[status(thm)],[[13]c_0_11,[14]c_0_1])).

 cnf(c_0_37,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(cons(sk5,X2))
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_35,c_0_12]),c_0_36])])).

 cnf(c_0_38,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[16]c_0_37,c_0_5]),c_0_6])])).

 cnf(c_0_39,plain,
     ( frontsegP(app(X1,X2),X1)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]),[17]c_0_12])).

 cnf(c_0_40,negated_conjecture,
     ( ssItem(X1) ),
     inference(spm,[status(thm)],[[18]c_0_38,[19]c_0_14])).

 cnf(c_0_41,negated_conjecture,
     ( frontsegP(sk3,sk6)
     | ~ ssList(cons(sk5,sk7)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[20]c_0_39,c_0_33]),c_0_3])])).

 cnf(c_0_42,negated_conjecture,
     ( app(tl(sk6),cons(sk5,sk7)) = tl(sk3)
     | sk6 = nil
     | ~ ssList(cons(sk5,sk7)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[21]c_0_15,c_0_33]),c_0_3])])).

 cnf(c_0_43,plain,
     ( ssList(cons(X1,X2))
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[22]c_0_5,c_0_40])])).

 cnf(c_0_44,negated_conjecture,
     ( frontsegP(sk3,sk6) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[23]c_0_41,c_0_5]),c_0_6]),c_0_4])])).

 cnf(c_0_45,negated_conjecture,
     ( app(tl(sk6),cons(sk5,sk7)) = tl(sk3)
     | sk6 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[24]c_0_42,c_0_43]),c_0_4])])).

 cnf(c_0_46,negated_conjecture,
     ( app(sk6,skaf45(sk3,sk6)) = sk3 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_16,c_0_44]),c_0_3]),c_0_36])])).

 cnf(c_0_47,negated_conjecture,
     ( cons(sk5,sk7) = nil
     | sk6 = nil
     | tl(sk3) != nil
     | ~ ssList(cons(sk5,sk7))
     | ~ ssList(tl(sk6)) ),
     inference(spm,[status(thm)],[[26]c_0_17,[27]c_0_45])).

 cnf(c_0_48,negated_conjecture,
     ( tl(sk3) = nil
     | sk3 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[28]c_0_18,c_0_19]),c_0_7])]),[29]c_0_20])).

 cnf(c_0_49,negated_conjecture,
     ( sk6 = nil
     | sk3 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[30]c_0_21,c_0_46]),c_0_22]),c_0_3])])).

 cnf(c_0_50,negated_conjecture,
     ( cons(sk5,sk7) = nil
     | sk6 = nil
     | ~ ssList(cons(sk5,sk7))
     | ~ ssList(tl(sk6)) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[31]c_0_47,c_0_48]),[32]c_0_49])).

 cnf(c_0_51,plain,
     ( memberP(cons(X1,X2),X1)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(er,[status(thm)],[[33]c_0_23])).

 cnf(c_0_52,negated_conjecture,
     ( cons(sk5,sk7) = nil
     | sk6 = nil
     | ~ ssList(cons(sk5,sk7)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[34]c_0_50,c_0_24]),c_0_3])])).

 cnf(c_0_53,plain,
     ( memberP(cons(X1,X2),X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[35]c_0_51,c_0_40])])).

 cnf(c_0_54,negated_conjecture,
     ( cons(sk5,sk7) = nil
     | sk6 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[36]c_0_52,c_0_43]),c_0_4])])).

 cnf(c_0_55,plain,
     ( ~ memberP(nil,X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[37]c_0_25,c_0_40])])).

 cnf(c_0_56,negated_conjecture,
     ( sk6 = nil ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[38]c_0_53,c_0_54]),c_0_4])]),[39]c_0_55])).

 cnf(c_0_57,negated_conjecture,
     ( app(nil,cons(sk5,sk7)) = sk3 ),
     inference(rw,[status(thm)],[[40]c_0_33,[41]c_0_56])).

 cnf(c_0_58,negated_conjecture,
     ( cons(sk5,sk7) = sk3
     | ~ ssList(cons(sk5,sk7)) ),
     inference(spm,[status(thm)],[[42]c_0_26,[43]c_0_57])).

 cnf(c_0_59,plain,
     ( cons(X1,X2) != nil
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[44]c_0_27,c_0_40])])).

 cnf(c_0_60,negated_conjecture,
     ( cons(sk5,sk7) = sk3 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[45]c_0_58,c_0_43]),c_0_4])])).

 cnf(c_0_61,negated_conjecture,
     ( sk3 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[46]c_0_59,c_0_60]),c_0_4])])).

 cnf(c_0_62,plain,
     ( tl(cons(X1,X2)) = X2
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[47]c_0_18,c_0_40])])).

 cnf(c_0_63,negated_conjecture,
     ( cons(sk8,nil) = sk3 ),
     inference(sr,[status(thm)],[[48]c_0_19,[49]c_0_61])).

 cnf(c_0_64,negated_conjecture,
     ( tl(sk3) = sk7 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[50]c_0_62,c_0_60]),c_0_4])])).

 cnf(c_0_65,negated_conjecture,
     ( memberP(sk7,sk5) ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[[51]c_0_28,c_0_56]),[52]c_0_55])).

 cnf(c_0_66,negated_conjecture,
     ( sk7 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[53]c_0_62,c_0_63]),c_0_64]),c_0_7])])).

 cnf(c_0_67,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[[54]c_0_65,c_0_66]),[55]c_0_55]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC183-1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.23  % Computer   : n124.star.cs.uiowa.edu
 % 0.02/0.23  % Model      : x86_64 x86_64
 % 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.23  % Memory     : 32218.75MB
 % 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.23  % CPULimit   : 300
 % 0.02/0.23  % DateTime   : Tue May 10 20:36:38 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.02/0.29  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
 % 0.02/0.29  # and selection function SelectCQArNTEqFirst.
 % 0.02/0.29  #
 % 0.02/0.29  # Preprocessing time       : 0.009 s
 % 0.02/0.29  # Presaturation interreduction done
 % 0.02/0.29
 % 0.02/0.29  # Proof found!
 % 0.02/0.29  # SZS status Unsatisfiable
 % 0.02/0.29  # SZS output start CNFRefutation
 % 0.02/0.29  cnf(c_0_0, negated_conjecture, (app(app(sk6,cons(sk5,nil)),sk7)=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_10)).
 % 0.02/0.29  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
 % 0.02/0.29  cnf(c_0_2, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
 % 0.02/0.29  cnf(c_0_3, negated_conjecture, (ssList(sk6)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
 % 0.02/0.29  cnf(c_0_4, negated_conjecture, (ssList(sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
 % 0.02/0.29  cnf(c_0_5, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
 % 0.02/0.29  cnf(c_0_6, negated_conjecture, (ssItem(sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
 % 0.02/0.29  cnf(c_0_7, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
 % 0.02/0.29  cnf(c_0_8, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
 % 0.02/0.29  cnf(c_0_9, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
 % 0.02/0.29  cnf(c_0_10, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
 % 0.02/0.29  cnf(c_0_11, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
 % 0.02/0.29  cnf(c_0_12, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
 % 0.02/0.29  cnf(c_0_13, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause144)).
 % 0.02/0.29  cnf(c_0_14, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause53)).
 % 0.02/0.29  cnf(c_0_15, axiom, (nil=X2|tl(app(X2,X1))=app(tl(X2),X1)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause133)).
 % 0.02/0.29  cnf(c_0_16, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause132)).
 % 0.02/0.29  cnf(c_0_17, axiom, (nil=X2|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause119)).
 % 0.02/0.29  cnf(c_0_18, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause96)).
 % 0.02/0.29  cnf(c_0_19, negated_conjecture, (cons(sk8,nil)=sk3|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_17)).
 % 0.02/0.29  cnf(c_0_20, negated_conjecture, (ssItem(sk8)|nil=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
 % 0.02/0.29  cnf(c_0_21, axiom, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause118)).
 % 0.02/0.29  cnf(c_0_22, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause51)).
 % 0.02/0.29  cnf(c_0_23, axiom, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause138)).
 % 0.02/0.29  cnf(c_0_24, axiom, (ssList(tl(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause75)).
 % 0.02/0.29  cnf(c_0_25, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause71)).
 % 0.02/0.29  cnf(c_0_26, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause74)).
 % 0.02/0.29  cnf(c_0_27, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause98)).
 % 0.02/0.29  cnf(c_0_28, negated_conjecture, (memberP(sk6,sk5)|memberP(sk7,sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_11)).
 % 0.02/0.29  cnf(c_0_29, negated_conjecture, (app(app(sk6,cons(sk5,nil)),sk7)=sk3), inference(rw,[status(thm)],[c_0_0, c_0_1])).
 % 0.02/0.29  cnf(c_0_30, negated_conjecture, (app(sk6,app(cons(sk5,nil),sk7))=sk3|~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_2]), c_0_3]), c_0_4])])).
 % 0.02/0.29  cnf(c_0_31, negated_conjecture, (app(sk6,app(cons(sk5,nil),sk7))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_5]), c_0_6]), c_0_7])])).
 % 0.02/0.29  cnf(c_0_32, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_8])).
 % 0.02/0.29  cnf(c_0_33, negated_conjecture, (app(sk6,cons(sk5,sk7))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_9]), c_0_6]), c_0_4])])).
 % 0.02/0.29  cnf(c_0_34, negated_conjecture, (~ssList(app(sk3,cons(sk5,X1)))|~ssList(X1)|~duplicatefreeP(app(sk3,cons(sk5,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_6]), c_0_4]), c_0_3])])).
 % 0.02/0.29  cnf(c_0_35, negated_conjecture, (ssItem(X1)|~ssList(app(sk3,cons(sk5,X2)))|~ssList(X2)), inference(spm,[status(thm)],[c_0_34, c_0_10])).
 % 0.02/0.29  cnf(c_0_36, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_11, c_0_1])).
 % 0.02/0.29  cnf(c_0_37, negated_conjecture, (ssItem(X1)|~ssList(cons(sk5,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_12]), c_0_36])])).
 % 0.02/0.29  cnf(c_0_38, negated_conjecture, (ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_5]), c_0_6])])).
 % 0.02/0.29  cnf(c_0_39, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_12])).
 % 0.02/0.29  cnf(c_0_40, negated_conjecture, (ssItem(X1)), inference(spm,[status(thm)],[c_0_38, c_0_14])).
 % 0.02/0.29  cnf(c_0_41, negated_conjecture, (frontsegP(sk3,sk6)|~ssList(cons(sk5,sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_33]), c_0_3])])).
 % 0.02/0.29  cnf(c_0_42, negated_conjecture, (app(tl(sk6),cons(sk5,sk7))=tl(sk3)|sk6=nil|~ssList(cons(sk5,sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_33]), c_0_3])])).
 % 0.02/0.29  cnf(c_0_43, plain, (ssList(cons(X1,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_40])])).
 % 0.02/0.29  cnf(c_0_44, negated_conjecture, (frontsegP(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_5]), c_0_6]), c_0_4])])).
 % 0.02/0.29  cnf(c_0_45, negated_conjecture, (app(tl(sk6),cons(sk5,sk7))=tl(sk3)|sk6=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_4])])).
 % 0.02/0.29  cnf(c_0_46, negated_conjecture, (app(sk6,skaf45(sk3,sk6))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_44]), c_0_3]), c_0_36])])).
 % 0.02/0.29  cnf(c_0_47, negated_conjecture, (cons(sk5,sk7)=nil|sk6=nil|tl(sk3)!=nil|~ssList(cons(sk5,sk7))|~ssList(tl(sk6))), inference(spm,[status(thm)],[c_0_17, c_0_45])).
 % 0.02/0.29  cnf(c_0_48, negated_conjecture, (tl(sk3)=nil|sk3=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_7])]), c_0_20])).
 % 0.02/0.29  cnf(c_0_49, negated_conjecture, (sk6=nil|sk3!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_46]), c_0_22]), c_0_3])])).
 % 0.02/0.29  cnf(c_0_50, negated_conjecture, (cons(sk5,sk7)=nil|sk6=nil|~ssList(cons(sk5,sk7))|~ssList(tl(sk6))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
 % 0.02/0.29  cnf(c_0_51, plain, (memberP(cons(X1,X2),X1)|~ssItem(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_23])).
 % 0.02/0.29  cnf(c_0_52, negated_conjecture, (cons(sk5,sk7)=nil|sk6=nil|~ssList(cons(sk5,sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_24]), c_0_3])])).
 % 0.02/0.29  cnf(c_0_53, plain, (memberP(cons(X1,X2),X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_40])])).
 % 0.02/0.29  cnf(c_0_54, negated_conjecture, (cons(sk5,sk7)=nil|sk6=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_43]), c_0_4])])).
 % 0.02/0.29  cnf(c_0_55, plain, (~memberP(nil,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_40])])).
 % 0.02/0.29  cnf(c_0_56, negated_conjecture, (sk6=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_4])]), c_0_55])).
 % 0.02/0.29  cnf(c_0_57, negated_conjecture, (app(nil,cons(sk5,sk7))=sk3), inference(rw,[status(thm)],[c_0_33, c_0_56])).
 % 0.02/0.29  cnf(c_0_58, negated_conjecture, (cons(sk5,sk7)=sk3|~ssList(cons(sk5,sk7))), inference(spm,[status(thm)],[c_0_26, c_0_57])).
 % 0.02/0.29  cnf(c_0_59, plain, (cons(X1,X2)!=nil|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_40])])).
 % 0.02/0.29  cnf(c_0_60, negated_conjecture, (cons(sk5,sk7)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_43]), c_0_4])])).
 % 0.02/0.29  cnf(c_0_61, negated_conjecture, (sk3!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_4])])).
 % 0.02/0.29  cnf(c_0_62, plain, (tl(cons(X1,X2))=X2|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_40])])).
 % 0.02/0.29  cnf(c_0_63, negated_conjecture, (cons(sk8,nil)=sk3), inference(sr,[status(thm)],[c_0_19, c_0_61])).
 % 0.02/0.29  cnf(c_0_64, negated_conjecture, (tl(sk3)=sk7), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_60]), c_0_4])])).
 % 0.02/0.29  cnf(c_0_65, negated_conjecture, (memberP(sk7,sk5)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_56]), c_0_55])).
 % 0.02/0.29  cnf(c_0_66, negated_conjecture, (sk7=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_7])])).
 % 0.02/0.29  cnf(c_0_67, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_66]), c_0_55]), ['proof']).
 % 0.02/0.29  # SZS output end CNFRefutation
 % 0.02/0.29  # Training examples: 0 positive, 0 negative
 % 0.02/0.29
 % 0.02/0.29  # -------------------------------------------------
 % 0.02/0.29  # User time                : 0.052 s
 % 0.02/0.29  # System time              : 0.005 s
 % 0.02/0.29  # Total time               : 0.058 s
 % 0.02/0.29  # Maximum resident set size: 1864 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
