   --------------------------------------------------------------------------

                   TSTP Solution File: SWC113-1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC113-1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n147.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:42 EDT 2016

 % Result     : Unsatisfiable 154.00s
 % Output     : CNFRefutation 154.00s
 % Verified   :
 % Statistics : Number of clauses        :   80 ( 598 expanded)
 %              Number of leaves         :   41 ( 369 expanded)
 %              Depth                    :   18
 %              Number of atoms          :  206 (1104 expanded)
 %              Number of equality atoms :   79 ( 583 expanded)
 %              Maximal clause size      :    6 (   3 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 cnf(c_0_0,negated_conjecture,
     ( cons(sk5,nil) = sk3
     | nil = sk4 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

 cnf(c_0_1,negated_conjecture,
     ( sk2 = sk4 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

 cnf(c_0_2,negated_conjecture,
     ( sk1 = sk3 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

 cnf(c_0_3,negated_conjecture,
     ( ssItem(sk5)
     | nil = sk4 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

 cnf(c_0_4,axiom,
     ( cons(X1,X2) != X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause99)).

 cnf(c_0_5,axiom,
     ( ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

 cnf(c_0_6,axiom,
     ( X1 != X2
     | ~ neq(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause115)).

 cnf(c_0_7,negated_conjecture,
     ( neq(sk2,nil)
     | nil != sk1 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_15)).

 cnf(c_0_8,negated_conjecture,
     ( cons(sk5,nil) = sk3
     | nil = sk3 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_11)).

 cnf(c_0_9,negated_conjecture,
     ( ssItem(sk5)
     | nil = sk3 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

 cnf(c_0_10,axiom,
     ( cons(X1,X2) != nil
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause98)).

 cnf(c_0_11,axiom,
     ( cons(skaf44(X1),nil) = X1
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause101)).

 cnf(c_0_12,axiom,
     ( ssItem(skaf44(X1)) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause47)).

 cnf(c_0_13,axiom,
     ( singletonP(X2)
     | cons(X1,nil) != X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause116)).

 cnf(c_0_14,axiom,
     ( X4 = X2
     | cons(X1,X2) != cons(X3,X4)
     | ~ ssItem(X3)
     | ~ ssItem(X1)
     | ~ ssList(X4)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause171)).

 cnf(c_0_15,axiom,
     ( app(cons(X1,nil),X2) = cons(X1,X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause120)).

 cnf(c_0_16,axiom,
     ( tl(cons(X1,X2)) = X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause96)).

 cnf(c_0_17,axiom,
     ( cons(skaf83(X1),skaf82(X1)) = X1
     | nil = X1
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause109)).

 cnf(c_0_18,axiom,
     ( ssItem(skaf83(X1)) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause12)).

 cnf(c_0_19,axiom,
     ( ssList(skaf82(X1)) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause13)).

 cnf(c_0_20,axiom,
     ( nil = X2
     | tl(app(X2,X1)) = app(tl(X2),X1)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause133)).

 cnf(c_0_21,negated_conjecture,
     ( ssList(sk1) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

 cnf(c_0_22,axiom,
     ( ssList(app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

 cnf(c_0_23,negated_conjecture,
     ( memberP(sk4,sk5)
     | nil = sk3 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_12)).

 cnf(c_0_24,axiom,
     ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause157)).

 cnf(c_0_25,axiom,
     ( memberP(app(X3,X1),X2)
     | ~ memberP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssItem(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause141)).

 cnf(c_0_26,axiom,
     ( ~ memberP(nil,X1)
     | ~ ssItem(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause71)).

 cnf(c_0_27,axiom,
     ( app(X1,X2) = nil
     | nil != X1
     | nil != X2
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause145)).

 cnf(c_0_28,negated_conjecture,
     ( nil = sk2
     | ~ neq(sk1,nil)
     | ~ segmentP(sk2,sk1) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_14)).

 cnf(c_0_29,axiom,
     ( neq(X2,X1)
     | X2 = X1
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause100)).

 cnf(c_0_30,negated_conjecture,
     ( ssList(sk2) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

 cnf(c_0_31,axiom,
     ( segmentP(X1,X3)
     | ~ segmentP(X1,X2)
     | ~ segmentP(X2,X3)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause152)).

 cnf(c_0_32,axiom,
     ( segmentP(X4,X2)
     | app(app(X1,X2),X3) != X4
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ ssList(X4) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause173)).

 cnf(c_0_33,axiom,
     ( segmentP(X1,X1)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause57)).

 cnf(c_0_34,axiom,
     ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause149)).

 cnf(c_0_35,axiom,
     ( app(X2,skaf45(X1,X2)) = X1
     | ~ frontsegP(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause132)).

 cnf(c_0_36,axiom,
     ( ssList(skaf45(X1,X2)) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause51)).

 cnf(c_0_37,axiom,
     ( frontsegP(X1,nil)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause60)).

 cnf(c_0_38,axiom,
     ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
     | ~ memberP(X1,X2)
     | ~ ssItem(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause169)).

 cnf(c_0_39,axiom,
     ( ssList(skaf42(X1,X2)) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause53)).

 cnf(c_0_40,axiom,
     ( ssList(skaf43(X1,X2)) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause52)).

 cnf(c_0_41,negated_conjecture,
     ( cons(sk5,nil) = sk1
     | sk2 = nil ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[4]c_0_0,c_0_1]),[5]c_0_2])).

 cnf(c_0_42,negated_conjecture,
     ( sk2 = nil
     | ssItem(sk5) ),
     inference(rw,[status(thm)],[[6]c_0_3,[7]c_0_1])).

 cnf(c_0_43,negated_conjecture,
     ( sk2 = nil
     | sk1 != nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[8]c_0_4,c_0_41]),c_0_5])]),[9]c_0_42])).

 cnf(c_0_44,plain,
     ( ~ neq(X1,X1)
     | ~ ssList(X1) ),
     inference(er,[status(thm)],[[10]c_0_6])).

 cnf(c_0_45,negated_conjecture,
     ( neq(nil,nil)
     | sk1 != nil ),
     inference(spm,[status(thm)],[[11]c_0_7,[12]c_0_43])).

 cnf(c_0_46,negated_conjecture,
     ( cons(sk5,nil) = sk1
     | sk1 = nil ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[13]c_0_8,c_0_2]),[14]c_0_2])).

 cnf(c_0_47,negated_conjecture,
     ( sk1 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_44,c_0_45]),c_0_5])])).

 cnf(c_0_48,negated_conjecture,
     ( sk1 = nil
     | ssItem(sk5) ),
     inference(rw,[status(thm)],[[16]c_0_9,[17]c_0_2])).

 cnf(c_0_49,plain,
     ( X1 != nil
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_10,c_0_11]),c_0_12]),c_0_5])])).

 cnf(c_0_50,negated_conjecture,
     ( cons(sk5,nil) = sk1 ),
     inference(sr,[status(thm)],[[19]c_0_46,[20]c_0_47])).

 cnf(c_0_51,negated_conjecture,
     ( ssItem(sk5) ),
     inference(sr,[status(thm)],[[21]c_0_48,[22]c_0_47])).

 cnf(c_0_52,plain,
     ( cons(X1,nil) != X2
     | X2 != nil
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(spm,[status(thm)],[[23]c_0_49,[24]c_0_13])).

 cnf(c_0_53,negated_conjecture,
     ( X1 = nil
     | cons(X2,X1) != sk1
     | ~ ssItem(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_14,c_0_50]),c_0_51]),c_0_5])])).

 cnf(c_0_54,negated_conjecture,
     ( sk1 != X1
     | X1 != nil
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_52,c_0_50]),c_0_51])])).

 cnf(c_0_55,negated_conjecture,
     ( app(sk1,X1) = cons(sk5,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_15,c_0_50]),c_0_51])])).

 cnf(c_0_56,negated_conjecture,
     ( tl(sk1) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[28]c_0_16,c_0_50]),c_0_51]),c_0_5])])).

 cnf(c_0_57,plain,
     ( skaf82(X1) = tl(X1)
     | nil = X1
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[29]c_0_16,c_0_17]),c_0_18]),c_0_19])])).

 cnf(c_0_58,negated_conjecture,
     ( skaf82(X1) = nil
     | X1 != sk1
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[30]c_0_53,c_0_17]),c_0_18]),c_0_19])]),[31]c_0_54])).

 cnf(c_0_59,negated_conjecture,
     ( tl(cons(sk5,X1)) = app(nil,X1)
     | ~ ssList(X1) ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[32]c_0_20,c_0_55]),c_0_56]),c_0_21])]),[33]c_0_47])).

 cnf(c_0_60,negated_conjecture,
     ( tl(X1) = nil
     | X1 != sk1
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[34]c_0_57,c_0_58]),[35]c_0_54])).

 cnf(c_0_61,negated_conjecture,
     ( ssList(cons(sk5,X1))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[36]c_0_22,c_0_55]),c_0_21])])).

 cnf(c_0_62,negated_conjecture,
     ( app(nil,X1) = nil
     | cons(sk5,X1) != sk1
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[37]c_0_59,c_0_60]),[38]c_0_61])).

 cnf(c_0_63,negated_conjecture,
     ( sk1 = nil
     | memberP(sk2,sk5) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[39]c_0_23,c_0_2]),[40]c_0_1])).

 cnf(c_0_64,negated_conjecture,
     ( cons(sk5,app(nil,X1)) = app(sk1,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[41]c_0_24,c_0_50]),c_0_51]),c_0_5])])).

 cnf(c_0_65,negated_conjecture,
     ( cons(sk5,X1) != sk1
     | ~ memberP(X1,X2)
     | ~ ssItem(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[42]c_0_25,c_0_62]),c_0_5])]),[43]c_0_26])).

 cnf(c_0_66,negated_conjecture,
     ( memberP(sk2,sk5) ),
     inference(sr,[status(thm)],[[44]c_0_63,[45]c_0_47])).

 cnf(c_0_67,negated_conjecture,
     ( app(sk1,X1) = sk1
     | nil != X1
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[46]c_0_64,c_0_27]),c_0_50]),c_0_5])])).

 cnf(c_0_68,negated_conjecture,
     ( sk2 = nil
     | ~ segmentP(sk2,sk1) ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[47]c_0_28,c_0_29]),c_0_21]),c_0_5])]),[48]c_0_47])).

 cnf(c_0_69,negated_conjecture,
     ( cons(sk5,sk2) != sk1 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[49]c_0_65,c_0_66]),c_0_51]),c_0_30])])).

 cnf(c_0_70,negated_conjecture,
     ( cons(sk5,X1) = sk1
     | nil != X1
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[50]c_0_55,[51]c_0_67])).

 cnf(c_0_71,negated_conjecture,
     ( sk2 = nil
     | ~ segmentP(X1,sk1)
     | ~ segmentP(sk2,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[52]c_0_68,c_0_31]),c_0_21]),c_0_30])])).

 cnf(c_0_72,negated_conjecture,
     ( sk2 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[53]c_0_69,c_0_70]),c_0_30])])).

 cnf(c_0_73,negated_conjecture,
     ( app(app(X1,X2),X3) != sk2
     | ~ segmentP(X2,sk1)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ ssList(X1) ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[54]c_0_71,c_0_32]),c_0_30])]),[55]c_0_72])).

 cnf(c_0_74,negated_conjecture,
     ( app(app(X1,sk1),X2) != sk2
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[56]c_0_73,c_0_33]),c_0_21])])).

 cnf(c_0_75,negated_conjecture,
     ( app(X1,app(sk1,X2)) != sk2
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[57]c_0_74,c_0_34]),c_0_21])])).

 cnf(c_0_76,negated_conjecture,
     ( app(sk1,skaf45(X1,nil)) = cons(sk5,X1)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[58]c_0_64,c_0_35]),c_0_36]),c_0_5])]),[59]c_0_37])).

 cnf(c_0_77,negated_conjecture,
     ( app(X1,cons(sk5,X2)) != sk2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[60]c_0_75,c_0_76]),c_0_36])])).

 cnf(c_0_78,negated_conjecture,
     ( X1 != sk2
     | ~ memberP(X1,sk5)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[61]c_0_77,c_0_38]),c_0_39]),c_0_40]),c_0_51])])).

 cnf(c_0_79,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[62]c_0_78,c_0_66]),c_0_30])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC113-1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.22  % Computer   : n147.star.cs.uiowa.edu
 % 0.03/0.22  % Model      : x86_64 x86_64
 % 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.22  % Memory     : 32218.75MB
 % 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.22  % CPULimit   : 300
 % 0.03/0.22  % DateTime   : Tue May 10 20:35:53 CDT 2016
 % 0.03/0.23  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 151.12/151.31  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
 % 151.12/151.31  # and selection function SelectCQArNTEqFirst.
 % 151.12/151.31  #
 % 151.12/151.31  # Preprocessing time       : 0.009 s
 % 151.12/151.31  # Presaturation interreduction done
 % 151.41/151.62  # No success with AutoSched0
 % 151.41/151.62  # Trying AutoSched1 for 73 seconds
 % 154.00/154.24  # AutoSched1-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S00
 % 154.00/154.24  # and selection function NoSelection.
 % 154.00/154.24  #
 % 154.00/154.24  # Preprocessing time       : 0.010 s
 % 154.00/154.24  # Presaturation interreduction done
 % 154.00/154.24
 % 154.00/154.24  # Proof found!
 % 154.00/154.24  # SZS status Unsatisfiable
 % 154.00/154.24  # SZS output start CNFRefutation
 % 154.00/154.24  cnf(c_0_0, negated_conjecture, (cons(sk5,nil)=sk3|nil=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
 % 154.00/154.24  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
 % 154.00/154.24  cnf(c_0_2, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
 % 154.00/154.24  cnf(c_0_3, negated_conjecture, (ssItem(sk5)|nil=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
 % 154.00/154.24  cnf(c_0_4, axiom, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause99)).
 % 154.00/154.24  cnf(c_0_5, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
 % 154.00/154.24  cnf(c_0_6, axiom, (X1!=X2|~neq(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause115)).
 % 154.00/154.24  cnf(c_0_7, negated_conjecture, (neq(sk2,nil)|nil!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_15)).
 % 154.00/154.24  cnf(c_0_8, negated_conjecture, (cons(sk5,nil)=sk3|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_11)).
 % 154.00/154.24  cnf(c_0_9, negated_conjecture, (ssItem(sk5)|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
 % 154.00/154.24  cnf(c_0_10, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause98)).
 % 154.00/154.24  cnf(c_0_11, axiom, (cons(skaf44(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause101)).
 % 154.00/154.24  cnf(c_0_12, axiom, (ssItem(skaf44(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause47)).
 % 154.00/154.24  cnf(c_0_13, axiom, (singletonP(X2)|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause116)).
 % 154.00/154.24  cnf(c_0_14, axiom, (X4=X2|cons(X1,X2)!=cons(X3,X4)|~ssItem(X3)|~ssItem(X1)|~ssList(X4)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause171)).
 % 154.00/154.24  cnf(c_0_15, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause120)).
 % 154.00/154.24  cnf(c_0_16, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause96)).
 % 154.00/154.24  cnf(c_0_17, axiom, (cons(skaf83(X1),skaf82(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause109)).
 % 154.00/154.24  cnf(c_0_18, axiom, (ssItem(skaf83(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause12)).
 % 154.00/154.24  cnf(c_0_19, axiom, (ssList(skaf82(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause13)).
 % 154.00/154.24  cnf(c_0_20, axiom, (nil=X2|tl(app(X2,X1))=app(tl(X2),X1)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause133)).
 % 154.00/154.24  cnf(c_0_21, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
 % 154.00/154.24  cnf(c_0_22, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
 % 154.00/154.24  cnf(c_0_23, negated_conjecture, (memberP(sk4,sk5)|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_12)).
 % 154.00/154.24  cnf(c_0_24, axiom, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause157)).
 % 154.00/154.24  cnf(c_0_25, axiom, (memberP(app(X3,X1),X2)|~memberP(X1,X2)|~ssList(X1)|~ssList(X3)|~ssItem(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause141)).
 % 154.00/154.24  cnf(c_0_26, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause71)).
 % 154.00/154.24  cnf(c_0_27, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause145)).
 % 154.00/154.24  cnf(c_0_28, negated_conjecture, (nil=sk2|~neq(sk1,nil)|~segmentP(sk2,sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_14)).
 % 154.00/154.24  cnf(c_0_29, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause100)).
 % 154.00/154.24  cnf(c_0_30, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
 % 154.00/154.24  cnf(c_0_31, axiom, (segmentP(X1,X3)|~segmentP(X1,X2)|~segmentP(X2,X3)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause152)).
 % 154.00/154.24  cnf(c_0_32, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause173)).
 % 154.00/154.24  cnf(c_0_33, axiom, (segmentP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause57)).
 % 154.00/154.24  cnf(c_0_34, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause149)).
 % 154.00/154.24  cnf(c_0_35, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause132)).
 % 154.00/154.24  cnf(c_0_36, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause51)).
 % 154.00/154.24  cnf(c_0_37, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause60)).
 % 154.00/154.24  cnf(c_0_38, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause169)).
 % 154.00/154.24  cnf(c_0_39, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause53)).
 % 154.00/154.24  cnf(c_0_40, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause52)).
 % 154.00/154.24  cnf(c_0_41, negated_conjecture, (cons(sk5,nil)=sk1|sk2=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_0, c_0_1]), c_0_2])).
 % 154.00/154.24  cnf(c_0_42, negated_conjecture, (sk2=nil|ssItem(sk5)), inference(rw,[status(thm)],[c_0_3, c_0_1])).
 % 154.00/154.24  cnf(c_0_43, negated_conjecture, (sk2=nil|sk1!=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_41]), c_0_5])]), c_0_42])).
 % 154.00/154.24  cnf(c_0_44, plain, (~neq(X1,X1)|~ssList(X1)), inference(er,[status(thm)],[c_0_6])).
 % 154.00/154.24  cnf(c_0_45, negated_conjecture, (neq(nil,nil)|sk1!=nil), inference(spm,[status(thm)],[c_0_7, c_0_43])).
 % 154.00/154.24  cnf(c_0_46, negated_conjecture, (cons(sk5,nil)=sk1|sk1=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_2]), c_0_2])).
 % 154.00/154.24  cnf(c_0_47, negated_conjecture, (sk1!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_5])])).
 % 154.00/154.24  cnf(c_0_48, negated_conjecture, (sk1=nil|ssItem(sk5)), inference(rw,[status(thm)],[c_0_9, c_0_2])).
 % 154.00/154.24  cnf(c_0_49, plain, (X1!=nil|~singletonP(X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_5])])).
 % 154.00/154.24  cnf(c_0_50, negated_conjecture, (cons(sk5,nil)=sk1), inference(sr,[status(thm)],[c_0_46, c_0_47])).
 % 154.00/154.24  cnf(c_0_51, negated_conjecture, (ssItem(sk5)), inference(sr,[status(thm)],[c_0_48, c_0_47])).
 % 154.00/154.24  cnf(c_0_52, plain, (cons(X1,nil)!=X2|X2!=nil|~ssItem(X1)|~ssList(X2)), inference(spm,[status(thm)],[c_0_49, c_0_13])).
 % 154.00/154.24  cnf(c_0_53, negated_conjecture, (X1=nil|cons(X2,X1)!=sk1|~ssItem(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_50]), c_0_51]), c_0_5])])).
 % 154.00/154.24  cnf(c_0_54, negated_conjecture, (sk1!=X1|X1!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_50]), c_0_51])])).
 % 154.00/154.24  cnf(c_0_55, negated_conjecture, (app(sk1,X1)=cons(sk5,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_50]), c_0_51])])).
 % 154.00/154.24  cnf(c_0_56, negated_conjecture, (tl(sk1)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_50]), c_0_51]), c_0_5])])).
 % 154.00/154.24  cnf(c_0_57, plain, (skaf82(X1)=tl(X1)|nil=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19])])).
 % 154.00/154.24  cnf(c_0_58, negated_conjecture, (skaf82(X1)=nil|X1!=sk1|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_17]), c_0_18]), c_0_19])]), c_0_54])).
 % 154.00/154.24  cnf(c_0_59, negated_conjecture, (tl(cons(sk5,X1))=app(nil,X1)|~ssList(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_55]), c_0_56]), c_0_21])]), c_0_47])).
 % 154.00/154.24  cnf(c_0_60, negated_conjecture, (tl(X1)=nil|X1!=sk1|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_54])).
 % 154.00/154.24  cnf(c_0_61, negated_conjecture, (ssList(cons(sk5,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_55]), c_0_21])])).
 % 154.00/154.24  cnf(c_0_62, negated_conjecture, (app(nil,X1)=nil|cons(sk5,X1)!=sk1|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])).
 % 154.00/154.24  cnf(c_0_63, negated_conjecture, (sk1=nil|memberP(sk2,sk5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_2]), c_0_1])).
 % 154.00/154.24  cnf(c_0_64, negated_conjecture, (cons(sk5,app(nil,X1))=app(sk1,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_50]), c_0_51]), c_0_5])])).
 % 154.00/154.24  cnf(c_0_65, negated_conjecture, (cons(sk5,X1)!=sk1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_62]), c_0_5])]), c_0_26])).
 % 154.00/154.24  cnf(c_0_66, negated_conjecture, (memberP(sk2,sk5)), inference(sr,[status(thm)],[c_0_63, c_0_47])).
 % 154.00/154.24  cnf(c_0_67, negated_conjecture, (app(sk1,X1)=sk1|nil!=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_27]), c_0_50]), c_0_5])])).
 % 154.00/154.24  cnf(c_0_68, negated_conjecture, (sk2=nil|~segmentP(sk2,sk1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_21]), c_0_5])]), c_0_47])).
 % 154.00/154.24  cnf(c_0_69, negated_conjecture, (cons(sk5,sk2)!=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_51]), c_0_30])])).
 % 154.00/154.24  cnf(c_0_70, negated_conjecture, (cons(sk5,X1)=sk1|nil!=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_55, c_0_67])).
 % 154.00/154.24  cnf(c_0_71, negated_conjecture, (sk2=nil|~segmentP(X1,sk1)|~segmentP(sk2,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_31]), c_0_21]), c_0_30])])).
 % 154.00/154.24  cnf(c_0_72, negated_conjecture, (sk2!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_30])])).
 % 154.00/154.24  cnf(c_0_73, negated_conjecture, (app(app(X1,X2),X3)!=sk2|~segmentP(X2,sk1)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_32]), c_0_30])]), c_0_72])).
 % 154.00/154.24  cnf(c_0_74, negated_conjecture, (app(app(X1,sk1),X2)!=sk2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_33]), c_0_21])])).
 % 154.00/154.24  cnf(c_0_75, negated_conjecture, (app(X1,app(sk1,X2))!=sk2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_34]), c_0_21])])).
 % 154.00/154.24  cnf(c_0_76, negated_conjecture, (app(sk1,skaf45(X1,nil))=cons(sk5,X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_35]), c_0_36]), c_0_5])]), c_0_37])).
 % 154.00/154.24  cnf(c_0_77, negated_conjecture, (app(X1,cons(sk5,X2))!=sk2|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_36])])).
 % 154.00/154.24  cnf(c_0_78, negated_conjecture, (X1!=sk2|~memberP(X1,sk5)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_38]), c_0_39]), c_0_40]), c_0_51])])).
 % 154.00/154.24  cnf(c_0_79, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_66]), c_0_30])]), ['proof']).
 % 154.00/154.24  # SZS output end CNFRefutation
 % 154.00/154.24  # Training examples: 0 positive, 0 negative
 % 154.00/154.25
 % 154.00/154.25  # -------------------------------------------------
 % 154.00/154.25  # User time                : 150.559 s
 % 154.00/154.25  # System time              : 3.467 s
 % 154.00/154.25  # Total time               : 154.026 s
 % 154.00/154.25  # Maximum resident set size: 1860 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
