   --------------------------------------------------------------------------

                   TSTP Solution File: SWC198-1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC198-1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n146.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:01 EDT 2016

 % Result     : Unsatisfiable 0.06s
 % Output     : CNFRefutation 0.06s
 % Verified   :
 % Statistics : Number of clauses        :   39 (  53 expanded)
 %              Number of leaves         :   22 (  34 expanded)
 %              Depth                    :   10
 %              Number of atoms          :   88 ( 110 expanded)
 %              Number of equality atoms :   26 (  32 expanded)
 %              Maximal clause size      :    7 (   2 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 cnf(c_0_0,negated_conjecture,
     ( ssList(sk2) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

 cnf(c_0_1,negated_conjecture,
     ( sk2 = sk4 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

 cnf(c_0_2,axiom,
     ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2)
     | ~ duplicatefreeP(X5)
     | ~ ssList(X5) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause179)).

 cnf(c_0_3,axiom,
     ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
     | ~ memberP(X1,X2)
     | ~ ssItem(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause169)).

 cnf(c_0_4,negated_conjecture,
     ( memberP(sk4,sk6)
     | nil = sk3 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_16)).

 cnf(c_0_5,negated_conjecture,
     ( ssItem(sk6)
     | nil = sk3 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_11)).

 cnf(c_0_6,axiom,
     ( memberP(X2,X3)
     | X3 = X1
     | ~ memberP(cons(X1,X2),X3)
     | ~ ssList(X2)
     | ~ ssItem(X1)
     | ~ ssItem(X3) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause161)).

 cnf(c_0_7,negated_conjecture,
     ( cons(sk6,nil) = sk3
     | nil = sk4 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_12)).

 cnf(c_0_8,axiom,
     ( ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

 cnf(c_0_9,negated_conjecture,
     ( ssItem(sk6)
     | nil = sk4 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

 cnf(c_0_10,axiom,
     ( ~ memberP(nil,X1)
     | ~ ssItem(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause71)).

 cnf(c_0_11,negated_conjecture,
     ( memberP(sk1,sk5(X1))
     | ~ ssItem(X1) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

 cnf(c_0_12,negated_conjecture,
     ( sk1 = sk3 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

 cnf(c_0_13,axiom,
     ( ssList(skaf43(X1,X2)) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause52)).

 cnf(c_0_14,axiom,
     ( ssList(skaf42(X1,X2)) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause53)).

 cnf(c_0_15,negated_conjecture,
     ( ssItem(sk5(X1))
     | ~ ssItem(X1) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

 cnf(c_0_16,negated_conjecture,
     ( cons(sk6,nil) = sk3
     | nil = sk3 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_15)).

 cnf(c_0_17,negated_conjecture,
     ( ~ ssItem(X1)
     | X1 != sk5(X1) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

 cnf(c_0_18,axiom,
     ( duplicatefreeP(X1)
     | ssItem(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause72)).

 cnf(c_0_19,negated_conjecture,
     ( ssList(sk1) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

 cnf(c_0_20,axiom,
     ( ssList(app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

 cnf(c_0_21,axiom,
     ( ssItem(skac3) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause9)).

 cnf(c_0_22,negated_conjecture,
     ( ssList(sk4) ),
     inference(rw,[status(thm)],[[4]c_0_0,[5]c_0_1])).

 cnf(c_0_23,plain,
     ( ~ ssItem(X1)
     | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
     inference(er,[status(thm)],[[6]c_0_2])).

 cnf(c_0_24,negated_conjecture,
     ( app(skaf42(sk4,sk6),cons(sk6,skaf43(sk6,sk4))) = sk4
     | sk3 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[7]c_0_3,c_0_4]),c_0_22])]),[8]c_0_5])).

 cnf(c_0_25,negated_conjecture,
     ( sk4 = nil
     | sk6 = X1
     | ~ memberP(sk3,X1)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[9]c_0_6,c_0_7]),c_0_8])]),c_0_9]),[10]c_0_10])).

 cnf(c_0_26,negated_conjecture,
     ( memberP(sk3,sk5(X1))
     | ~ ssItem(X1) ),
     inference(rw,[status(thm)],[[11]c_0_11,[12]c_0_12])).

 cnf(c_0_27,negated_conjecture,
     ( sk3 = nil
     | ~ ssList(app(sk4,cons(sk6,X1)))
     | ~ ssList(X1)
     | ~ duplicatefreeP(app(sk4,cons(sk6,X1))) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[13]c_0_23,c_0_24]),c_0_13]),c_0_14])]),[14]c_0_5])).

 cnf(c_0_28,negated_conjecture,
     ( sk5(X1) = sk6
     | sk4 = nil
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_25,c_0_26]),[16]c_0_15])).

 cnf(c_0_29,negated_conjecture,
     ( sk3 = nil
     | ~ ssList(app(sk4,sk3))
     | ~ duplicatefreeP(app(sk4,sk3)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_27,c_0_16]),c_0_8])])).

 cnf(c_0_30,negated_conjecture,
     ( sk4 = nil ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_17,c_0_28])]),[19]c_0_9])).

 cnf(c_0_31,negated_conjecture,
     ( sk3 = nil
     | ssItem(X1)
     | ~ ssList(app(sk4,sk3)) ),
     inference(spm,[status(thm)],[[20]c_0_29,[21]c_0_18])).

 cnf(c_0_32,negated_conjecture,
     ( ssList(sk3) ),
     inference(rw,[status(thm)],[[22]c_0_19,[23]c_0_12])).

 cnf(c_0_33,negated_conjecture,
     ( sk3 = nil
     | memberP(nil,sk6) ),
     inference(rw,[status(thm)],[[24]c_0_4,[25]c_0_30])).

 cnf(c_0_34,negated_conjecture,
     ( sk3 = nil
     | ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_31,c_0_20]),c_0_22]),c_0_32])])).

 cnf(c_0_35,negated_conjecture,
     ( sk3 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_10,c_0_33]),[28]c_0_34])).

 cnf(c_0_36,negated_conjecture,
     ( memberP(nil,sk5(X1))
     | ~ ssItem(X1) ),
     inference(rw,[status(thm)],[[29]c_0_26,[30]c_0_35])).

 cnf(c_0_37,negated_conjecture,
     ( ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[31]c_0_10,c_0_36]),[32]c_0_15])).

 cnf(c_0_38,plain,
     ( $false ),
     inference(sr,[status(thm)],[[33]c_0_21,[34]c_0_37]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC198-1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.23  % Computer   : n146.star.cs.uiowa.edu
 % 0.02/0.23  % Model      : x86_64 x86_64
 % 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.23  % Memory     : 32218.75MB
 % 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.23  % CPULimit   : 300
 % 0.02/0.23  % DateTime   : Tue May 10 20:37:22 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.06/0.34  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
 % 0.06/0.34  # and selection function SelectCQArNTEqFirst.
 % 0.06/0.34  #
 % 0.06/0.34  # Preprocessing time       : 0.010 s
 % 0.06/0.34  # Presaturation interreduction done
 % 0.06/0.34
 % 0.06/0.34  # Proof found!
 % 0.06/0.34  # SZS status Unsatisfiable
 % 0.06/0.34  # SZS output start CNFRefutation
 % 0.06/0.34  cnf(c_0_0, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
 % 0.06/0.34  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
 % 0.06/0.34  cnf(c_0_2, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause179)).
 % 0.06/0.34  cnf(c_0_3, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause169)).
 % 0.06/0.34  cnf(c_0_4, negated_conjecture, (memberP(sk4,sk6)|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_16)).
 % 0.06/0.34  cnf(c_0_5, negated_conjecture, (ssItem(sk6)|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_11)).
 % 0.06/0.34  cnf(c_0_6, axiom, (memberP(X2,X3)|X3=X1|~memberP(cons(X1,X2),X3)|~ssList(X2)|~ssItem(X1)|~ssItem(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause161)).
 % 0.06/0.34  cnf(c_0_7, negated_conjecture, (cons(sk6,nil)=sk3|nil=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_12)).
 % 0.06/0.34  cnf(c_0_8, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
 % 0.06/0.34  cnf(c_0_9, negated_conjecture, (ssItem(sk6)|nil=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
 % 0.06/0.34  cnf(c_0_10, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause71)).
 % 0.06/0.34  cnf(c_0_11, negated_conjecture, (memberP(sk1,sk5(X1))|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
 % 0.06/0.34  cnf(c_0_12, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
 % 0.06/0.34  cnf(c_0_13, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause52)).
 % 0.06/0.34  cnf(c_0_14, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause53)).
 % 0.06/0.34  cnf(c_0_15, negated_conjecture, (ssItem(sk5(X1))|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
 % 0.06/0.34  cnf(c_0_16, negated_conjecture, (cons(sk6,nil)=sk3|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_15)).
 % 0.06/0.34  cnf(c_0_17, negated_conjecture, (~ssItem(X1)|X1!=sk5(X1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
 % 0.06/0.34  cnf(c_0_18, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause72)).
 % 0.06/0.34  cnf(c_0_19, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
 % 0.06/0.34  cnf(c_0_20, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
 % 0.06/0.34  cnf(c_0_21, axiom, (ssItem(skac3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause9)).
 % 0.06/0.34  cnf(c_0_22, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_0, c_0_1])).
 % 0.06/0.34  cnf(c_0_23, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_2])).
 % 0.06/0.34  cnf(c_0_24, negated_conjecture, (app(skaf42(sk4,sk6),cons(sk6,skaf43(sk6,sk4)))=sk4|sk3=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_3, c_0_4]), c_0_22])]), c_0_5])).
 % 0.06/0.34  cnf(c_0_25, negated_conjecture, (sk4=nil|sk6=X1|~memberP(sk3,X1)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])]), c_0_9]), c_0_10])).
 % 0.06/0.34  cnf(c_0_26, negated_conjecture, (memberP(sk3,sk5(X1))|~ssItem(X1)), inference(rw,[status(thm)],[c_0_11, c_0_12])).
 % 0.06/0.34  cnf(c_0_27, negated_conjecture, (sk3=nil|~ssList(app(sk4,cons(sk6,X1)))|~ssList(X1)|~duplicatefreeP(app(sk4,cons(sk6,X1)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_13]), c_0_14])]), c_0_5])).
 % 0.06/0.34  cnf(c_0_28, negated_conjecture, (sk5(X1)=sk6|sk4=nil|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_15])).
 % 0.06/0.34  cnf(c_0_29, negated_conjecture, (sk3=nil|~ssList(app(sk4,sk3))|~duplicatefreeP(app(sk4,sk3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_16]), c_0_8])])).
 % 0.06/0.34  cnf(c_0_30, negated_conjecture, (sk4=nil), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_28])]), c_0_9])).
 % 0.06/0.34  cnf(c_0_31, negated_conjecture, (sk3=nil|ssItem(X1)|~ssList(app(sk4,sk3))), inference(spm,[status(thm)],[c_0_29, c_0_18])).
 % 0.06/0.34  cnf(c_0_32, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_19, c_0_12])).
 % 0.06/0.34  cnf(c_0_33, negated_conjecture, (sk3=nil|memberP(nil,sk6)), inference(rw,[status(thm)],[c_0_4, c_0_30])).
 % 0.06/0.34  cnf(c_0_34, negated_conjecture, (sk3=nil|ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_20]), c_0_22]), c_0_32])])).
 % 0.06/0.34  cnf(c_0_35, negated_conjecture, (sk3=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_33]), c_0_34])).
 % 0.06/0.34  cnf(c_0_36, negated_conjecture, (memberP(nil,sk5(X1))|~ssItem(X1)), inference(rw,[status(thm)],[c_0_26, c_0_35])).
 % 0.06/0.34  cnf(c_0_37, negated_conjecture, (~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_36]), c_0_15])).
 % 0.06/0.34  cnf(c_0_38, plain, ($false), inference(sr,[status(thm)],[c_0_21, c_0_37]), ['proof']).
 % 0.06/0.34  # SZS output end CNFRefutation
 % 0.06/0.34  # Training examples: 0 positive, 0 negative
 % 0.06/0.34
 % 0.06/0.34  # -------------------------------------------------
 % 0.06/0.34  # User time                : 0.107 s
 % 0.06/0.34  # System time              : 0.004 s
 % 0.06/0.34  # Total time               : 0.111 s
 % 0.06/0.34  # Maximum resident set size: 1860 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
