   --------------------------------------------------------------------------

                   TSTP Solution File: SWC149-1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC149-1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n112.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:16:50 EDT 2016

 % Result     : Unsatisfiable 0.02s
 % Output     : CNFRefutation 0.02s
 % Verified   :
 % Statistics : Number of clauses        :   28 (  39 expanded)
 %              Number of leaves         :   14 (  25 expanded)
 %              Depth                    :   11
 %              Number of atoms          :   73 (  92 expanded)
 %              Number of equality atoms :   21 (  26 expanded)
 %              Maximal clause size      :    5 (   3 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 cnf(c_0_0,negated_conjecture,
     ( ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk4 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

 cnf(c_0_1,negated_conjecture,
     ( sk2 = sk4 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

 cnf(c_0_2,axiom,
     ( app(cons(X1,nil),X2) = cons(X1,X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause120)).

 cnf(c_0_3,axiom,
     ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause157)).

 cnf(c_0_4,axiom,
     ( cons(skaf83(X1),skaf82(X1)) = X1
     | nil = X1
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause109)).

 cnf(c_0_5,axiom,
     ( ssItem(skaf83(X1)) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause12)).

 cnf(c_0_6,axiom,
     ( ssList(skaf82(X1)) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause13)).

 cnf(c_0_7,negated_conjecture,
     ( ssList(sk2) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

 cnf(c_0_8,axiom,
     ( singletonP(X2)
     | cons(X1,nil) != X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause116)).

 cnf(c_0_9,axiom,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause86)).

 cnf(c_0_10,axiom,
     ( ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

 cnf(c_0_11,negated_conjecture,
     ( ~ singletonP(sk2) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

 cnf(c_0_12,axiom,
     ( X1 != X2
     | ~ neq(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause115)).

 cnf(c_0_13,negated_conjecture,
     ( neq(sk2,nil) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

 cnf(c_0_14,negated_conjecture,
     ( app(app(cons(X1,nil),cons(X2,nil)),X3) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(X3) ),
     inference(rw,[status(thm)],[[4]c_0_0,[5]c_0_1])).

 cnf(c_0_15,negated_conjecture,
     ( app(cons(X1,cons(X2,nil)),X3) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(cons(X2,nil))
     | ~ ssList(X3) ),
     inference(spm,[status(thm)],[[6]c_0_14,[7]c_0_2])).

 cnf(c_0_16,negated_conjecture,
     ( cons(X1,app(cons(X2,nil),X3)) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(cons(X2,nil))
     | ~ ssList(X3) ),
     inference(spm,[status(thm)],[[8]c_0_15,[9]c_0_3])).

 cnf(c_0_17,negated_conjecture,
     ( cons(X1,cons(X2,X3)) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(cons(X2,nil))
     | ~ ssList(X3) ),
     inference(spm,[status(thm)],[[10]c_0_16,[11]c_0_2])).

 cnf(c_0_18,negated_conjecture,
     ( nil = X1
     | cons(X2,X1) != sk2
     | ~ ssItem(X2)
     | ~ ssList(cons(skaf83(X1),nil))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[12]c_0_17,c_0_4]),c_0_5]),c_0_6])])).

 cnf(c_0_19,negated_conjecture,
     ( skaf82(sk2) = nil
     | sk2 = nil
     | ~ ssList(cons(skaf83(skaf82(sk2)),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[13]c_0_18,c_0_4]),c_0_5]),c_0_6])])]),c_0_7])])).

 cnf(c_0_20,plain,
     ( singletonP(cons(X1,nil))
     | ~ ssItem(X1)
     | ~ ssList(cons(X1,nil)) ),
     inference(er,[status(thm)],[[14]c_0_8])).

 cnf(c_0_21,negated_conjecture,
     ( skaf82(sk2) = nil
     | sk2 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_19,c_0_9]),c_0_5]),c_0_10])])).

 cnf(c_0_22,plain,
     ( singletonP(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[16]c_0_20,c_0_9]),c_0_10])])).

 cnf(c_0_23,negated_conjecture,
     ( cons(skaf83(sk2),nil) = sk2
     | sk2 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_4,c_0_21]),c_0_7])])).

 cnf(c_0_24,negated_conjecture,
     ( sk2 = nil ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_22,c_0_23]),c_0_5])]),[19]c_0_11])).

 cnf(c_0_25,plain,
     ( ~ neq(X1,X1)
     | ~ ssList(X1) ),
     inference(er,[status(thm)],[[20]c_0_12])).

 cnf(c_0_26,negated_conjecture,
     ( neq(nil,nil) ),
     inference(rw,[status(thm)],[[21]c_0_13,[22]c_0_24])).

 cnf(c_0_27,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[23]c_0_25,c_0_26]),c_0_10])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC149-1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.23  % Computer   : n112.star.cs.uiowa.edu
 % 0.02/0.23  % Model      : x86_64 x86_64
 % 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.23  % Memory     : 32218.75MB
 % 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.23  % CPULimit   : 300
 % 0.02/0.23  % DateTime   : Tue May 10 20:36:08 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.02/0.26  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
 % 0.02/0.26  # and selection function SelectCQArNTEqFirst.
 % 0.02/0.26  #
 % 0.02/0.26  # Preprocessing time       : 0.010 s
 % 0.02/0.26  # Presaturation interreduction done
 % 0.02/0.26
 % 0.02/0.26  # Proof found!
 % 0.02/0.26  # SZS status Unsatisfiable
 % 0.02/0.26  # SZS output start CNFRefutation
 % 0.02/0.26  cnf(c_0_0, negated_conjecture, (~ssItem(X1)|~ssItem(X2)|~ssList(X3)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
 % 0.02/0.26  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
 % 0.02/0.26  cnf(c_0_2, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause120)).
 % 0.02/0.26  cnf(c_0_3, axiom, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause157)).
 % 0.02/0.26  cnf(c_0_4, axiom, (cons(skaf83(X1),skaf82(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause109)).
 % 0.02/0.26  cnf(c_0_5, axiom, (ssItem(skaf83(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause12)).
 % 0.02/0.26  cnf(c_0_6, axiom, (ssList(skaf82(X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause13)).
 % 0.02/0.26  cnf(c_0_7, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
 % 0.02/0.26  cnf(c_0_8, axiom, (singletonP(X2)|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause116)).
 % 0.02/0.26  cnf(c_0_9, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause86)).
 % 0.02/0.26  cnf(c_0_10, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
 % 0.02/0.26  cnf(c_0_11, negated_conjecture, (~singletonP(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
 % 0.02/0.26  cnf(c_0_12, axiom, (X1!=X2|~neq(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause115)).
 % 0.02/0.26  cnf(c_0_13, negated_conjecture, (neq(sk2,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
 % 0.02/0.26  cnf(c_0_14, negated_conjecture, (app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(X3)), inference(rw,[status(thm)],[c_0_0, c_0_1])).
 % 0.02/0.26  cnf(c_0_15, negated_conjecture, (app(cons(X1,cons(X2,nil)),X3)!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(cons(X2,nil))|~ssList(X3)), inference(spm,[status(thm)],[c_0_14, c_0_2])).
 % 0.02/0.26  cnf(c_0_16, negated_conjecture, (cons(X1,app(cons(X2,nil),X3))!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(cons(X2,nil))|~ssList(X3)), inference(spm,[status(thm)],[c_0_15, c_0_3])).
 % 0.02/0.26  cnf(c_0_17, negated_conjecture, (cons(X1,cons(X2,X3))!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(cons(X2,nil))|~ssList(X3)), inference(spm,[status(thm)],[c_0_16, c_0_2])).
 % 0.02/0.26  cnf(c_0_18, negated_conjecture, (nil=X1|cons(X2,X1)!=sk2|~ssItem(X2)|~ssList(cons(skaf83(X1),nil))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_4]), c_0_5]), c_0_6])])).
 % 0.02/0.26  cnf(c_0_19, negated_conjecture, (skaf82(sk2)=nil|sk2=nil|~ssList(cons(skaf83(skaf82(sk2)),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_4]), c_0_5]), c_0_6])])]), c_0_7])])).
 % 0.02/0.26  cnf(c_0_20, plain, (singletonP(cons(X1,nil))|~ssItem(X1)|~ssList(cons(X1,nil))), inference(er,[status(thm)],[c_0_8])).
 % 0.02/0.26  cnf(c_0_21, negated_conjecture, (skaf82(sk2)=nil|sk2=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_9]), c_0_5]), c_0_10])])).
 % 0.02/0.26  cnf(c_0_22, plain, (singletonP(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_9]), c_0_10])])).
 % 0.02/0.26  cnf(c_0_23, negated_conjecture, (cons(skaf83(sk2),nil)=sk2|sk2=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_21]), c_0_7])])).
 % 0.02/0.26  cnf(c_0_24, negated_conjecture, (sk2=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_5])]), c_0_11])).
 % 0.02/0.26  cnf(c_0_25, plain, (~neq(X1,X1)|~ssList(X1)), inference(er,[status(thm)],[c_0_12])).
 % 0.02/0.26  cnf(c_0_26, negated_conjecture, (neq(nil,nil)), inference(rw,[status(thm)],[c_0_13, c_0_24])).
 % 0.02/0.26  cnf(c_0_27, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_10])]), ['proof']).
 % 0.02/0.26  # SZS output end CNFRefutation
 % 0.02/0.26  # Training examples: 0 positive, 0 negative
 % 0.02/0.26
 % 0.02/0.26  # -------------------------------------------------
 % 0.02/0.26  # User time                : 0.035 s
 % 0.02/0.26  # System time              : 0.003 s
 % 0.02/0.26  # Total time               : 0.038 s
 % 0.02/0.26  # Maximum resident set size: 1860 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
