   --------------------------------------------------------------------------

                   TSTP Solution File: SWC246+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC246+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n152.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:12 EDT 2016

 % Result     : Theorem 0.06s
 % Output     : CNFRefutation 0.06s
 % Verified   :
 % Statistics : -

 % Comments   :
 %------------------------------------------------------------------------------
 %----WARNING: Could not form TPTP format derivation
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC246+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.22  % Computer   : n152.star.cs.uiowa.edu
 % 0.03/0.22  % Model      : x86_64 x86_64
 % 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.22  % Memory     : 32218.75MB
 % 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.22  % CPULimit   : 300
 % 0.03/0.22  % DateTime   : Tue May 10 20:38:08 CDT 2016
 % 0.03/0.22  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.06/0.30  # AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
 % 0.06/0.30  # and selection function SelectMaxLComplexAvoidPosPred.
 % 0.06/0.30  #
 % 0.06/0.30  # Preprocessing time       : 0.014 s
 % 0.06/0.30  # Presaturation interreduction done
 % 0.06/0.30
 % 0.06/0.30  # Proof found!
 % 0.06/0.30  # SZS status Theorem
 % 0.06/0.30  # SZS output start CNFRefutation
 % 0.06/0.30  fof(c_0_0, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax38)).
 % 0.06/0.30  fof(c_0_1, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((((~(ssList(X4))|X2!=X4)|X1!=X3)|nil=X1)|?[X5]:(ssItem(X5)&?[X6]:(ssList(X6)&?[X7]:((ssList(X7)&app(app(X6,cons(X5,nil)),X7)=X1)&![X8]:((((~(ssItem(X8))|~(memberP(X6,X8)))|~(memberP(X7,X8)))|~(lt(X5,X8)))|leq(X5,X8))))))|![X9]:(ssList(X9)=>![X10]:((((~(ssList(X10))|app(app(X9,X3),X10)!=X4)|~(equalelemsP(X3)))|?[X11]:(ssItem(X11)&?[X12]:((ssList(X12)&app(X12,cons(X11,nil))=X9)&?[X13]:(ssList(X13)&app(cons(X11,nil),X13)=X3))))|?[X14]:(ssItem(X14)&?[X15]:((ssList(X15)&app(cons(X14,nil),X15)=X10)&?[X16]:(ssList(X16)&app(X16,cons(X14,nil))=X3))))))|(nil!=X4&nil=X3)))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
 % 0.06/0.30  fof(c_0_2, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.06/0.30  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 0.06/0.30  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax81)).
 % 0.06/0.30  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.06/0.30  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax20)).
 % 0.06/0.30  fof(c_0_7, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_0])).
 % 0.06/0.30  fof(c_0_8, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((((~ssList(X4)|X2!=X4)|X1!=X3)|nil=X1)|?[X5]:(ssItem(X5)&?[X6]:(ssList(X6)&?[X7]:((ssList(X7)&app(app(X6,cons(X5,nil)),X7)=X1)&![X8]:((((~ssItem(X8)|~memberP(X6,X8))|~memberP(X7,X8))|~lt(X5,X8))|leq(X5,X8))))))|![X9]:(ssList(X9)=>![X10]:((((~ssList(X10)|app(app(X9,X3),X10)!=X4)|~equalelemsP(X3))|?[X11]:(ssItem(X11)&?[X12]:((ssList(X12)&app(X12,cons(X11,nil))=X9)&?[X13]:(ssList(X13)&app(cons(X11,nil),X13)=X3))))|?[X14]:(ssItem(X14)&?[X15]:((ssList(X15)&app(cons(X14,nil),X15)=X10)&?[X16]:(ssList(X16)&app(X16,cons(X14,nil))=X3))))))|(nil!=X4&nil=X3))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_1])])).
 % 0.06/0.30  fof(c_0_9, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
 % 0.06/0.30  fof(c_0_10, negated_conjecture,
 % (![X21]:![X22]:![X23]:![X27]:![X28]:![X29]:![X30]:![X31]:![X32]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&((((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&nil!=esk48_0)&(((((((ssItem(esk52_3(X21,X22,X23))|(~ssList(X23)|app(app(X22,cons(X21,nil)),X23)!=esk48_0))|~ssList(X22))|~ssItem(X21))&(((memberP(X22,esk52_3(X21,X22,X23))|(~ssList(X23)|app(app(X22,cons(X21,nil)),X23)!=esk48_0))|~ssList(X22))|~ssItem(X21)))&(((memberP(X23,esk52_3(X21,X22,X23))|(~ssList(X23)|app(app(X22,cons(X21,nil)),X23)!=esk48_0))|~ssList(X22))|~ssItem(X21)))&(((lt(X21,esk52_3(X21,X22,X23))|(~ssList(X23)|app(app(X22,cons(X21,nil)),X23)!=esk48_0))|~ssList(X22))|~ssItem(X21)))&(((~leq(X21,esk52_3(X21,X22,X23))|(~ssList(X23)|app(app(X22,cons(X21,nil)),X23)!=esk48_0))|~ssList(X22))|~ssItem(X21))))&(ssList(esk53_0)&((((ssList(esk54_0)&app(app(esk53_0,esk50_0),esk54_0)=esk51_0)&equalelemsP(esk50_0))&(~ssItem(X27)|((~ssList(X28)|app(X28,cons(X27,nil))!=esk53_0)|(~ssList(X29)|app(cons(X27,nil),X29)!=esk50_0))))&(~ssItem(X30)|((~ss
 % List(X31)|app(cons(X30,nil),X31)!=esk54_0)|(~ssList(X32)|app(X32,cons(X30,nil))!=esk50_0))))))&(nil=esk51_0|nil!=esk50_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])])).
 % 0.06/0.30  cnf(c_0_11, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.06/0.30  cnf(c_0_12, negated_conjecture, (memberP(X2,esk52_3(X1,X2,X3))|~ssItem(X1)|~ssList(X2)|app(app(X2,cons(X1,nil)),X3)!=esk48_0|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.06/0.30  cnf(c_0_13, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.06/0.30  cnf(c_0_14, negated_conjecture, (app(app(nil,cons(X1,nil)),X2)!=esk48_0|~ssList(X2)|~ssItem(esk52_3(X1,nil,X2))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
 % 0.06/0.30  cnf(c_0_15, negated_conjecture, (ssItem(esk52_3(X1,X2,X3))|~ssItem(X1)|~ssList(X2)|app(app(X2,cons(X1,nil)),X3)!=esk48_0|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.06/0.30  fof(c_0_16, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
 % 0.06/0.30  cnf(c_0_17, negated_conjecture, (app(app(nil,cons(X1,nil)),X2)!=esk48_0|~ssList(X2)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_13])])).
 % 0.06/0.30  cnf(c_0_18, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
 % 0.06/0.30  fof(c_0_19, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 0.06/0.30  cnf(c_0_20, negated_conjecture, (app(cons(X1,nil),X2)!=esk48_0|~ssList(cons(X1,nil))|~ssList(X2)|~ssItem(X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
 % 0.06/0.30  cnf(c_0_21, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.30  fof(c_0_22, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
 % 0.06/0.30  cnf(c_0_23, negated_conjecture, (cons(X1,X2)!=esk48_0|~ssList(cons(X1,nil))|~ssList(X2)|~ssItem(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
 % 0.06/0.30  cnf(c_0_24, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 0.06/0.30  fof(c_0_25, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
 % 0.06/0.30  cnf(c_0_26, negated_conjecture, (cons(X1,X2)!=esk48_0|~ssList(X2)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_13])])).
 % 0.06/0.30  cnf(c_0_27, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
 % 0.06/0.30  cnf(c_0_28, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.06/0.30  cnf(c_0_29, negated_conjecture, (nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.06/0.30  cnf(c_0_30, negated_conjecture, (~ssList(esk44_1(esk48_0))|~ssItem(esk45_1(esk48_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27])]), c_0_28])]), c_0_29])).
 % 0.06/0.30  cnf(c_0_31, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
 % 0.06/0.30  cnf(c_0_32, negated_conjecture, (~ssList(esk44_1(esk48_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_28])]), c_0_29])).
 % 0.06/0.30  cnf(c_0_33, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
 % 0.06/0.30  cnf(c_0_34, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_28])]), c_0_29]), ['proof']).
 % 0.06/0.30  # SZS output end CNFRefutation
 % 0.06/0.30  # Training examples: 0 positive, 0 negative
 % 0.06/0.30
 % 0.06/0.30  # -------------------------------------------------
 % 0.06/0.30  # User time                : 0.070 s
 % 0.06/0.30  # System time              : 0.004 s
 % 0.06/0.30  # Total time               : 0.074 s
 % 0.06/0.30  # Maximum resident set size: 1892 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
