   --------------------------------------------------------------------------

                   TSTP Solution File: SWC219+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC219+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n103.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:06 EDT 2016

 % Result     : Theorem 0.06s
 % Output     : CNFRefutation 0.06s
 % Verified   :
 % Statistics : Number of formulae       :   41 ( 158 expanded)
 %              Number of clauses        :   24 (  55 expanded)
 %              Number of leaves         :    8 (  40 expanded)
 %              Depth                    :    7
 %              Number of atoms          :  196 (3045 expanded)
 %              Number of equality atoms :   66 (1065 expanded)
 %              Maximal formula depth    :   29 (   5 average)
 %              Maximal clause size      :   54 (   2 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | nil = X1
                     | ? [X5] :
                         ( ssItem(X5)
                         & ? [X6] :
                             ( ssList(X6)
                             & ? [X7] :
                                 ( ssList(X7)
                                 & app(app(X6,cons(X5,nil)),X7) = X1
                                 & ! [X8] :
                                     ( ssItem(X8)
                                    => ( ~ memberP(X6,X8)
                                       | ~ memberP(X7,X8)
                                       | ~ leq(X5,X8)
                                       | leq(X8,X5) ) ) ) ) )
                     | ( ! [X9] :
                           ( ssItem(X9)
                          => ( cons(X9,nil) != X3
                             | ~ memberP(X4,X9)
                             | ? [X10] :
                                 ( ssItem(X10)
                                 & X9 != X10
                                 & memberP(X4,X10)
                                 & leq(X10,X9) ) ) )
                       & ( nil != X4
                         | nil != X3 ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ~ memberP(nil,X1) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax38)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(nil,X1) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssItem(X3)
              => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

 fof(c_0_6,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(X1,nil) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

 fof(c_0_8,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | nil = X1
                       | ? [X5] :
                           ( ssItem(X5)
                           & ? [X6] :
                               ( ssList(X6)
                               & ? [X7] :
                                   ( ssList(X7)
                                   & app(app(X6,cons(X5,nil)),X7) = X1
                                   & ! [X8] :
                                       ( ssItem(X8)
                                      => ( ~ memberP(X6,X8)
                                         | ~ memberP(X7,X8)
                                         | ~ leq(X5,X8)
                                         | leq(X8,X5) ) ) ) ) )
                       | ( ! [X9] :
                             ( ssItem(X9)
                            => ( cons(X9,nil) != X3
                               | ~ memberP(X4,X9)
                               | ? [X10] :
                                   ( ssItem(X10)
                                   & X9 != X10
                                   & memberP(X4,X10)
                                   & leq(X10,X9) ) ) )
                         & ( nil != X4
                           | nil != X3 ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_9,plain,(
     ! [X1] :
       ( ssItem(X1)
      => ~ memberP(nil,X1) ) ),
     inference(fof_simplification,[status(thm)],[[4]c_0_1])).

 fof(c_0_10,negated_conjecture,(
     ! [X15,X16,X17,X20] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & nil != esk48_0
       & ( ssItem(esk52_3(X15,X16,X17))
         | ~ ssList(X17)
         | app(app(X16,cons(X15,nil)),X17) != esk48_0
         | ~ ssList(X16)
         | ~ ssItem(X15) )
       & ( memberP(X16,esk52_3(X15,X16,X17))
         | ~ ssList(X17)
         | app(app(X16,cons(X15,nil)),X17) != esk48_0
         | ~ ssList(X16)
         | ~ ssItem(X15) )
       & ( memberP(X17,esk52_3(X15,X16,X17))
         | ~ ssList(X17)
         | app(app(X16,cons(X15,nil)),X17) != esk48_0
         | ~ ssList(X16)
         | ~ ssItem(X15) )
       & ( leq(X15,esk52_3(X15,X16,X17))
         | ~ ssList(X17)
         | app(app(X16,cons(X15,nil)),X17) != esk48_0
         | ~ ssList(X16)
         | ~ ssItem(X15) )
       & ( ~ leq(esk52_3(X15,X16,X17),X15)
         | ~ ssList(X17)
         | app(app(X16,cons(X15,nil)),X17) != esk48_0
         | ~ ssList(X16)
         | ~ ssItem(X15) )
       & ( nil = esk51_0
         | ssItem(esk53_0) )
       & ( nil = esk50_0
         | ssItem(esk53_0) )
       & ( nil = esk51_0
         | cons(esk53_0,nil) = esk50_0 )
       & ( nil = esk50_0
         | cons(esk53_0,nil) = esk50_0 )
       & ( nil = esk51_0
         | memberP(esk51_0,esk53_0) )
       & ( nil = esk50_0
         | memberP(esk51_0,esk53_0) )
       & ( nil = esk51_0
         | ~ ssItem(X20)
         | esk53_0 = X20
         | ~ memberP(esk51_0,X20)
         | ~ leq(X20,esk53_0) )
       & ( nil = esk50_0
         | ~ ssItem(X20)
         | esk53_0 = X20
         | ~ memberP(esk51_0,X20)
         | ~ leq(X20,esk53_0) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).

 fof(c_0_11,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | ~ memberP(nil,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).

 fof(c_0_12,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

 fof(c_0_13,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(nil,X2) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

 fof(c_0_14,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

 cnf(c_0_15,negated_conjecture,
     ( cons(esk53_0,nil) = esk50_0
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_10])).

 cnf(c_0_16,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_10])).

 cnf(c_0_17,negated_conjecture,
     ( nil != esk48_0 ),
     inference(split_conjunct,[status(thm)],[[7]c_0_10])).

 cnf(c_0_18,negated_conjecture,
     ( ssItem(esk53_0)
     | nil = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[8]c_0_10])).

 fof(c_0_19,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssItem(X6)
       | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

 cnf(c_0_20,plain,
     ( ~ memberP(nil,X1)
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[9]c_0_11])).

 cnf(c_0_21,negated_conjecture,
     ( memberP(X3,esk52_3(X1,X2,X3))
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | app(app(X2,cons(X1,nil)),X3) != esk48_0
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[10]c_0_10])).

 cnf(c_0_22,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_6])).

 cnf(c_0_23,plain,
     ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_12])).

 cnf(c_0_24,plain,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[13]c_0_13])).

 cnf(c_0_25,plain,
     ( cons(X1,X2) = app(cons(X1,nil),X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_14])).

 cnf(c_0_26,negated_conjecture,
     ( cons(esk53_0,nil) = esk48_0 ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[15]c_0_15,c_0_16]),c_0_16]),[16]c_0_17])).

 cnf(c_0_27,negated_conjecture,
     ( ssItem(esk53_0) ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[[17]c_0_18,c_0_16]),[18]c_0_17])).

 cnf(c_0_28,plain,
     ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_19])).

 fof(c_0_29,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(X2,nil) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

 cnf(c_0_30,negated_conjecture,
     ( app(app(X1,cons(X2,nil)),nil) != esk48_0
     | ~ ssList(X1)
     | ~ ssItem(esk52_3(X2,X1,nil))
     | ~ ssItem(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[20]c_0_20,c_0_21]),c_0_22])])).

 cnf(c_0_31,negated_conjecture,
     ( ssItem(esk52_3(X1,X2,X3))
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | app(app(X2,cons(X1,nil)),X3) != esk48_0
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[21]c_0_10])).

 cnf(c_0_32,plain,
     ( app(nil,app(X1,X2)) = app(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[22]c_0_23,c_0_24]),c_0_22])])).

 cnf(c_0_33,negated_conjecture,
     ( app(esk48_0,X1) = cons(esk53_0,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[23]c_0_25,c_0_26]),c_0_27])])).

 cnf(c_0_34,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[24]c_0_10])).

 cnf(c_0_35,negated_conjecture,
     ( cons(esk53_0,app(nil,X1)) = app(esk48_0,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_28,c_0_26]),c_0_22]),c_0_27])])).

 cnf(c_0_36,plain,
     ( app(X1,nil) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[26]c_0_29])).

 cnf(c_0_37,negated_conjecture,
     ( app(app(X1,cons(X2,nil)),nil) != esk48_0
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_30,c_0_31]),c_0_22])])).

 cnf(c_0_38,negated_conjecture,
     ( app(nil,cons(esk53_0,X1)) = cons(esk53_0,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[28]c_0_32,c_0_33]),c_0_34])])).

 cnf(c_0_39,negated_conjecture,
     ( app(esk48_0,nil) = esk48_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[29]c_0_35,c_0_36]),c_0_26]),c_0_22])])).

 cnf(c_0_40,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[30]c_0_37,c_0_38]),c_0_26]),c_0_39]),c_0_22]),c_0_27])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC219+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.23  % Computer   : n103.star.cs.uiowa.edu
 % 0.03/0.23  % Model      : x86_64 x86_64
 % 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.23  % Memory     : 32218.75MB
 % 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.23  % CPULimit   : 300
 % 0.03/0.23  % DateTime   : Tue May 10 20:37:23 CDT 2016
 % 0.03/0.23  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.06/0.66  # AutoSched0-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
 % 0.06/0.66  # and selection function SelectMaxLComplexAvoidPosPred.
 % 0.06/0.66  #
 % 0.06/0.66  # Preprocessing time       : 0.014 s
 % 0.06/0.66  # Presaturation interreduction done
 % 0.06/0.66
 % 0.06/0.66  # Proof found!
 % 0.06/0.66  # SZS status Theorem
 % 0.06/0.66  # SZS output start CNFRefutation
 % 0.06/0.66  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|nil=X1)|?[X5]:(ssItem(X5)&?[X6]:(ssList(X6)&?[X7]:((ssList(X7)&app(app(X6,cons(X5,nil)),X7)=X1)&![X8]:(ssItem(X8)=>(((~(memberP(X6,X8))|~(memberP(X7,X8)))|~(leq(X5,X8)))|leq(X8,X5)))))))|(![X9]:(ssItem(X9)=>((cons(X9,nil)!=X3|~(memberP(X4,X9)))|?[X10]:(((ssItem(X10)&X9!=X10)&memberP(X4,X10))&leq(X10,X9))))&(nil!=X4|nil!=X3)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.06/0.66  fof(c_0_1, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax38)).
 % 0.06/0.66  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
 % 0.06/0.66  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 0.06/0.66  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
 % 0.06/0.66  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
 % 0.06/0.66  fof(c_0_6, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.06/0.66  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
 % 0.06/0.66  fof(c_0_8, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|nil=X1)|?[X5]:(ssItem(X5)&?[X6]:(ssList(X6)&?[X7]:((ssList(X7)&app(app(X6,cons(X5,nil)),X7)=X1)&![X8]:(ssItem(X8)=>(((~memberP(X6,X8)|~memberP(X7,X8))|~leq(X5,X8))|leq(X8,X5)))))))|(![X9]:(ssItem(X9)=>((cons(X9,nil)!=X3|~memberP(X4,X9))|?[X10]:(((ssItem(X10)&X9!=X10)&memberP(X4,X10))&leq(X10,X9))))&(nil!=X4|nil!=X3))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.06/0.66  fof(c_0_9, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_1])).
 % 0.06/0.66  fof(c_0_10, negated_conjecture,
 % (![X15]:![X16]:![X17]:![X20]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&nil!=esk48_0)&((((ssItem(esk52_3(X15,X16,X17))|(~ssList(X17)|app(app(X16,cons(X15,nil)),X17)!=esk48_0))|~ssList(X16))|~ssItem(X15))&((((((memberP(X16,esk52_3(X15,X16,X17))|(~ssList(X17)|app(app(X16,cons(X15,nil)),X17)!=esk48_0))|~ssList(X16))|~ssItem(X15))&(((memberP(X17,esk52_3(X15,X16,X17))|(~ssList(X17)|app(app(X16,cons(X15,nil)),X17)!=esk48_0))|~ssList(X16))|~ssItem(X15)))&(((leq(X15,esk52_3(X15,X16,X17))|(~ssList(X17)|app(app(X16,cons(X15,nil)),X17)!=esk48_0))|~ssList(X16))|~ssItem(X15)))&(((~leq(esk52_3(X15,X16,X17),X15)|(~ssList(X17)|app(app(X16,cons(X15,nil)),X17)!=esk48_0))|~ssList(X16))|~ssItem(X15)))))&(((nil=esk51_0|ssItem(esk53_0))&(nil=esk50_0|ssItem(esk53_0)))&((((nil=esk51_0|cons(esk53_0,nil)=esk50_0)&(nil=esk50_0|cons(esk53_0,nil)=esk50_0))&((nil=esk51_0|memberP(esk51_0,esk53_0))&(nil=esk50_0|memberP(esk51_0,esk53_0))))&((nil=esk51_0|(((~ssItem(X20)|esk53_0=X20
 % )|~memberP(esk51_0,X20))|~leq(X20,esk53_0)))&(nil=esk50_0|(((~ssItem(X20)|esk53_0=X20)|~memberP(esk51_0,X20))|~leq(X20,esk53_0)))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])])).
 % 0.06/0.66  fof(c_0_11, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])).
 % 0.06/0.66  fof(c_0_12, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
 % 0.06/0.66  fof(c_0_13, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
 % 0.06/0.66  fof(c_0_14, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 0.06/0.66  cnf(c_0_15, negated_conjecture, (cons(esk53_0,nil)=esk50_0|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.06/0.66  cnf(c_0_16, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.06/0.66  cnf(c_0_17, negated_conjecture, (nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.06/0.66  cnf(c_0_18, negated_conjecture, (ssItem(esk53_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.06/0.66  fof(c_0_19, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
 % 0.06/0.66  cnf(c_0_20, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
 % 0.06/0.66  cnf(c_0_21, negated_conjecture, (memberP(X3,esk52_3(X1,X2,X3))|~ssItem(X1)|~ssList(X2)|app(app(X2,cons(X1,nil)),X3)!=esk48_0|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.06/0.66  cnf(c_0_22, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.06/0.66  cnf(c_0_23, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
 % 0.06/0.66  cnf(c_0_24, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.06/0.66  cnf(c_0_25, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
 % 0.06/0.66  cnf(c_0_26, negated_conjecture, (cons(esk53_0,nil)=esk48_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_16]), c_0_16]), c_0_17])).
 % 0.06/0.66  cnf(c_0_27, negated_conjecture, (ssItem(esk53_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_16]), c_0_17])).
 % 0.06/0.66  cnf(c_0_28, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
 % 0.06/0.66  fof(c_0_29, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
 % 0.06/0.66  cnf(c_0_30, negated_conjecture, (app(app(X1,cons(X2,nil)),nil)!=esk48_0|~ssList(X1)|~ssItem(esk52_3(X2,X1,nil))|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).
 % 0.06/0.66  cnf(c_0_31, negated_conjecture, (ssItem(esk52_3(X1,X2,X3))|~ssItem(X1)|~ssList(X2)|app(app(X2,cons(X1,nil)),X3)!=esk48_0|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.06/0.66  cnf(c_0_32, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_22])])).
 % 0.06/0.66  cnf(c_0_33, negated_conjecture, (app(esk48_0,X1)=cons(esk53_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
 % 0.06/0.66  cnf(c_0_34, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.06/0.66  cnf(c_0_35, negated_conjecture, (cons(esk53_0,app(nil,X1))=app(esk48_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_26]), c_0_22]), c_0_27])])).
 % 0.06/0.66  cnf(c_0_36, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
 % 0.06/0.66  cnf(c_0_37, negated_conjecture, (app(app(X1,cons(X2,nil)),nil)!=esk48_0|~ssList(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_22])])).
 % 0.06/0.66  cnf(c_0_38, negated_conjecture, (app(nil,cons(esk53_0,X1))=cons(esk53_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])])).
 % 0.06/0.66  cnf(c_0_39, negated_conjecture, (app(esk48_0,nil)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_26]), c_0_22])])).
 % 0.06/0.66  cnf(c_0_40, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_26]), c_0_39]), c_0_22]), c_0_27])]), ['proof']).
 % 0.06/0.66  # SZS output end CNFRefutation
 % 0.06/0.66  # Training examples: 0 positive, 0 negative
 % 0.06/0.66
 % 0.06/0.66  # -------------------------------------------------
 % 0.06/0.66  # User time                : 0.424 s
 % 0.06/0.66  # System time              : 0.010 s
 % 0.06/0.66  # Total time               : 0.434 s
 % 0.06/0.66  # Maximum resident set size: 1880 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
