   --------------------------------------------------------------------------

                   TSTP Solution File: SWC221-1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC221-1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n180.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:06 EDT 2016

 % Result     : Unsatisfiable 0.03s
 % Output     : CNFRefutation 0.03s
 % Verified   :
 % Statistics : Number of clauses        :   57 ( 857 expanded)
 %              Number of leaves         :   25 ( 590 expanded)
 %              Depth                    :   18
 %              Number of atoms          :  135 (1376 expanded)
 %              Number of equality atoms :   34 ( 618 expanded)
 %              Maximal clause size      :    7 (   2 average)
 %              Maximal term depth       :    4 (   2 average)

 % Comments   :
 %------------------------------------------------------------------------------
 cnf(c_0_0,negated_conjecture,
     ( nil = sk3
     | app(app(sk7,cons(sk6,nil)),sk8) = sk3 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_16)).

 cnf(c_0_1,negated_conjecture,
     ( sk1 = sk3 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

 cnf(c_0_2,negated_conjecture,
     ( nil != sk1 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

 cnf(c_0_3,negated_conjecture,
     ( nil = sk3
     | ssList(sk7) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_14)).

 cnf(c_0_4,negated_conjecture,
     ( nil = sk3
     | ssList(sk8) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_15)).

 cnf(c_0_5,axiom,
     ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

 cnf(c_0_6,negated_conjecture,
     ( nil = sk3
     | ssItem(sk6) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

 cnf(c_0_7,axiom,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

 cnf(c_0_8,axiom,
     ( ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

 cnf(c_0_9,axiom,
     ( frontsegP(X3,X1)
     | app(X1,X2) != X3
     | ~ ssList(X2)
     | ~ ssList(X1)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause144)).

 cnf(c_0_10,axiom,
     ( ssList(app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

 cnf(c_0_11,axiom,
     ( app(cons(X1,nil),X2) = cons(X1,X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

 cnf(c_0_12,axiom,
     ( app(X2,skaf45(X1,X2)) = X1
     | ~ frontsegP(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause132)).

 cnf(c_0_13,negated_conjecture,
     ( ssList(sk1) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

 cnf(c_0_14,axiom,
     ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2)
     | ~ duplicatefreeP(X5)
     | ~ ssList(X5) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

 cnf(c_0_15,axiom,
     ( X2 = X3
     | app(X1,X2) != app(X1,X3)
     | ~ ssList(X2)
     | ~ ssList(X1)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause150)).

 cnf(c_0_16,axiom,
     ( ssList(skaf45(X1,X2)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause51)).

 cnf(c_0_17,axiom,
     ( duplicatefreeP(X1)
     | ssItem(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

 cnf(c_0_18,negated_conjecture,
     ( memberP(X2,sk5(X3,X2,X1))
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | app(app(X2,cons(X1,nil)),X3) != sk1 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

 cnf(c_0_19,axiom,
     ( ~ memberP(nil,X1)
     | ~ ssItem(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause71)).

 cnf(c_0_20,axiom,
     ( singletonP(X2)
     | cons(X1,nil) != X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause116)).

 cnf(c_0_21,axiom,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause74)).

 cnf(c_0_22,axiom,
     ( cons(skaf44(X1),nil) = X1
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause101)).

 cnf(c_0_23,axiom,
     ( cons(skaf83(X1),skaf82(X1)) = X1
     | nil = X1
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause109)).

 cnf(c_0_24,axiom,
     ( ssList(skaf82(X1)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause13)).

 cnf(c_0_25,negated_conjecture,
     ( app(app(sk7,cons(sk6,nil)),sk8) = sk1 ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[4]c_0_0,c_0_1]),c_0_1]),[5]c_0_2])).

 cnf(c_0_26,negated_conjecture,
     ( ssList(sk7) ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[[6]c_0_3,c_0_1]),[7]c_0_2])).

 cnf(c_0_27,negated_conjecture,
     ( ssList(sk8) ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[[8]c_0_4,c_0_1]),[9]c_0_2])).

 cnf(c_0_28,negated_conjecture,
     ( app(sk7,app(cons(sk6,nil),sk8)) = sk1
     | ~ ssList(cons(sk6,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[10]c_0_25,c_0_5]),c_0_26]),c_0_27])])).

 cnf(c_0_29,negated_conjecture,
     ( ssItem(sk6) ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[[11]c_0_6,c_0_1]),[12]c_0_2])).

 cnf(c_0_30,negated_conjecture,
     ( app(sk7,app(cons(sk6,nil),sk8)) = sk1 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[13]c_0_28,c_0_7]),c_0_29]),c_0_8])])).

 cnf(c_0_31,plain,
     ( frontsegP(app(X1,X2),X1)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_9]),[14]c_0_10])).

 cnf(c_0_32,negated_conjecture,
     ( app(sk7,cons(sk6,sk8)) = sk1 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_30,c_0_11]),c_0_29]),c_0_27])])).

 cnf(c_0_33,negated_conjecture,
     ( frontsegP(sk1,sk7)
     | ~ ssList(cons(sk6,sk8)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[16]c_0_31,c_0_32]),c_0_26])])).

 cnf(c_0_34,negated_conjecture,
     ( frontsegP(sk1,sk7) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_33,c_0_7]),c_0_29]),c_0_27])])).

 cnf(c_0_35,negated_conjecture,
     ( app(sk7,skaf45(sk1,sk7)) = sk1 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_12,c_0_34]),c_0_26]),c_0_13])])).

 cnf(c_0_36,plain,
     ( ~ ssItem(X1)
     | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
     inference(er,[status(thm)],[[19]c_0_14])).

 cnf(c_0_37,negated_conjecture,
     ( X1 = skaf45(sk1,sk7)
     | app(sk7,X1) != sk1
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[20]c_0_15,c_0_35]),c_0_16]),c_0_26])])).

 cnf(c_0_38,negated_conjecture,
     ( ~ ssList(app(sk1,cons(sk6,X1)))
     | ~ ssList(X1)
     | ~ duplicatefreeP(app(sk1,cons(sk6,X1))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[21]c_0_36,c_0_32]),c_0_29]),c_0_27]),c_0_26])])).

 cnf(c_0_39,negated_conjecture,
     ( skaf45(sk1,sk7) = cons(sk6,sk8)
     | ~ ssList(cons(sk6,sk8)) ),
     inference(spm,[status(thm)],[[22]c_0_37,[23]c_0_32])).

 cnf(c_0_40,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(app(sk1,cons(sk6,X2)))
     | ~ ssList(X2) ),
     inference(spm,[status(thm)],[[24]c_0_38,[25]c_0_17])).

 cnf(c_0_41,negated_conjecture,
     ( skaf45(sk1,sk7) = cons(sk6,sk8) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_39,c_0_7]),c_0_29]),c_0_27])])).

 cnf(c_0_42,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(cons(sk6,X2))
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_40,c_0_10]),c_0_13])])).

 cnf(c_0_43,negated_conjecture,
     ( ssList(cons(sk6,sk8)) ),
     inference(spm,[status(thm)],[[28]c_0_16,[29]c_0_41])).

 cnf(c_0_44,negated_conjecture,
     ( ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[30]c_0_42,c_0_43]),c_0_27])])).

 cnf(c_0_45,negated_conjecture,
     ( memberP(X1,sk5(X2,X1,X3))
     | app(app(X1,cons(X3,nil)),X2) != sk1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[31]c_0_18,c_0_44])])).

 cnf(c_0_46,plain,
     ( ~ memberP(nil,X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[32]c_0_19,c_0_44])])).

 cnf(c_0_47,plain,
     ( singletonP(cons(X1,nil))
     | ~ ssItem(X1)
     | ~ ssList(cons(X1,nil)) ),
     inference(er,[status(thm)],[[33]c_0_20])).

 cnf(c_0_48,negated_conjecture,
     ( app(cons(X1,nil),X2) != sk1
     | ~ ssList(cons(X1,nil))
     | ~ ssList(X2) ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[34]c_0_45,c_0_21]),c_0_8])]),[35]c_0_46])).

 cnf(c_0_49,plain,
     ( singletonP(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[36]c_0_47,c_0_7]),c_0_8])])).

 cnf(c_0_50,negated_conjecture,
     ( app(X1,X2) != sk1
     | ~ singletonP(X1)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(spm,[status(thm)],[[37]c_0_48,[38]c_0_22])).

 cnf(c_0_51,plain,
     ( app(cons(X1,nil),X2) = cons(X1,X2)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[39]c_0_11,c_0_44])])).

 cnf(c_0_52,plain,
     ( singletonP(cons(X1,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[40]c_0_49,c_0_44])])).

 cnf(c_0_53,negated_conjecture,
     ( cons(X1,X2) != sk1
     | ~ ssList(cons(X1,nil))
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[41]c_0_50,c_0_51]),c_0_52])])).

 cnf(c_0_54,negated_conjecture,
     ( ~ ssList(cons(skaf83(sk1),nil)) ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[42]c_0_53,c_0_23]),c_0_24])])]),c_0_13])]),[43]c_0_2])).

 cnf(c_0_55,plain,
     ( ssList(cons(X1,X2))
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[44]c_0_7,c_0_44])])).

 cnf(c_0_56,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[45]c_0_54,c_0_55]),c_0_8])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC221-1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.22  % Computer   : n180.star.cs.uiowa.edu
 % 0.03/0.22  % Model      : x86_64 x86_64
 % 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.22  % Memory     : 32218.75MB
 % 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.22  % CPULimit   : 300
 % 0.03/0.22  % DateTime   : Tue May 10 20:37:23 CDT 2016
 % 0.03/0.22  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.03/0.29  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4d
 % 0.03/0.29  # and selection function SelectCQIPrecWNTNp.
 % 0.03/0.29  #
 % 0.03/0.29  # Preprocessing time       : 0.009 s
 % 0.03/0.29  # Presaturation interreduction done
 % 0.03/0.29
 % 0.03/0.29  # Proof found!
 % 0.03/0.29  # SZS status Unsatisfiable
 % 0.03/0.29  # SZS output start CNFRefutation
 % 0.03/0.29  cnf(c_0_0, negated_conjecture, (nil=sk3|app(app(sk7,cons(sk6,nil)),sk8)=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_16)).
 % 0.03/0.29  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
 % 0.03/0.29  cnf(c_0_2, negated_conjecture, (nil!=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
 % 0.03/0.29  cnf(c_0_3, negated_conjecture, (nil=sk3|ssList(sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_14)).
 % 0.03/0.29  cnf(c_0_4, negated_conjecture, (nil=sk3|ssList(sk8)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_15)).
 % 0.03/0.29  cnf(c_0_5, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
 % 0.03/0.29  cnf(c_0_6, negated_conjecture, (nil=sk3|ssItem(sk6)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
 % 0.03/0.29  cnf(c_0_7, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
 % 0.03/0.29  cnf(c_0_8, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
 % 0.03/0.29  cnf(c_0_9, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause144)).
 % 0.03/0.29  cnf(c_0_10, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
 % 0.03/0.29  cnf(c_0_11, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
 % 0.03/0.29  cnf(c_0_12, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause132)).
 % 0.03/0.29  cnf(c_0_13, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
 % 0.03/0.29  cnf(c_0_14, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
 % 0.03/0.29  cnf(c_0_15, axiom, (X2=X3|app(X1,X2)!=app(X1,X3)|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause150)).
 % 0.03/0.29  cnf(c_0_16, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause51)).
 % 0.03/0.29  cnf(c_0_17, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
 % 0.03/0.29  cnf(c_0_18, negated_conjecture, (memberP(X2,sk5(X3,X2,X1))|~ssItem(X1)|~ssList(X2)|~ssList(X3)|app(app(X2,cons(X1,nil)),X3)!=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
 % 0.03/0.29  cnf(c_0_19, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause71)).
 % 0.03/0.29  cnf(c_0_20, axiom, (singletonP(X2)|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause116)).
 % 0.03/0.29  cnf(c_0_21, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause74)).
 % 0.03/0.29  cnf(c_0_22, axiom, (cons(skaf44(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause101)).
 % 0.03/0.29  cnf(c_0_23, axiom, (cons(skaf83(X1),skaf82(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause109)).
 % 0.03/0.29  cnf(c_0_24, axiom, (ssList(skaf82(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause13)).
 % 0.03/0.29  cnf(c_0_25, negated_conjecture, (app(app(sk7,cons(sk6,nil)),sk8)=sk1), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_0, c_0_1]), c_0_1]), c_0_2])).
 % 0.03/0.29  cnf(c_0_26, negated_conjecture, (ssList(sk7)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_1]), c_0_2])).
 % 0.03/0.29  cnf(c_0_27, negated_conjecture, (ssList(sk8)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_1]), c_0_2])).
 % 0.03/0.29  cnf(c_0_28, negated_conjecture, (app(sk7,app(cons(sk6,nil),sk8))=sk1|~ssList(cons(sk6,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_5]), c_0_26]), c_0_27])])).
 % 0.03/0.29  cnf(c_0_29, negated_conjecture, (ssItem(sk6)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_1]), c_0_2])).
 % 0.03/0.29  cnf(c_0_30, negated_conjecture, (app(sk7,app(cons(sk6,nil),sk8))=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_7]), c_0_29]), c_0_8])])).
 % 0.03/0.29  cnf(c_0_31, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_9]), c_0_10])).
 % 0.03/0.29  cnf(c_0_32, negated_conjecture, (app(sk7,cons(sk6,sk8))=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_11]), c_0_29]), c_0_27])])).
 % 0.03/0.29  cnf(c_0_33, negated_conjecture, (frontsegP(sk1,sk7)|~ssList(cons(sk6,sk8))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_26])])).
 % 0.03/0.29  cnf(c_0_34, negated_conjecture, (frontsegP(sk1,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_7]), c_0_29]), c_0_27])])).
 % 0.03/0.29  cnf(c_0_35, negated_conjecture, (app(sk7,skaf45(sk1,sk7))=sk1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_34]), c_0_26]), c_0_13])])).
 % 0.03/0.29  cnf(c_0_36, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_14])).
 % 0.03/0.29  cnf(c_0_37, negated_conjecture, (X1=skaf45(sk1,sk7)|app(sk7,X1)!=sk1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_35]), c_0_16]), c_0_26])])).
 % 0.03/0.29  cnf(c_0_38, negated_conjecture, (~ssList(app(sk1,cons(sk6,X1)))|~ssList(X1)|~duplicatefreeP(app(sk1,cons(sk6,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_32]), c_0_29]), c_0_27]), c_0_26])])).
 % 0.03/0.29  cnf(c_0_39, negated_conjecture, (skaf45(sk1,sk7)=cons(sk6,sk8)|~ssList(cons(sk6,sk8))), inference(spm,[status(thm)],[c_0_37, c_0_32])).
 % 0.03/0.29  cnf(c_0_40, negated_conjecture, (ssItem(X1)|~ssList(app(sk1,cons(sk6,X2)))|~ssList(X2)), inference(spm,[status(thm)],[c_0_38, c_0_17])).
 % 0.03/0.29  cnf(c_0_41, negated_conjecture, (skaf45(sk1,sk7)=cons(sk6,sk8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_7]), c_0_29]), c_0_27])])).
 % 0.03/0.29  cnf(c_0_42, negated_conjecture, (ssItem(X1)|~ssList(cons(sk6,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_10]), c_0_13])])).
 % 0.03/0.29  cnf(c_0_43, negated_conjecture, (ssList(cons(sk6,sk8))), inference(spm,[status(thm)],[c_0_16, c_0_41])).
 % 0.03/0.29  cnf(c_0_44, negated_conjecture, (ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_27])])).
 % 0.03/0.29  cnf(c_0_45, negated_conjecture, (memberP(X1,sk5(X2,X1,X3))|app(app(X1,cons(X3,nil)),X2)!=sk1|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_44])])).
 % 0.03/0.29  cnf(c_0_46, plain, (~memberP(nil,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_44])])).
 % 0.03/0.29  cnf(c_0_47, plain, (singletonP(cons(X1,nil))|~ssItem(X1)|~ssList(cons(X1,nil))), inference(er,[status(thm)],[c_0_20])).
 % 0.03/0.29  cnf(c_0_48, negated_conjecture, (app(cons(X1,nil),X2)!=sk1|~ssList(cons(X1,nil))|~ssList(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_21]), c_0_8])]), c_0_46])).
 % 0.03/0.29  cnf(c_0_49, plain, (singletonP(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_7]), c_0_8])])).
 % 0.03/0.29  cnf(c_0_50, negated_conjecture, (app(X1,X2)!=sk1|~singletonP(X1)|~ssList(X1)|~ssList(X2)), inference(spm,[status(thm)],[c_0_48, c_0_22])).
 % 0.03/0.29  cnf(c_0_51, plain, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_44])])).
 % 0.03/0.29  cnf(c_0_52, plain, (singletonP(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_44])])).
 % 0.03/0.29  cnf(c_0_53, negated_conjecture, (cons(X1,X2)!=sk1|~ssList(cons(X1,nil))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])])).
 % 0.03/0.29  cnf(c_0_54, negated_conjecture, (~ssList(cons(skaf83(sk1),nil))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_23]), c_0_24])])]), c_0_13])]), c_0_2])).
 % 0.03/0.29  cnf(c_0_55, plain, (ssList(cons(X1,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_44])])).
 % 0.03/0.29  cnf(c_0_56, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_8])]), ['proof']).
 % 0.03/0.29  # SZS output end CNFRefutation
 % 0.03/0.29  # Training examples: 0 positive, 0 negative
 % 0.03/0.30
 % 0.03/0.30  # -------------------------------------------------
 % 0.03/0.30  # User time                : 0.066 s
 % 0.03/0.30  # System time              : 0.006 s
 % 0.03/0.30  # Total time               : 0.071 s
 % 0.03/0.30  # Maximum resident set size: 1868 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
