   --------------------------------------------------------------------------

                   TSTP Solution File: SWC218+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC218+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n140.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:05 EDT 2016

 % Result     : Theorem 0.02s
 % Output     : CNFRefutation 0.02s
 % Verified   :
 % Statistics : Number of formulae       :   53 ( 276 expanded)
 %              Number of clauses        :   37 ( 107 expanded)
 %              Number of leaves         :    8 (  72 expanded)
 %              Depth                    :   17
 %              Number of atoms          :  154 (1757 expanded)
 %              Number of equality atoms :   63 ( 597 expanded)
 %              Maximal formula depth    :   21 (   4 average)
 %              Maximal clause size      :   13 (   2 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ! [X5] :
                         ( ssList(X5)
                        => ! [X6] :
                             ( ssList(X6)
                            => ( app(X5,X6) != X4
                               | ! [X7] :
                                   ( ssItem(X7)
                                  => app(app(X5,cons(X7,nil)),X6) != X3 ) ) ) )
                     | neq(X1,nil) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_5,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(X1,nil) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) != X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax18)).

 fof(c_0_8,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ! [X5] :
                           ( ssList(X5)
                          => ! [X6] :
                               ( ssList(X6)
                              => ( app(X5,X6) != X4
                                 | ! [X7] :
                                     ( ssItem(X7)
                                    => app(app(X5,cons(X7,nil)),X6) != X3 ) ) ) )
                       | neq(X1,nil) ) ) ) ) ) ),
     inference(assume_negation,[status(cth)],[[4]c_0_0])).

 fof(c_0_9,negated_conjecture,
     ( ssList(esk48_0)
     & ssList(esk49_0)
     & ssList(esk50_0)
     & ssList(esk51_0)
     & esk49_0 = esk51_0
     & esk48_0 = esk50_0
     & ssList(esk52_0)
     & ssList(esk53_0)
     & app(esk52_0,esk53_0) = esk51_0
     & ssItem(esk54_0)
     & app(app(esk52_0,cons(esk54_0,nil)),esk53_0) = esk50_0
     & ~ neq(esk48_0,nil) ),
     inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

 fof(c_0_10,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).

 cnf(c_0_11,negated_conjecture,
     ( app(app(esk52_0,cons(esk54_0,nil)),esk53_0) = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_9])).

 cnf(c_0_12,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_9])).

 cnf(c_0_13,plain,
     ( nil = X2
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[7]c_0_10])).

 cnf(c_0_14,negated_conjecture,
     ( app(app(esk52_0,cons(esk54_0,nil)),esk53_0) = esk48_0 ),
     inference(rw,[status(thm)],[[8]c_0_11,[9]c_0_12])).

 cnf(c_0_15,negated_conjecture,
     ( ssList(esk53_0) ),
     inference(split_conjunct,[status(thm)],[[10]c_0_9])).

 fof(c_0_16,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

 cnf(c_0_17,negated_conjecture,
     ( esk53_0 = nil
     | esk48_0 != nil
     | ~ ssList(app(esk52_0,cons(esk54_0,nil))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[11]c_0_13,c_0_14]),c_0_15])])).

 cnf(c_0_18,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_16])).

 cnf(c_0_19,negated_conjecture,
     ( ssList(esk52_0) ),
     inference(split_conjunct,[status(thm)],[[13]c_0_9])).

 fof(c_0_20,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

 fof(c_0_21,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).

 cnf(c_0_22,negated_conjecture,
     ( esk53_0 = nil
     | esk48_0 != nil
     | ~ ssList(cons(esk54_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[14]c_0_17,c_0_18]),c_0_19])])).

 cnf(c_0_23,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_20])).

 cnf(c_0_24,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_5])).

 cnf(c_0_25,negated_conjecture,
     ( ssItem(esk54_0) ),
     inference(split_conjunct,[status(thm)],[[17]c_0_9])).

 cnf(c_0_26,negated_conjecture,
     ( ~ neq(esk48_0,nil) ),
     inference(split_conjunct,[status(thm)],[[18]c_0_9])).

 cnf(c_0_27,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_21])).

 cnf(c_0_28,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[20]c_0_9])).

 cnf(c_0_29,plain,
     ( nil = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[21]c_0_10])).

 cnf(c_0_30,negated_conjecture,
     ( esk53_0 = nil
     | esk48_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[22]c_0_22,c_0_23]),c_0_24]),c_0_25])])).

 cnf(c_0_31,negated_conjecture,
     ( esk48_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[23]c_0_26,c_0_27]),c_0_24]),c_0_28])])).

 cnf(c_0_32,negated_conjecture,
     ( app(esk52_0,cons(esk54_0,nil)) = nil
     | esk48_0 != nil
     | ~ ssList(app(esk52_0,cons(esk54_0,nil))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[24]c_0_29,c_0_14]),c_0_15])])).

 fof(c_0_33,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(X2,nil) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).

 cnf(c_0_34,negated_conjecture,
     ( app(esk52_0,esk53_0) = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[25]c_0_9])).

 cnf(c_0_35,negated_conjecture,
     ( esk53_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[26]c_0_30,c_0_31])])).

 cnf(c_0_36,negated_conjecture,
     ( app(esk52_0,cons(esk54_0,nil)) = nil
     | esk48_0 != nil
     | ~ ssList(cons(esk54_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_32,c_0_18]),c_0_19])])).

 cnf(c_0_37,plain,
     ( app(X1,nil) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[28]c_0_33])).

 cnf(c_0_38,negated_conjecture,
     ( app(esk52_0,nil) = esk51_0 ),
     inference(rw,[status(thm)],[[29]c_0_34,[30]c_0_35])).

 cnf(c_0_39,negated_conjecture,
     ( app(esk52_0,cons(esk54_0,nil)) = nil
     | esk48_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[31]c_0_36,c_0_23]),c_0_24]),c_0_25])])).

 cnf(c_0_40,negated_conjecture,
     ( esk52_0 = esk51_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[32]c_0_37,c_0_38]),c_0_19])])).

 cnf(c_0_41,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[33]c_0_9])).

 cnf(c_0_42,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[34]c_0_9])).

 cnf(c_0_43,negated_conjecture,
     ( app(esk51_0,cons(esk54_0,nil)) = nil ),
     inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[[35]c_0_39,c_0_31])]),[36]c_0_40])).

 cnf(c_0_44,negated_conjecture,
     ( ssList(esk51_0) ),
     inference(rw,[status(thm)],[[37]c_0_41,[38]c_0_42])).

 cnf(c_0_45,negated_conjecture,
     ( esk51_0 = nil
     | ~ ssList(cons(esk54_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[39]c_0_29,c_0_43]),c_0_44])])).

 cnf(c_0_46,negated_conjecture,
     ( esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[40]c_0_45,c_0_23]),c_0_24]),c_0_25])])).

 cnf(c_0_47,negated_conjecture,
     ( app(nil,cons(esk54_0,nil)) = nil ),
     inference(rw,[status(thm)],[[41]c_0_43,[42]c_0_46])).

 fof(c_0_48,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) != X3 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).

 cnf(c_0_49,negated_conjecture,
     ( cons(esk54_0,nil) = nil
     | ~ ssList(cons(esk54_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[43]c_0_13,c_0_47]),c_0_24])])).

 cnf(c_0_50,plain,
     ( cons(X1,X2) != X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[44]c_0_48])).

 cnf(c_0_51,negated_conjecture,
     ( cons(esk54_0,nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[45]c_0_49,c_0_23]),c_0_24]),c_0_25])])).

 cnf(c_0_52,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[46]c_0_50,c_0_51]),c_0_24]),c_0_25])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC218+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.22  % Computer   : n140.star.cs.uiowa.edu
 % 0.02/0.22  % Model      : x86_64 x86_64
 % 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.22  % Memory     : 32218.75MB
 % 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.22  % CPULimit   : 300
 % 0.02/0.22  % DateTime   : Tue May 10 20:37:23 CDT 2016
 % 0.02/0.22  % CPUTime    :
 % 0.02/0.22  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.02/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.02/0.25  # and selection function SelectCQPrecWNTNp.
 % 0.02/0.25  #
 % 0.02/0.25  # Preprocessing time       : 0.014 s
 % 0.02/0.25  # Presaturation interreduction done
 % 0.02/0.25
 % 0.02/0.25  # Proof found!
 % 0.02/0.25  # SZS status Theorem
 % 0.02/0.25  # SZS output start CNFRefutation
 % 0.02/0.25  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(X5,X6)!=X4|![X7]:(ssItem(X7)=>app(app(X5,cons(X7,nil)),X6)!=X3)))))|neq(X1,nil))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.02/0.25  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 0.02/0.25  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 0.02/0.25  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.02/0.25  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 0.02/0.25  fof(c_0_5, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.02/0.25  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
 % 0.02/0.25  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax18)).
 % 0.02/0.25  fof(c_0_8, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(X5,X6)!=X4|![X7]:(ssItem(X7)=>app(app(X5,cons(X7,nil)),X6)!=X3)))))|neq(X1,nil)))))))), inference(assume_negation,[status(cth)],[c_0_0])).
 % 0.02/0.25  fof(c_0_9, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((esk49_0=esk51_0&esk48_0=esk50_0)&(ssList(esk52_0)&(ssList(esk53_0)&(app(esk52_0,esk53_0)=esk51_0&(ssItem(esk54_0)&app(app(esk52_0,cons(esk54_0,nil)),esk53_0)=esk50_0)))))&~neq(esk48_0,nil))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
 % 0.02/0.25  fof(c_0_10, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
 % 0.02/0.25  cnf(c_0_11, negated_conjecture, (app(app(esk52_0,cons(esk54_0,nil)),esk53_0)=esk50_0), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.02/0.25  cnf(c_0_12, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.02/0.25  cnf(c_0_13, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.02/0.25  cnf(c_0_14, negated_conjecture, (app(app(esk52_0,cons(esk54_0,nil)),esk53_0)=esk48_0), inference(rw,[status(thm)],[c_0_11, c_0_12])).
 % 0.02/0.25  cnf(c_0_15, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.02/0.25  fof(c_0_16, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
 % 0.02/0.25  cnf(c_0_17, negated_conjecture, (esk53_0=nil|esk48_0!=nil|~ssList(app(esk52_0,cons(esk54_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
 % 0.02/0.25  cnf(c_0_18, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
 % 0.02/0.25  cnf(c_0_19, negated_conjecture, (ssList(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.02/0.25  fof(c_0_20, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
 % 0.02/0.25  fof(c_0_21, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
 % 0.02/0.25  cnf(c_0_22, negated_conjecture, (esk53_0=nil|esk48_0!=nil|~ssList(cons(esk54_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
 % 0.02/0.25  cnf(c_0_23, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
 % 0.02/0.25  cnf(c_0_24, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.02/0.25  cnf(c_0_25, negated_conjecture, (ssItem(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.02/0.25  cnf(c_0_26, negated_conjecture, (~neq(esk48_0,nil)), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.02/0.25  cnf(c_0_27, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
 % 0.02/0.25  cnf(c_0_28, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.02/0.25  cnf(c_0_29, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
 % 0.02/0.25  cnf(c_0_30, negated_conjecture, (esk53_0=nil|esk48_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25])])).
 % 0.02/0.25  cnf(c_0_31, negated_conjecture, (esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_24]), c_0_28])])).
 % 0.02/0.25  cnf(c_0_32, negated_conjecture, (app(esk52_0,cons(esk54_0,nil))=nil|esk48_0!=nil|~ssList(app(esk52_0,cons(esk54_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_14]), c_0_15])])).
 % 0.02/0.25  fof(c_0_33, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
 % 0.02/0.25  cnf(c_0_34, negated_conjecture, (app(esk52_0,esk53_0)=esk51_0), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.02/0.25  cnf(c_0_35, negated_conjecture, (esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31])])).
 % 0.02/0.25  cnf(c_0_36, negated_conjecture, (app(esk52_0,cons(esk54_0,nil))=nil|esk48_0!=nil|~ssList(cons(esk54_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_18]), c_0_19])])).
 % 0.02/0.25  cnf(c_0_37, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
 % 0.02/0.25  cnf(c_0_38, negated_conjecture, (app(esk52_0,nil)=esk51_0), inference(rw,[status(thm)],[c_0_34, c_0_35])).
 % 0.02/0.25  cnf(c_0_39, negated_conjecture, (app(esk52_0,cons(esk54_0,nil))=nil|esk48_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_23]), c_0_24]), c_0_25])])).
 % 0.02/0.25  cnf(c_0_40, negated_conjecture, (esk52_0=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_19])])).
 % 0.02/0.25  cnf(c_0_41, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.02/0.25  cnf(c_0_42, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_9])).
 % 0.02/0.25  cnf(c_0_43, negated_conjecture, (app(esk51_0,cons(esk54_0,nil))=nil), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_31])]), c_0_40])).
 % 0.02/0.25  cnf(c_0_44, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_41, c_0_42])).
 % 0.02/0.25  cnf(c_0_45, negated_conjecture, (esk51_0=nil|~ssList(cons(esk54_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_43]), c_0_44])])).
 % 0.02/0.25  cnf(c_0_46, negated_conjecture, (esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_23]), c_0_24]), c_0_25])])).
 % 0.02/0.25  cnf(c_0_47, negated_conjecture, (app(nil,cons(esk54_0,nil))=nil), inference(rw,[status(thm)],[c_0_43, c_0_46])).
 % 0.02/0.25  fof(c_0_48, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
 % 0.02/0.25  cnf(c_0_49, negated_conjecture, (cons(esk54_0,nil)=nil|~ssList(cons(esk54_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_47]), c_0_24])])).
 % 0.02/0.25  cnf(c_0_50, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
 % 0.02/0.25  cnf(c_0_51, negated_conjecture, (cons(esk54_0,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_23]), c_0_24]), c_0_25])])).
 % 0.02/0.25  cnf(c_0_52, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_24]), c_0_25])]), ['proof']).
 % 0.02/0.25  # SZS output end CNFRefutation
 % 0.02/0.25  # Training examples: 0 positive, 0 negative
 % 0.02/0.25
 % 0.02/0.25  # -------------------------------------------------
 % 0.02/0.25  # User time                : 0.025 s
 % 0.02/0.25  # System time              : 0.005 s
 % 0.02/0.25  # Total time               : 0.030 s
 % 0.02/0.25  # Maximum resident set size: 1872 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
