   --------------------------------------------------------------------------

                   TSTP Solution File: SWC315-1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC315-1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n168.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:28 EDT 2016

 % Result     : Unsatisfiable 0.07s
 % Output     : CNFRefutation 0.07s
 % Verified   :
 % Statistics : Number of clauses        :   49 ( 840 expanded)
 %              Number of leaves         :   21 ( 546 expanded)
 %              Depth                    :   11
 %              Number of atoms          :   92 (1490 expanded)
 %              Number of equality atoms :   60 ( 721 expanded)
 %              Maximal clause size      :    4 (   2 average)
 %              Maximal term depth       :    3 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 cnf(c_0_0,negated_conjecture,
     ( ssList(sk2) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

 cnf(c_0_1,negated_conjecture,
     ( sk2 = sk4 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

 cnf(c_0_2,negated_conjecture,
     ( app(cons(sk7,nil),sk8) = sk4
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_18)).

 cnf(c_0_3,axiom,
     ( neq(X2,X1)
     | X2 = X1
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause100)).

 cnf(c_0_4,axiom,
     ( ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

 cnf(c_0_5,negated_conjecture,
     ( ssList(sk8)
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_17)).

 cnf(c_0_6,negated_conjecture,
     ( ssItem(sk7)
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_16)).

 cnf(c_0_7,negated_conjecture,
     ( app(cons(sk5,nil),sk6) = sk2
     | nil = sk2 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_11)).

 cnf(c_0_8,negated_conjecture,
     ( ssList(sk6)
     | nil = sk2 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

 cnf(c_0_9,negated_conjecture,
     ( ssItem(sk5)
     | nil = sk2 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

 cnf(c_0_10,axiom,
     ( app(cons(X1,nil),X2) = cons(X1,X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause120)).

 cnf(c_0_11,negated_conjecture,
     ( app(sk8,cons(sk7,nil)) = sk3
     | ~ neq(sk4,nil) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_19)).

 cnf(c_0_12,negated_conjecture,
     ( sk1 = sk3 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

 cnf(c_0_13,axiom,
     ( tl(cons(X1,X2)) = X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause96)).

 cnf(c_0_14,axiom,
     ( hd(cons(X1,X2)) = X1
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause97)).

 cnf(c_0_15,negated_conjecture,
     ( nil = sk2
     | app(sk6,cons(sk5,nil)) != sk1 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_12)).

 cnf(c_0_16,negated_conjecture,
     ( nil = sk3
     | nil != sk4 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

 cnf(c_0_17,negated_conjecture,
     ( app(cons(sk5,nil),sk6) = sk2
     | nil != sk1 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_14)).

 cnf(c_0_18,negated_conjecture,
     ( ssItem(sk5)
     | nil != sk1 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

 cnf(c_0_19,negated_conjecture,
     ( ssList(sk6)
     | nil != sk1 ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_13)).

 cnf(c_0_20,axiom,
     ( cons(X1,X2) != nil
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause98)).

 cnf(c_0_21,negated_conjecture,
     ( ssList(sk4) ),
     inference(rw,[status(thm)],[[4]c_0_0,[5]c_0_1])).

 cnf(c_0_22,negated_conjecture,
     ( app(cons(sk7,nil),sk8) = sk4
     | sk4 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[6]c_0_2,c_0_3]),c_0_21]),c_0_4])])).

 cnf(c_0_23,negated_conjecture,
     ( sk4 = nil
     | ssList(sk8) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[7]c_0_5,c_0_3]),c_0_21]),c_0_4])])).

 cnf(c_0_24,negated_conjecture,
     ( sk4 = nil
     | ssItem(sk7) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[8]c_0_6,c_0_3]),c_0_21]),c_0_4])])).

 cnf(c_0_25,negated_conjecture,
     ( app(cons(sk5,nil),sk6) = sk4
     | sk4 = nil ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[9]c_0_7,c_0_1]),[10]c_0_1])).

 cnf(c_0_26,negated_conjecture,
     ( sk4 = nil
     | ssList(sk6) ),
     inference(rw,[status(thm)],[[11]c_0_8,[12]c_0_1])).

 cnf(c_0_27,negated_conjecture,
     ( sk4 = nil
     | ssItem(sk5) ),
     inference(rw,[status(thm)],[[13]c_0_9,[14]c_0_1])).

 cnf(c_0_28,negated_conjecture,
     ( cons(sk7,sk8) = sk4
     | sk4 = nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_22,c_0_10]),c_0_23]),[16]c_0_24])).

 cnf(c_0_29,negated_conjecture,
     ( cons(sk5,sk6) = sk4
     | sk4 = nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_25,c_0_10]),c_0_26]),[18]c_0_27])).

 cnf(c_0_30,negated_conjecture,
     ( app(sk8,cons(sk7,nil)) = sk1
     | ~ neq(sk4,nil) ),
     inference(rw,[status(thm)],[[19]c_0_11,[20]c_0_12])).

 cnf(c_0_31,negated_conjecture,
     ( tl(sk4) = sk8
     | sk4 = nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[21]c_0_13,c_0_28]),c_0_23]),[22]c_0_24])).

 cnf(c_0_32,negated_conjecture,
     ( tl(sk4) = sk6
     | sk4 = nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[23]c_0_13,c_0_29]),c_0_26]),[24]c_0_27])).

 cnf(c_0_33,negated_conjecture,
     ( app(sk8,cons(sk7,nil)) = sk1
     | sk4 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_30,c_0_3]),c_0_21]),c_0_4])])).

 cnf(c_0_34,negated_conjecture,
     ( sk4 = nil
     | sk8 = sk6 ),
     inference(spm,[status(thm)],[[26]c_0_31,[27]c_0_32])).

 cnf(c_0_35,negated_conjecture,
     ( hd(sk4) = sk7
     | sk4 = nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[28]c_0_14,c_0_28]),c_0_23]),[29]c_0_24])).

 cnf(c_0_36,negated_conjecture,
     ( hd(sk4) = sk5
     | sk4 = nil ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[30]c_0_14,c_0_29]),c_0_26]),[31]c_0_27])).

 cnf(c_0_37,negated_conjecture,
     ( app(sk6,cons(sk7,nil)) = sk1
     | sk4 = nil ),
     inference(spm,[status(thm)],[[32]c_0_33,[33]c_0_34])).

 cnf(c_0_38,negated_conjecture,
     ( sk4 = nil
     | sk7 = sk5 ),
     inference(spm,[status(thm)],[[34]c_0_35,[35]c_0_36])).

 cnf(c_0_39,negated_conjecture,
     ( sk4 = nil
     | app(sk6,cons(sk5,nil)) != sk1 ),
     inference(rw,[status(thm)],[[36]c_0_15,[37]c_0_1])).

 cnf(c_0_40,negated_conjecture,
     ( sk1 = nil
     | sk4 != nil ),
     inference(rw,[status(thm)],[[38]c_0_16,[39]c_0_12])).

 cnf(c_0_41,negated_conjecture,
     ( sk4 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[40]c_0_37,c_0_38]),[41]c_0_39])).

 cnf(c_0_42,negated_conjecture,
     ( app(cons(sk5,nil),sk6) = sk4
     | sk1 != nil ),
     inference(rw,[status(thm)],[[42]c_0_17,[43]c_0_1])).

 cnf(c_0_43,negated_conjecture,
     ( sk1 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[44]c_0_40,c_0_41])])).

 cnf(c_0_44,negated_conjecture,
     ( app(cons(sk5,nil),sk6) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[45]c_0_42,c_0_41]),c_0_43])])).

 cnf(c_0_45,negated_conjecture,
     ( ssItem(sk5) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[46]c_0_18,c_0_43])])).

 cnf(c_0_46,negated_conjecture,
     ( ssList(sk6) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[47]c_0_19,c_0_43])])).

 cnf(c_0_47,negated_conjecture,
     ( cons(sk5,sk6) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[48]c_0_10,c_0_44]),c_0_45]),c_0_46])])).

 cnf(c_0_48,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[49]c_0_20,c_0_47]),c_0_45]),c_0_46])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.07  % Problem    : SWC315-1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.08  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.07/0.30  % Computer   : n168.star.cs.uiowa.edu
 % 0.07/0.30  % Model      : x86_64 x86_64
 % 0.07/0.30  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.07/0.30  % Memory     : 32218.75MB
 % 0.07/0.30  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.07/0.30  % CPULimit   : 300
 % 0.07/0.30  % DateTime   : Tue May 10 20:40:53 CDT 2016
 % 0.07/0.30  % CPUTime    :
 % 0.07/0.30  # Version: 1.9.1pre011
 % 0.07/0.31  # No SInE strategy applied
 % 0.07/0.31  # Trying AutoSched0 for 151 seconds
 % 0.07/0.37  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
 % 0.07/0.37  # and selection function SelectCQArNTEqFirst.
 % 0.07/0.37  #
 % 0.07/0.37  # Preprocessing time       : 0.019 s
 % 0.07/0.37  # Presaturation interreduction done
 % 0.07/0.37
 % 0.07/0.37  # Proof found!
 % 0.07/0.37  # SZS status Unsatisfiable
 % 0.07/0.37  # SZS output start CNFRefutation
 % 0.07/0.37  cnf(c_0_0, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
 % 0.07/0.37  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
 % 0.07/0.37  cnf(c_0_2, negated_conjecture, (app(cons(sk7,nil),sk8)=sk4|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_18)).
 % 0.07/0.37  cnf(c_0_3, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause100)).
 % 0.07/0.37  cnf(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
 % 0.07/0.37  cnf(c_0_5, negated_conjecture, (ssList(sk8)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_17)).
 % 0.07/0.37  cnf(c_0_6, negated_conjecture, (ssItem(sk7)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_16)).
 % 0.07/0.37  cnf(c_0_7, negated_conjecture, (app(cons(sk5,nil),sk6)=sk2|nil=sk2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_11)).
 % 0.07/0.37  cnf(c_0_8, negated_conjecture, (ssList(sk6)|nil=sk2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
 % 0.07/0.37  cnf(c_0_9, negated_conjecture, (ssItem(sk5)|nil=sk2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
 % 0.07/0.37  cnf(c_0_10, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause120)).
 % 0.07/0.37  cnf(c_0_11, negated_conjecture, (app(sk8,cons(sk7,nil))=sk3|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_19)).
 % 0.07/0.37  cnf(c_0_12, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
 % 0.07/0.37  cnf(c_0_13, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause96)).
 % 0.07/0.37  cnf(c_0_14, axiom, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause97)).
 % 0.07/0.37  cnf(c_0_15, negated_conjecture, (nil=sk2|app(sk6,cons(sk5,nil))!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_12)).
 % 0.07/0.37  cnf(c_0_16, negated_conjecture, (nil=sk3|nil!=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
 % 0.07/0.37  cnf(c_0_17, negated_conjecture, (app(cons(sk5,nil),sk6)=sk2|nil!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_14)).
 % 0.07/0.37  cnf(c_0_18, negated_conjecture, (ssItem(sk5)|nil!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
 % 0.07/0.37  cnf(c_0_19, negated_conjecture, (ssList(sk6)|nil!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_13)).
 % 0.07/0.37  cnf(c_0_20, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause98)).
 % 0.07/0.37  cnf(c_0_21, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_0, c_0_1])).
 % 0.07/0.37  cnf(c_0_22, negated_conjecture, (app(cons(sk7,nil),sk8)=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_3]), c_0_21]), c_0_4])])).
 % 0.07/0.37  cnf(c_0_23, negated_conjecture, (sk4=nil|ssList(sk8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_3]), c_0_21]), c_0_4])])).
 % 0.07/0.37  cnf(c_0_24, negated_conjecture, (sk4=nil|ssItem(sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_3]), c_0_21]), c_0_4])])).
 % 0.07/0.37  cnf(c_0_25, negated_conjecture, (app(cons(sk5,nil),sk6)=sk4|sk4=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_1]), c_0_1])).
 % 0.07/0.37  cnf(c_0_26, negated_conjecture, (sk4=nil|ssList(sk6)), inference(rw,[status(thm)],[c_0_8, c_0_1])).
 % 0.07/0.37  cnf(c_0_27, negated_conjecture, (sk4=nil|ssItem(sk5)), inference(rw,[status(thm)],[c_0_9, c_0_1])).
 % 0.07/0.37  cnf(c_0_28, negated_conjecture, (cons(sk7,sk8)=sk4|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_10]), c_0_23]), c_0_24])).
 % 0.07/0.37  cnf(c_0_29, negated_conjecture, (cons(sk5,sk6)=sk4|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_10]), c_0_26]), c_0_27])).
 % 0.07/0.37  cnf(c_0_30, negated_conjecture, (app(sk8,cons(sk7,nil))=sk1|~neq(sk4,nil)), inference(rw,[status(thm)],[c_0_11, c_0_12])).
 % 0.07/0.37  cnf(c_0_31, negated_conjecture, (tl(sk4)=sk8|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_28]), c_0_23]), c_0_24])).
 % 0.07/0.37  cnf(c_0_32, negated_conjecture, (tl(sk4)=sk6|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_29]), c_0_26]), c_0_27])).
 % 0.07/0.37  cnf(c_0_33, negated_conjecture, (app(sk8,cons(sk7,nil))=sk1|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_3]), c_0_21]), c_0_4])])).
 % 0.07/0.37  cnf(c_0_34, negated_conjecture, (sk4=nil|sk8=sk6), inference(spm,[status(thm)],[c_0_31, c_0_32])).
 % 0.07/0.37  cnf(c_0_35, negated_conjecture, (hd(sk4)=sk7|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_28]), c_0_23]), c_0_24])).
 % 0.07/0.37  cnf(c_0_36, negated_conjecture, (hd(sk4)=sk5|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_29]), c_0_26]), c_0_27])).
 % 0.07/0.37  cnf(c_0_37, negated_conjecture, (app(sk6,cons(sk7,nil))=sk1|sk4=nil), inference(spm,[status(thm)],[c_0_33, c_0_34])).
 % 0.07/0.37  cnf(c_0_38, negated_conjecture, (sk4=nil|sk7=sk5), inference(spm,[status(thm)],[c_0_35, c_0_36])).
 % 0.07/0.37  cnf(c_0_39, negated_conjecture, (sk4=nil|app(sk6,cons(sk5,nil))!=sk1), inference(rw,[status(thm)],[c_0_15, c_0_1])).
 % 0.07/0.37  cnf(c_0_40, negated_conjecture, (sk1=nil|sk4!=nil), inference(rw,[status(thm)],[c_0_16, c_0_12])).
 % 0.07/0.37  cnf(c_0_41, negated_conjecture, (sk4=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
 % 0.07/0.37  cnf(c_0_42, negated_conjecture, (app(cons(sk5,nil),sk6)=sk4|sk1!=nil), inference(rw,[status(thm)],[c_0_17, c_0_1])).
 % 0.07/0.37  cnf(c_0_43, negated_conjecture, (sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41])])).
 % 0.07/0.37  cnf(c_0_44, negated_conjecture, (app(cons(sk5,nil),sk6)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_41]), c_0_43])])).
 % 0.07/0.37  cnf(c_0_45, negated_conjecture, (ssItem(sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_43])])).
 % 0.07/0.37  cnf(c_0_46, negated_conjecture, (ssList(sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_43])])).
 % 0.07/0.37  cnf(c_0_47, negated_conjecture, (cons(sk5,sk6)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_44]), c_0_45]), c_0_46])])).
 % 0.07/0.37  cnf(c_0_48, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_47]), c_0_45]), c_0_46])]), ['proof']).
 % 0.07/0.37  # SZS output end CNFRefutation
 % 0.07/0.37  # Training examples: 0 positive, 0 negative
 % 0.07/0.37
 % 0.07/0.37  # -------------------------------------------------
 % 0.07/0.37  # User time                : 0.063 s
 % 0.07/0.37  # System time              : 0.004 s
 % 0.07/0.37  # Total time               : 0.067 s
 % 0.07/0.37  # Maximum resident set size: 1864 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
