   --------------------------------------------------------------------------

                   TSTP Solution File: SWC323+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC323+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n168.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:29 EDT 2016

 % Result     : Theorem 0.07s
 % Output     : CNFRefutation 0.07s
 % Verified   :
 % Statistics : Number of formulae       :   33 ( 224 expanded)
 %              Number of clauses        :   23 (  88 expanded)
 %              Number of leaves         :    5 (  61 expanded)
 %              Depth                    :    8
 %              Number of atoms          :  134 (2061 expanded)
 %              Number of equality atoms :   55 ( 922 expanded)
 %              Maximal formula depth    :   21 (   5 average)
 %              Maximal clause size      :   20 (   2 average)
 %              Maximal term depth       :    3 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ? [X5] :
                         ( ssList(X5)
                         & ? [X6] :
                             ( ssList(X6)
                             & app(X5,X6) = X2
                             & app(X6,X5) = X1 ) )
                     | ( nil != X3
                       & nil = X4 )
                     | ( ! [X7] :
                           ( ssItem(X7)
                          => ! [X8] :
                               ( ssList(X8)
                              => ( app(cons(X7,nil),X8) != X4
                                 | app(X8,cons(X7,nil)) != X3 ) ) )
                       & neq(X4,nil) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_2,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_5,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ? [X5] :
                           ( ssList(X5)
                           & ? [X6] :
                               ( ssList(X6)
                               & app(X5,X6) = X2
                               & app(X6,X5) = X1 ) )
                       | ( nil != X3
                         & nil = X4 )
                       | ( ! [X7] :
                             ( ssItem(X7)
                            => ! [X8] :
                                 ( ssList(X8)
                                => ( app(cons(X7,nil),X8) != X4
                                   | app(X8,cons(X7,nil)) != X3 ) ) )
                         & neq(X4,nil) ) ) ) ) ) ) ),
     inference(assume_negation,[status(cth)],[[4]c_0_0])).

 fof(c_0_6,negated_conjecture,(
     ! [X13,X14] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ( ~ ssList(X13)
         | ~ ssList(X14)
         | app(X13,X14) != esk49_0
         | app(X14,X13) != esk48_0 )
       & ( nil = esk50_0
         | nil != esk51_0 )
       & ( ssItem(esk52_0)
         | ~ neq(esk51_0,nil) )
       & ( ssList(esk53_0)
         | ~ neq(esk51_0,nil) )
       & ( app(cons(esk52_0,nil),esk53_0) = esk51_0
         | ~ neq(esk51_0,nil) )
       & ( app(esk53_0,cons(esk52_0,nil)) = esk50_0
         | ~ neq(esk51_0,nil) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).

 fof(c_0_7,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).

 cnf(c_0_8,negated_conjecture,
     ( app(X1,X2) != esk48_0
     | app(X2,X1) != esk49_0
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[5]c_0_6])).

 cnf(c_0_9,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_6])).

 cnf(c_0_10,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[7]c_0_6])).

 cnf(c_0_11,plain,
     ( nil = app(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != X1
     | nil != X2 ),
     inference(split_conjunct,[status(thm)],[[8]c_0_7])).

 cnf(c_0_12,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[9]c_0_2])).

 fof(c_0_13,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).

 cnf(c_0_14,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[10]c_0_6])).

 cnf(c_0_15,negated_conjecture,
     ( app(X1,X2) != esk51_0
     | app(X2,X1) != esk50_0
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[11]c_0_8,c_0_9]),[12]c_0_10])).

 cnf(c_0_16,plain,
     ( app(nil,nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_11])]),c_0_12])])).

 cnf(c_0_17,negated_conjecture,
     ( nil = esk50_0
     | nil != esk51_0 ),
     inference(split_conjunct,[status(thm)],[[13]c_0_6])).

 cnf(c_0_18,negated_conjecture,
     ( app(cons(esk52_0,nil),esk53_0) = esk51_0
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_6])).

 cnf(c_0_19,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_13])).

 cnf(c_0_20,negated_conjecture,
     ( ssList(esk51_0) ),
     inference(rw,[status(thm)],[[16]c_0_14,[17]c_0_9])).

 cnf(c_0_21,negated_conjecture,
     ( esk51_0 != nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_15,c_0_16]),c_0_16]),c_0_12])]),[19]c_0_17])).

 cnf(c_0_22,negated_conjecture,
     ( app(esk53_0,cons(esk52_0,nil)) = esk50_0
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[20]c_0_6])).

 cnf(c_0_23,negated_conjecture,
     ( ssList(esk53_0)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[21]c_0_6])).

 cnf(c_0_24,negated_conjecture,
     ( app(cons(esk52_0,nil),esk53_0) = esk51_0 ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[22]c_0_18,c_0_19]),c_0_12]),c_0_20])]),[23]c_0_21])).

 cnf(c_0_25,negated_conjecture,
     ( app(esk53_0,cons(esk52_0,nil)) = esk50_0 ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[24]c_0_22,c_0_19]),c_0_12]),c_0_20])]),[25]c_0_21])).

 cnf(c_0_26,negated_conjecture,
     ( ssList(esk53_0) ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_23,c_0_19]),c_0_12]),c_0_20])]),[27]c_0_21])).

 fof(c_0_27,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

 cnf(c_0_28,negated_conjecture,
     ( ssItem(esk52_0)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[28]c_0_6])).

 cnf(c_0_29,negated_conjecture,
     ( ~ ssList(cons(esk52_0,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[29]c_0_15,c_0_24]),c_0_25]),c_0_26])])).

 cnf(c_0_30,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[30]c_0_27])).

 cnf(c_0_31,negated_conjecture,
     ( ssItem(esk52_0) ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[31]c_0_28,c_0_19]),c_0_12]),c_0_20])]),[32]c_0_21])).

 cnf(c_0_32,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[33]c_0_29,c_0_30]),c_0_12]),c_0_31])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.01/0.07  % Problem    : SWC323+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.01/0.08  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.07/0.30  % Computer   : n168.star.cs.uiowa.edu
 % 0.07/0.30  % Model      : x86_64 x86_64
 % 0.07/0.30  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.07/0.30  % Memory     : 32218.75MB
 % 0.07/0.30  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.07/0.30  % CPULimit   : 300
 % 0.07/0.30  % DateTime   : Tue May 10 20:40:38 CDT 2016
 % 0.07/0.30  % CPUTime    :
 % 0.07/0.30  # Version: 1.9.1pre011
 % 0.07/0.31  # No SInE strategy applied
 % 0.07/0.31  # Trying AutoSched0 for 151 seconds
 % 0.07/0.36  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.07/0.36  # and selection function SelectCQPrecWNTNp.
 % 0.07/0.36  #
 % 0.07/0.36  # Preprocessing time       : 0.028 s
 % 0.07/0.36  # Presaturation interreduction done
 % 0.07/0.36
 % 0.07/0.36  # Proof found!
 % 0.07/0.36  # SZS status Theorem
 % 0.07/0.36  # SZS output start CNFRefutation
 % 0.07/0.36  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|?[X5]:(ssList(X5)&?[X6]:((ssList(X6)&app(X5,X6)=X2)&app(X6,X5)=X1)))|(nil!=X3&nil=X4))|(![X7]:(ssItem(X7)=>![X8]:(ssList(X8)=>(app(cons(X7,nil),X8)!=X4|app(X8,cons(X7,nil))!=X3)))&neq(X4,nil)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.07/0.36  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 0.07/0.36  fof(c_0_2, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.07/0.36  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 0.07/0.36  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.07/0.36  fof(c_0_5, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|?[X5]:(ssList(X5)&?[X6]:((ssList(X6)&app(X5,X6)=X2)&app(X6,X5)=X1)))|(nil!=X3&nil=X4))|(![X7]:(ssItem(X7)=>![X8]:(ssList(X8)=>(app(cons(X7,nil),X8)!=X4|app(X8,cons(X7,nil))!=X3)))&neq(X4,nil))))))))), inference(assume_negation,[status(cth)],[c_0_0])).
 % 0.07/0.36  fof(c_0_6, negated_conjecture, (![X13]:![X14]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(~ssList(X13)|((~ssList(X14)|app(X13,X14)!=esk49_0)|app(X14,X13)!=esk48_0)))&(nil=esk50_0|nil!=esk51_0))&((ssItem(esk52_0)|~neq(esk51_0,nil))&((ssList(esk53_0)|~neq(esk51_0,nil))&((app(cons(esk52_0,nil),esk53_0)=esk51_0|~neq(esk51_0,nil))&(app(esk53_0,cons(esk52_0,nil))=esk50_0|~neq(esk51_0,nil))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
 % 0.07/0.36  fof(c_0_7, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
 % 0.07/0.36  cnf(c_0_8, negated_conjecture, (app(X1,X2)!=esk48_0|app(X2,X1)!=esk49_0|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.07/0.36  cnf(c_0_9, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.07/0.36  cnf(c_0_10, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.07/0.36  cnf(c_0_11, plain, (nil=app(X1,X2)|~ssList(X1)|~ssList(X2)|nil!=X1|nil!=X2), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.07/0.36  cnf(c_0_12, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.07/0.36  fof(c_0_13, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
 % 0.07/0.36  cnf(c_0_14, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.07/0.36  cnf(c_0_15, negated_conjecture, (app(X1,X2)!=esk51_0|app(X2,X1)!=esk50_0|~ssList(X1)|~ssList(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_9]), c_0_10])).
 % 0.07/0.36  cnf(c_0_16, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_11])]), c_0_12])])).
 % 0.07/0.36  cnf(c_0_17, negated_conjecture, (nil=esk50_0|nil!=esk51_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.07/0.36  cnf(c_0_18, negated_conjecture, (app(cons(esk52_0,nil),esk53_0)=esk51_0|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.07/0.36  cnf(c_0_19, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.07/0.36  cnf(c_0_20, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_14, c_0_9])).
 % 0.07/0.36  cnf(c_0_21, negated_conjecture, (esk51_0!=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_16]), c_0_12])]), c_0_17])).
 % 0.07/0.36  cnf(c_0_22, negated_conjecture, (app(esk53_0,cons(esk52_0,nil))=esk50_0|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.07/0.36  cnf(c_0_23, negated_conjecture, (ssList(esk53_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.07/0.36  cnf(c_0_24, negated_conjecture, (app(cons(esk52_0,nil),esk53_0)=esk51_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_12]), c_0_20])]), c_0_21])).
 % 0.07/0.36  cnf(c_0_25, negated_conjecture, (app(esk53_0,cons(esk52_0,nil))=esk50_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_19]), c_0_12]), c_0_20])]), c_0_21])).
 % 0.07/0.36  cnf(c_0_26, negated_conjecture, (ssList(esk53_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_19]), c_0_12]), c_0_20])]), c_0_21])).
 % 0.07/0.36  fof(c_0_27, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 0.07/0.36  cnf(c_0_28, negated_conjecture, (ssItem(esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 0.07/0.36  cnf(c_0_29, negated_conjecture, (~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_24]), c_0_25]), c_0_26])])).
 % 0.07/0.36  cnf(c_0_30, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
 % 0.07/0.36  cnf(c_0_31, negated_conjecture, (ssItem(esk52_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_19]), c_0_12]), c_0_20])]), c_0_21])).
 % 0.07/0.36  cnf(c_0_32, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_12]), c_0_31])]), ['proof']).
 % 0.07/0.36  # SZS output end CNFRefutation
 % 0.07/0.36  # Training examples: 0 positive, 0 negative
 % 0.07/0.36
 % 0.07/0.36  # -------------------------------------------------
 % 0.07/0.36  # User time                : 0.052 s
 % 0.07/0.36  # System time              : 0.007 s
 % 0.07/0.36  # Total time               : 0.060 s
 % 0.07/0.36  # Maximum resident set size: 1876 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
