   --------------------------------------------------------------------------

                   TSTP Solution File: SWC340-1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC340-1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n170.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:33 EDT 2016

 % Result     : Unsatisfiable 224.77s
 % Output     : CNFRefutation 224.77s
 % Verified   :
 % Statistics : Number of clauses        :   53 ( 286 expanded)
 %              Number of leaves         :   27 ( 176 expanded)
 %              Depth                    :   14
 %              Number of atoms          :  154 ( 813 expanded)
 %              Number of equality atoms :    9 (  53 expanded)
 %              Maximal clause size      :    9 (   3 average)
 %              Maximal term depth       :    5 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 cnf(c_0_0,axiom,
     ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2)
     | ~ duplicatefreeP(X5)
     | ~ ssList(X5) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

 cnf(c_0_1,axiom,
     ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
     | ~ memberP(X1,X2)
     | ~ ssItem(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause169)).

 cnf(c_0_2,axiom,
     ( ssList(skaf43(X1,X2)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause52)).

 cnf(c_0_3,axiom,
     ( ssList(skaf42(X1,X2)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause53)).

 cnf(c_0_4,axiom,
     ( cons(skaf44(X1),nil) = X1
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause101)).

 cnf(c_0_5,axiom,
     ( ssItem(skaf44(X1)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause47)).

 cnf(c_0_6,axiom,
     ( ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

 cnf(c_0_7,axiom,
     ( ssList(app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

 cnf(c_0_8,axiom,
     ( memberP(cons(X2,X3),X1)
     | X1 != X2
     | ~ ssList(X3)
     | ~ ssItem(X2)
     | ~ ssItem(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause138)).

 cnf(c_0_9,axiom,
     ( duplicatefreeP(X1)
     | ssItem(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

 cnf(c_0_10,axiom,
     ( singletonP(X2)
     | cons(X1,nil) != X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause116)).

 cnf(c_0_11,axiom,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

 cnf(c_0_12,axiom,
     ( leq(X1,X2)
     | ~ lt(X1,X2)
     | ~ ssItem(X2)
     | ~ ssItem(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause103)).

 cnf(c_0_13,axiom,
     ( lt(X1,X3)
     | ~ leq(X1,X2)
     | ~ lt(X2,X3)
     | ~ ssItem(X3)
     | ~ ssItem(X2)
     | ~ ssItem(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause147)).

 cnf(c_0_14,axiom,
     ( lt(X2,X4)
     | app(app(X1,cons(X2,X3)),cons(X4,X5)) != X6
     | ~ ssList(X5)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X4)
     | ~ ssItem(X2)
     | ~ strictorderedP(X6)
     | ~ ssList(X6) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause181)).

 cnf(c_0_15,axiom,
     ( totalorderedP(X1)
     | ~ leq(skaf64(X1),skaf65(X1))
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause91)).

 cnf(c_0_16,axiom,
     ( leq(X1,X1)
     | ~ ssItem(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause62)).

 cnf(c_0_17,negated_conjecture,
     ( ~ segmentP(sk2,sk1)
     | ~ totalorderedP(sk1) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

 cnf(c_0_18,negated_conjecture,
     ( sk2 = sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

 cnf(c_0_19,negated_conjecture,
     ( segmentP(sk4,sk3) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

 cnf(c_0_20,negated_conjecture,
     ( sk1 = sk3 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

 cnf(c_0_21,axiom,
     ( totalorderedP(X1)
     | app(app(skaf66(X1),cons(skaf64(X1),skaf67(X1))),cons(skaf65(X1),skaf68(X1))) = X1
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause164)).

 cnf(c_0_22,axiom,
     ( ssList(skaf68(X1)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause27)).

 cnf(c_0_23,axiom,
     ( ssList(skaf67(X1)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause28)).

 cnf(c_0_24,axiom,
     ( ssList(skaf66(X1)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause29)).

 cnf(c_0_25,negated_conjecture,
     ( strictorderedP(sk3) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

 cnf(c_0_26,negated_conjecture,
     ( ssList(sk1) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

 cnf(c_0_27,plain,
     ( ~ ssItem(X1)
     | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
     inference(er,[status(thm)],[[4]c_0_0])).

 cnf(c_0_28,plain,
     ( ~ memberP(X1,X2)
     | ~ ssItem(X2)
     | ~ ssList(app(X1,cons(X2,X3)))
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ duplicatefreeP(app(X1,cons(X2,X3))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[5]c_0_27,c_0_1]),c_0_2]),c_0_3])])).

 cnf(c_0_29,plain,
     ( ~ memberP(X1,skaf44(X2))
     | ~ singletonP(X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ duplicatefreeP(app(X1,X2)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[6]c_0_28,c_0_4]),c_0_5]),c_0_6])]),[7]c_0_7])).

 cnf(c_0_30,plain,
     ( memberP(cons(X1,X2),X1)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(er,[status(thm)],[[8]c_0_8])).

 cnf(c_0_31,plain,
     ( ssItem(X1)
     | ~ memberP(X2,skaf44(X3))
     | ~ singletonP(X3)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[9]c_0_29,c_0_9]),[10]c_0_7])).

 cnf(c_0_32,plain,
     ( memberP(X1,skaf44(X1))
     | ~ singletonP(X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[11]c_0_30,c_0_4]),c_0_5]),c_0_6])])).

 cnf(c_0_33,plain,
     ( singletonP(cons(X1,nil))
     | ~ ssItem(X1)
     | ~ ssList(cons(X1,nil)) ),
     inference(er,[status(thm)],[[12]c_0_10])).

 cnf(c_0_34,plain,
     ( ssItem(X1)
     | ~ singletonP(X2)
     | ~ ssList(X2) ),
     inference(spm,[status(thm)],[[13]c_0_31,[14]c_0_32])).

 cnf(c_0_35,plain,
     ( singletonP(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_33,c_0_11]),c_0_6])])).

 cnf(c_0_36,plain,
     ( ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(cons(X2,nil)) ),
     inference(spm,[status(thm)],[[16]c_0_34,[17]c_0_35])).

 cnf(c_0_37,plain,
     ( ssItem(X1)
     | ~ ssItem(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_36,c_0_11]),c_0_6])])).

 cnf(c_0_38,plain,
     ( ssItem(X1) ),
     inference(spm,[status(thm)],[[19]c_0_37,[20]c_0_5])).

 cnf(c_0_39,plain,
     ( leq(X1,X2)
     | ~ lt(X1,X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[21]c_0_12,c_0_38]),c_0_38])])).

 cnf(c_0_40,plain,
     ( lt(X1,X2)
     | ~ lt(X3,X2)
     | ~ leq(X1,X3) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[22]c_0_13,c_0_38]),c_0_38]),c_0_38])])).

 cnf(c_0_41,plain,
     ( leq(X1,X2)
     | ~ lt(X3,X2)
     | ~ leq(X1,X3) ),
     inference(spm,[status(thm)],[[23]c_0_39,[24]c_0_40])).

 cnf(c_0_42,plain,
     ( lt(X1,X2)
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
     | ~ ssList(X5)
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5))) ),
     inference(er,[status(thm)],[[25]c_0_14])).

 cnf(c_0_43,plain,
     ( totalorderedP(X1)
     | ~ lt(X2,skaf65(X1))
     | ~ leq(skaf64(X1),X2)
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[26]c_0_15,[27]c_0_41])).

 cnf(c_0_44,plain,
     ( leq(X1,X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[28]c_0_16,c_0_38])])).

 cnf(c_0_45,negated_conjecture,
     ( ~ segmentP(sk4,sk1)
     | ~ totalorderedP(sk1) ),
     inference(rw,[status(thm)],[[29]c_0_17,[30]c_0_18])).

 cnf(c_0_46,negated_conjecture,
     ( segmentP(sk4,sk1) ),
     inference(rw,[status(thm)],[[31]c_0_19,[32]c_0_20])).

 cnf(c_0_47,plain,
     ( lt(X1,X2)
     | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
     | ~ ssList(X5)
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[33]c_0_42,c_0_38]),c_0_38])])).

 cnf(c_0_48,plain,
     ( totalorderedP(X1)
     | ~ lt(skaf64(X1),skaf65(X1))
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[34]c_0_43,[35]c_0_44])).

 cnf(c_0_49,negated_conjecture,
     ( ~ totalorderedP(sk1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[36]c_0_45,c_0_46])])).

 cnf(c_0_50,plain,
     ( totalorderedP(X1)
     | ~ ssList(X1)
     | ~ strictorderedP(X1) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[37]c_0_47,c_0_21]),c_0_22]),c_0_23]),c_0_24])]),[38]c_0_48])).

 cnf(c_0_51,negated_conjecture,
     ( strictorderedP(sk1) ),
     inference(rw,[status(thm)],[[39]c_0_25,[40]c_0_20])).

 cnf(c_0_52,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[41]c_0_49,c_0_50]),c_0_26]),c_0_51])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC340-1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.23  % Computer   : n170.star.cs.uiowa.edu
 % 0.03/0.23  % Model      : x86_64 x86_64
 % 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.23  % Memory     : 32218.75MB
 % 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.23  % CPULimit   : 300
 % 0.03/0.23  % DateTime   : Tue May 10 20:42:08 CDT 2016
 % 0.03/0.23  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 151.13/151.30  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
 % 151.13/151.30  # and selection function SelectCQArNTEqFirst.
 % 151.13/151.30  #
 % 151.13/151.30  # Preprocessing time       : 0.010 s
 % 151.13/151.30  # Presaturation interreduction done
 % 151.43/151.69  # No success with AutoSched0
 % 151.43/151.69  # Trying AutoSched1 for 73 seconds
 % 224.54/224.72  # AutoSched1-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S00
 % 224.54/224.72  # and selection function NoSelection.
 % 224.54/224.72  #
 % 224.54/224.72  # Preprocessing time       : 0.011 s
 % 224.54/224.72  # Presaturation interreduction done
 % 224.74/224.92  # No success with AutoSched1
 % 224.74/224.92  # Trying AutoSched2 for 23 seconds
 % 224.77/225.05  # AutoSched2-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2p
 % 224.77/225.05  # and selection function SelectCQIArNXTEqFirst.
 % 224.77/225.05  #
 % 224.77/225.05  # Preprocessing time       : 0.011 s
 % 224.77/225.05  # Presaturation interreduction done
 % 224.77/225.05
 % 224.77/225.05  # Proof found!
 % 224.77/225.05  # SZS status Unsatisfiable
 % 224.77/225.05  # SZS output start CNFRefutation
 % 224.77/225.05  cnf(c_0_0, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
 % 224.77/225.05  cnf(c_0_1, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause169)).
 % 224.77/225.05  cnf(c_0_2, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause52)).
 % 224.77/225.05  cnf(c_0_3, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause53)).
 % 224.77/225.05  cnf(c_0_4, axiom, (cons(skaf44(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause101)).
 % 224.77/225.05  cnf(c_0_5, axiom, (ssItem(skaf44(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause47)).
 % 224.77/225.05  cnf(c_0_6, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
 % 224.77/225.05  cnf(c_0_7, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
 % 224.77/225.05  cnf(c_0_8, axiom, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause138)).
 % 224.77/225.05  cnf(c_0_9, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
 % 224.77/225.05  cnf(c_0_10, axiom, (singletonP(X2)|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause116)).
 % 224.77/225.05  cnf(c_0_11, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
 % 224.77/225.05  cnf(c_0_12, axiom, (leq(X1,X2)|~lt(X1,X2)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause103)).
 % 224.77/225.05  cnf(c_0_13, axiom, (lt(X1,X3)|~leq(X1,X2)|~lt(X2,X3)|~ssItem(X3)|~ssItem(X2)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause147)).
 % 224.77/225.05  cnf(c_0_14, axiom, (lt(X2,X4)|app(app(X1,cons(X2,X3)),cons(X4,X5))!=X6|~ssList(X5)|~ssList(X3)|~ssList(X1)|~ssItem(X4)|~ssItem(X2)|~strictorderedP(X6)|~ssList(X6)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause181)).
 % 224.77/225.05  cnf(c_0_15, axiom, (totalorderedP(X1)|~leq(skaf64(X1),skaf65(X1))|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause91)).
 % 224.77/225.05  cnf(c_0_16, axiom, (leq(X1,X1)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause62)).
 % 224.77/225.05  cnf(c_0_17, negated_conjecture, (~segmentP(sk2,sk1)|~totalorderedP(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
 % 224.77/225.05  cnf(c_0_18, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
 % 224.77/225.05  cnf(c_0_19, negated_conjecture, (segmentP(sk4,sk3)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
 % 224.77/225.05  cnf(c_0_20, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
 % 224.77/225.05  cnf(c_0_21, axiom, (totalorderedP(X1)|app(app(skaf66(X1),cons(skaf64(X1),skaf67(X1))),cons(skaf65(X1),skaf68(X1)))=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause164)).
 % 224.77/225.05  cnf(c_0_22, axiom, (ssList(skaf68(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause27)).
 % 224.77/225.05  cnf(c_0_23, axiom, (ssList(skaf67(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause28)).
 % 224.77/225.05  cnf(c_0_24, axiom, (ssList(skaf66(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause29)).
 % 224.77/225.05  cnf(c_0_25, negated_conjecture, (strictorderedP(sk3)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
 % 224.77/225.05  cnf(c_0_26, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
 % 224.77/225.05  cnf(c_0_27, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_0])).
 % 224.77/225.05  cnf(c_0_28, plain, (~memberP(X1,X2)|~ssItem(X2)|~ssList(app(X1,cons(X2,X3)))|~ssList(X3)|~ssList(X1)|~duplicatefreeP(app(X1,cons(X2,X3)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_1]), c_0_2]), c_0_3])])).
 % 224.77/225.05  cnf(c_0_29, plain, (~memberP(X1,skaf44(X2))|~singletonP(X2)|~ssList(X1)|~ssList(X2)|~duplicatefreeP(app(X1,X2))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_4]), c_0_5]), c_0_6])]), c_0_7])).
 % 224.77/225.05  cnf(c_0_30, plain, (memberP(cons(X1,X2),X1)|~ssItem(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_8])).
 % 224.77/225.05  cnf(c_0_31, plain, (ssItem(X1)|~memberP(X2,skaf44(X3))|~singletonP(X3)|~ssList(X2)|~ssList(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_9]), c_0_7])).
 % 224.77/225.05  cnf(c_0_32, plain, (memberP(X1,skaf44(X1))|~singletonP(X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_4]), c_0_5]), c_0_6])])).
 % 224.77/225.05  cnf(c_0_33, plain, (singletonP(cons(X1,nil))|~ssItem(X1)|~ssList(cons(X1,nil))), inference(er,[status(thm)],[c_0_10])).
 % 224.77/225.05  cnf(c_0_34, plain, (ssItem(X1)|~singletonP(X2)|~ssList(X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
 % 224.77/225.05  cnf(c_0_35, plain, (singletonP(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_11]), c_0_6])])).
 % 224.77/225.05  cnf(c_0_36, plain, (ssItem(X1)|~ssItem(X2)|~ssList(cons(X2,nil))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
 % 224.77/225.05  cnf(c_0_37, plain, (ssItem(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_11]), c_0_6])])).
 % 224.77/225.05  cnf(c_0_38, plain, (ssItem(X1)), inference(spm,[status(thm)],[c_0_37, c_0_5])).
 % 224.77/225.05  cnf(c_0_39, plain, (leq(X1,X2)|~lt(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_38]), c_0_38])])).
 % 224.77/225.05  cnf(c_0_40, plain, (lt(X1,X2)|~lt(X3,X2)|~leq(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_38]), c_0_38]), c_0_38])])).
 % 224.77/225.05  cnf(c_0_41, plain, (leq(X1,X2)|~lt(X3,X2)|~leq(X1,X3)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
 % 224.77/225.05  cnf(c_0_42, plain, (lt(X1,X2)|~ssItem(X2)|~ssItem(X1)|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))), inference(er,[status(thm)],[c_0_14])).
 % 224.77/225.05  cnf(c_0_43, plain, (totalorderedP(X1)|~lt(X2,skaf65(X1))|~leq(skaf64(X1),X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_15, c_0_41])).
 % 224.77/225.05  cnf(c_0_44, plain, (leq(X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_38])])).
 % 224.77/225.05  cnf(c_0_45, negated_conjecture, (~segmentP(sk4,sk1)|~totalorderedP(sk1)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
 % 224.77/225.05  cnf(c_0_46, negated_conjecture, (segmentP(sk4,sk1)), inference(rw,[status(thm)],[c_0_19, c_0_20])).
 % 224.77/225.05  cnf(c_0_47, plain, (lt(X1,X2)|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~strictorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_38]), c_0_38])])).
 % 224.77/225.05  cnf(c_0_48, plain, (totalorderedP(X1)|~lt(skaf64(X1),skaf65(X1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
 % 224.77/225.05  cnf(c_0_49, negated_conjecture, (~totalorderedP(sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46])])).
 % 224.77/225.05  cnf(c_0_50, plain, (totalorderedP(X1)|~ssList(X1)|~strictorderedP(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_21]), c_0_22]), c_0_23]), c_0_24])]), c_0_48])).
 % 224.77/225.05  cnf(c_0_51, negated_conjecture, (strictorderedP(sk1)), inference(rw,[status(thm)],[c_0_25, c_0_20])).
 % 224.77/225.05  cnf(c_0_52, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_26]), c_0_51])]), ['proof']).
 % 224.77/225.05  # SZS output end CNFRefutation
 % 224.77/225.05  # Training examples: 0 positive, 0 negative
 % 224.77/225.05
 % 224.77/225.05  # -------------------------------------------------
 % 224.77/225.05  # User time                : 219.050 s
 % 224.77/225.05  # System time              : 5.799 s
 % 224.77/225.05  # Total time               : 224.849 s
 % 224.77/225.05  # Maximum resident set size: 1856 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
