   --------------------------------------------------------------------------

                   TSTP Solution File: SWC379+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC379+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n138.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:43 EDT 2016

 % Result     : Theorem 0.03s
 % Output     : CNFRefutation 0.03s
 % Verified   :
 % Statistics : -

 % Comments   :
 %------------------------------------------------------------------------------
 %----WARNING: Could not form TPTP format derivation
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC379+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.22  % Computer   : n138.star.cs.uiowa.edu
 % 0.03/0.22  % Model      : x86_64 x86_64
 % 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.22  % Memory     : 32218.75MB
 % 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.22  % CPULimit   : 300
 % 0.03/0.22  % DateTime   : Tue May 10 20:42:53 CDT 2016
 % 0.03/0.23  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.03/0.28  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.03/0.28  # and selection function SelectCQPrecWNTNp.
 % 0.03/0.28  #
 % 0.03/0.28  # Preprocessing time       : 0.014 s
 % 0.03/0.28  # Presaturation interreduction done
 % 0.03/0.28
 % 0.03/0.28  # Proof found!
 % 0.03/0.28  # SZS status Theorem
 % 0.03/0.28  # SZS output start CNFRefutation
 % 0.03/0.28  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|?[X5]:(ssItem(X5)&(((~(memberP(X3,X5))&![X6]:(ssItem(X6)=>((~(memberP(X4,X6))|~(leq(X6,X5)))|X5=X6)))&memberP(X4,X5))|(memberP(X3,X5)&(~(memberP(X4,X5))|?[X6]:(((ssItem(X6)&X5!=X6)&memberP(X4,X6))&leq(X6,X5)))))))|![X7]:(ssItem(X7)=>((~(memberP(X1,X7))&(~(memberP(X2,X7))|?[X8]:(((ssItem(X8)&X7!=X8)&memberP(X2,X8))&leq(X8,X7))))|((![X8]:(ssItem(X8)=>((~(memberP(X2,X8))|~(leq(X8,X7)))|X7=X8))&memberP(X2,X7))&memberP(X1,X7)))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
 % 0.03/0.28  fof(c_0_1, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|?[X5]:(ssItem(X5)&(((~memberP(X3,X5)&![X6]:(ssItem(X6)=>((~memberP(X4,X6)|~leq(X6,X5))|X5=X6)))&memberP(X4,X5))|(memberP(X3,X5)&(~memberP(X4,X5)|?[X6]:(((ssItem(X6)&X5!=X6)&memberP(X4,X6))&leq(X6,X5)))))))|![X7]:(ssItem(X7)=>((~memberP(X1,X7)&(~memberP(X2,X7)|?[X8]:(((ssItem(X8)&X7!=X8)&memberP(X2,X8))&leq(X8,X7))))|((![X8]:(ssItem(X8)=>((~memberP(X2,X8)|~leq(X8,X7))|X7=X8))&memberP(X2,X7))&memberP(X1,X7))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.03/0.28  fof(c_0_2, negated_conjecture,
 % (![X13]:![X15]:![X17]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((esk49_0=esk51_0&esk48_0=esk50_0)&(((((ssItem(esk52_1(X13))|memberP(esk50_0,X13))|~memberP(esk51_0,X13))|~ssItem(X13))&(((((memberP(esk51_0,esk52_1(X13))|memberP(esk50_0,X13))|~memberP(esk51_0,X13))|~ssItem(X13))&(((leq(esk52_1(X13),X13)|memberP(esk50_0,X13))|~memberP(esk51_0,X13))|~ssItem(X13)))&(((X13!=esk52_1(X13)|memberP(esk50_0,X13))|~memberP(esk51_0,X13))|~ssItem(X13))))&(((memberP(esk51_0,X13)|~memberP(esk50_0,X13))|~ssItem(X13))&(((((~ssItem(X15)|X13=X15)|~memberP(esk51_0,X15))|~leq(X15,X13))|~memberP(esk50_0,X13))|~ssItem(X13)))))&(ssItem(esk53_0)&(((memberP(esk49_0,esk53_0)|memberP(esk48_0,esk53_0))&((((~ssItem(X17)|esk53_0=X17)|~memberP(esk49_0,X17))|~leq(X17,esk53_0))|memberP(esk48_0,esk53_0)))&(((ssItem(esk54_0)|~memberP(esk49_0,esk53_0))|~memberP(esk48_0,esk53_0))&((((memberP(esk49_0,esk54_0)|~memberP(esk49_0,esk53_0))|~memberP(esk48_0,esk53_0))&((leq(esk54_0,esk53_0)|~memberP(esk49_0,esk53_0))|~memberP(e
 % sk48_0,esk53_0)))&((esk53_0!=esk54_0|~memberP(esk49_0,esk53_0))|~memberP(esk48_0,esk53_0)))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
 % 0.03/0.28  cnf(c_0_3, negated_conjecture, (memberP(esk51_0,X1)|~ssItem(X1)|~memberP(esk50_0,X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_4, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_5, negated_conjecture, (memberP(esk48_0,esk53_0)|memberP(esk49_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_6, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_7, negated_conjecture, (memberP(esk50_0,X1)|memberP(esk51_0,esk52_1(X1))|~ssItem(X1)|~memberP(esk51_0,X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_8, negated_conjecture, (memberP(esk51_0,X1)|~memberP(esk48_0,X1)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_3, c_0_4])).
 % 0.03/0.28  cnf(c_0_9, negated_conjecture, (memberP(esk51_0,esk53_0)|memberP(esk48_0,esk53_0)), inference(rw,[status(thm)],[c_0_5, c_0_6])).
 % 0.03/0.28  cnf(c_0_10, negated_conjecture, (ssItem(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_11, negated_conjecture, (memberP(esk50_0,X1)|ssItem(esk52_1(X1))|~ssItem(X1)|~memberP(esk51_0,X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_12, negated_conjecture, (memberP(esk50_0,X1)|~ssItem(X1)|~memberP(esk51_0,X1)|X1!=esk52_1(X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_13, negated_conjecture, (memberP(esk48_0,esk53_0)|esk53_0=X1|~leq(X1,esk53_0)|~memberP(esk49_0,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_14, negated_conjecture, (memberP(esk51_0,esk52_1(X1))|memberP(esk48_0,X1)|~memberP(esk51_0,X1)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_7, c_0_4])).
 % 0.03/0.28  cnf(c_0_15, negated_conjecture, (memberP(esk51_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])])).
 % 0.03/0.28  cnf(c_0_16, negated_conjecture, (memberP(esk48_0,X1)|ssItem(esk52_1(X1))|~memberP(esk51_0,X1)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_11, c_0_4])).
 % 0.03/0.28  cnf(c_0_17, negated_conjecture, (memberP(esk48_0,X1)|esk52_1(X1)!=X1|~memberP(esk51_0,X1)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_12, c_0_4])).
 % 0.03/0.28  cnf(c_0_18, negated_conjecture, (memberP(esk50_0,X1)|leq(esk52_1(X1),X1)|~ssItem(X1)|~memberP(esk51_0,X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_19, negated_conjecture, (esk53_0=X1|memberP(esk48_0,esk53_0)|~leq(X1,esk53_0)|~memberP(esk51_0,X1)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_13, c_0_6])).
 % 0.03/0.28  cnf(c_0_20, negated_conjecture, (memberP(esk51_0,esk52_1(esk53_0))|memberP(esk48_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_10])])).
 % 0.03/0.28  cnf(c_0_21, negated_conjecture, (memberP(esk48_0,esk53_0)|ssItem(esk52_1(esk53_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_15]), c_0_10])])).
 % 0.03/0.28  cnf(c_0_22, negated_conjecture, (memberP(esk48_0,esk53_0)|esk52_1(esk53_0)!=esk53_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_15]), c_0_10])])).
 % 0.03/0.28  cnf(c_0_23, negated_conjecture, (memberP(esk49_0,esk54_0)|~memberP(esk48_0,esk53_0)|~memberP(esk49_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_24, negated_conjecture, (leq(esk54_0,esk53_0)|~memberP(esk48_0,esk53_0)|~memberP(esk49_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_25, negated_conjecture, (ssItem(esk54_0)|~memberP(esk48_0,esk53_0)|~memberP(esk49_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_26, negated_conjecture, (~memberP(esk48_0,esk53_0)|~memberP(esk49_0,esk53_0)|esk53_0!=esk54_0), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_27, negated_conjecture, (X1=X2|~ssItem(X1)|~memberP(esk50_0,X1)|~leq(X2,X1)|~memberP(esk51_0,X2)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.03/0.28  cnf(c_0_28, negated_conjecture, (leq(esk52_1(X1),X1)|memberP(esk48_0,X1)|~memberP(esk51_0,X1)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_18, c_0_4])).
 % 0.03/0.28  cnf(c_0_29, negated_conjecture, (memberP(esk48_0,esk53_0)|~leq(esk52_1(esk53_0),esk53_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22])).
 % 0.03/0.28  cnf(c_0_30, negated_conjecture, (memberP(esk51_0,esk54_0)|~memberP(esk48_0,esk53_0)|~memberP(esk51_0,esk53_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_6]), c_0_6])).
 % 0.03/0.28  cnf(c_0_31, negated_conjecture, (leq(esk54_0,esk53_0)|~memberP(esk48_0,esk53_0)|~memberP(esk51_0,esk53_0)), inference(rw,[status(thm)],[c_0_24, c_0_6])).
 % 0.03/0.28  cnf(c_0_32, negated_conjecture, (ssItem(esk54_0)|~memberP(esk48_0,esk53_0)|~memberP(esk51_0,esk53_0)), inference(rw,[status(thm)],[c_0_25, c_0_6])).
 % 0.03/0.28  cnf(c_0_33, negated_conjecture, (esk54_0!=esk53_0|~memberP(esk48_0,esk53_0)|~memberP(esk51_0,esk53_0)), inference(rw,[status(thm)],[c_0_26, c_0_6])).
 % 0.03/0.28  cnf(c_0_34, negated_conjecture, (X1=X2|~leq(X2,X1)|~memberP(esk48_0,X1)|~memberP(esk51_0,X2)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_27, c_0_4])).
 % 0.03/0.28  cnf(c_0_35, negated_conjecture, (memberP(esk48_0,esk53_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_15]), c_0_10])]), c_0_29])).
 % 0.03/0.28  cnf(c_0_36, negated_conjecture, (memberP(esk51_0,esk54_0)|~memberP(esk48_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_15])])).
 % 0.03/0.28  cnf(c_0_37, negated_conjecture, (leq(esk54_0,esk53_0)|~memberP(esk48_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_15])])).
 % 0.03/0.28  cnf(c_0_38, negated_conjecture, (ssItem(esk54_0)|~memberP(esk48_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_15])])).
 % 0.03/0.28  cnf(c_0_39, negated_conjecture, (esk54_0!=esk53_0|~memberP(esk48_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_15])])).
 % 0.03/0.28  cnf(c_0_40, negated_conjecture, (esk53_0=X1|~leq(X1,esk53_0)|~memberP(esk51_0,X1)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_10])])).
 % 0.03/0.28  cnf(c_0_41, negated_conjecture, (memberP(esk51_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_35])])).
 % 0.03/0.28  cnf(c_0_42, negated_conjecture, (leq(esk54_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_35])])).
 % 0.03/0.28  cnf(c_0_43, negated_conjecture, (ssItem(esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_35])])).
 % 0.03/0.28  cnf(c_0_44, negated_conjecture, (esk54_0!=esk53_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_35])])).
 % 0.03/0.28  cnf(c_0_45, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_43])]), c_0_44]), ['proof']).
 % 0.03/0.28  # SZS output end CNFRefutation
 % 0.03/0.28  # Training examples: 0 positive, 0 negative
 % 0.03/0.28
 % 0.03/0.28  # -------------------------------------------------
 % 0.03/0.28  # User time                : 0.048 s
 % 0.03/0.28  # System time              : 0.006 s
 % 0.03/0.28  # Total time               : 0.054 s
 % 0.03/0.28  # Maximum resident set size: 1880 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
