   --------------------------------------------------------------------------

                   TSTP Solution File: SWC327+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC327+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n128.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:30 EDT 2016

 % Result     : Theorem 0.02s
 % Output     : CNFRefutation 0.02s
 % Verified   :
 % Statistics : Number of formulae       :   66 (2657 expanded)
 %              Number of clauses        :   54 ( 945 expanded)
 %              Number of leaves         :    6 ( 581 expanded)
 %              Depth                    :   16
 %              Number of atoms          :  299 (67899 expanded)
 %              Number of equality atoms :  149 (31838 expanded)
 %              Maximal formula depth    :   30 (   5 average)
 %              Maximal clause size      :   66 (   3 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ! [X5] :
                         ( ssList(X5)
                        => ( app(X3,X5) != X4
                           | ~ equalelemsP(X3)
                           | ? [X6] :
                               ( ssItem(X6)
                               & ? [X7] :
                                   ( ssList(X7)
                                   & app(cons(X6,nil),X7) = X5
                                   & ? [X8] :
                                       ( ssList(X8)
                                       & app(X8,cons(X6,nil)) = X3 ) ) ) ) )
                     | ( nil != X4
                       & nil = X3 )
                     | ( ? [X9] :
                           ( ssList(X9)
                           & app(X1,X9) = X2
                           & ! [X10] :
                               ( ssItem(X10)
                              => ! [X11] :
                                   ( ssList(X11)
                                  => ( app(cons(X10,nil),X11) != X9
                                     | ! [X12] :
                                         ( ssList(X12)
                                        => app(X12,cons(X10,nil)) != X1 ) ) ) )
                           & equalelemsP(X1) )
                       & ( nil != X1
                         | nil = X2 ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(X1,nil) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

 fof(c_0_2,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) != X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax18)).

 fof(c_0_6,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ! [X5] :
                           ( ssList(X5)
                          => ( app(X3,X5) != X4
                             | ~ equalelemsP(X3)
                             | ? [X6] :
                                 ( ssItem(X6)
                                 & ? [X7] :
                                     ( ssList(X7)
                                     & app(cons(X6,nil),X7) = X5
                                     & ? [X8] :
                                         ( ssList(X8)
                                         & app(X8,cons(X6,nil)) = X3 ) ) ) ) )
                       | ( nil != X4
                         & nil = X3 )
                       | ( ? [X9] :
                             ( ssList(X9)
                             & app(X1,X9) = X2
                             & ! [X10] :
                                 ( ssItem(X10)
                                => ! [X11] :
                                     ( ssList(X11)
                                    => ( app(cons(X10,nil),X11) != X9
                                       | ! [X12] :
                                           ( ssList(X12)
                                          => app(X12,cons(X10,nil)) != X1 ) ) ) )
                             & equalelemsP(X1) )
                         & ( nil != X1
                           | nil = X2 ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_7,negated_conjecture,(
     ! [X18,X19,X20,X21] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ssList(esk52_0)
       & app(esk50_0,esk52_0) = esk51_0
       & equalelemsP(esk50_0)
       & ( ~ ssItem(X18)
         | ~ ssList(X19)
         | app(cons(X18,nil),X19) != esk52_0
         | ~ ssList(X20)
         | app(X20,cons(X18,nil)) != esk50_0 )
       & ( nil = esk51_0
         | nil != esk50_0 )
       & ( nil = esk48_0
         | ssItem(esk53_1(X21))
         | ~ ssList(X21)
         | app(esk48_0,X21) != esk49_0
         | ~ equalelemsP(esk48_0) )
       & ( nil != esk49_0
         | ssItem(esk53_1(X21))
         | ~ ssList(X21)
         | app(esk48_0,X21) != esk49_0
         | ~ equalelemsP(esk48_0) )
       & ( nil = esk48_0
         | ssList(esk54_1(X21))
         | ~ ssList(X21)
         | app(esk48_0,X21) != esk49_0
         | ~ equalelemsP(esk48_0) )
       & ( nil != esk49_0
         | ssList(esk54_1(X21))
         | ~ ssList(X21)
         | app(esk48_0,X21) != esk49_0
         | ~ equalelemsP(esk48_0) )
       & ( nil = esk48_0
         | app(cons(esk53_1(X21),nil),esk54_1(X21)) = X21
         | ~ ssList(X21)
         | app(esk48_0,X21) != esk49_0
         | ~ equalelemsP(esk48_0) )
       & ( nil != esk49_0
         | app(cons(esk53_1(X21),nil),esk54_1(X21)) = X21
         | ~ ssList(X21)
         | app(esk48_0,X21) != esk49_0
         | ~ equalelemsP(esk48_0) )
       & ( nil = esk48_0
         | ssList(esk55_1(X21))
         | ~ ssList(X21)
         | app(esk48_0,X21) != esk49_0
         | ~ equalelemsP(esk48_0) )
       & ( nil != esk49_0
         | ssList(esk55_1(X21))
         | ~ ssList(X21)
         | app(esk48_0,X21) != esk49_0
         | ~ equalelemsP(esk48_0) )
       & ( nil = esk48_0
         | app(esk55_1(X21),cons(esk53_1(X21),nil)) = esk48_0
         | ~ ssList(X21)
         | app(esk48_0,X21) != esk49_0
         | ~ equalelemsP(esk48_0) )
       & ( nil != esk49_0
         | app(esk55_1(X21),cons(esk53_1(X21),nil)) = esk48_0
         | ~ ssList(X21)
         | app(esk48_0,X21) != esk49_0
         | ~ equalelemsP(esk48_0) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])])).

 cnf(c_0_8,negated_conjecture,
     ( equalelemsP(esk50_0) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_7])).

 cnf(c_0_9,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_7])).

 cnf(c_0_10,negated_conjecture,
     ( app(esk55_1(X1),cons(esk53_1(X1),nil)) = esk48_0
     | nil = esk48_0
     | ~ equalelemsP(esk48_0)
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[6]c_0_7])).

 cnf(c_0_11,negated_conjecture,
     ( equalelemsP(esk48_0) ),
     inference(rw,[status(thm)],[[7]c_0_8,[8]c_0_9])).

 cnf(c_0_12,negated_conjecture,
     ( app(esk50_0,esk52_0) = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[9]c_0_7])).

 cnf(c_0_13,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[10]c_0_7])).

 cnf(c_0_14,negated_conjecture,
     ( ssItem(esk53_1(X1))
     | nil = esk48_0
     | ~ equalelemsP(esk48_0)
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_7])).

 cnf(c_0_15,negated_conjecture,
     ( ssList(esk55_1(X1))
     | nil = esk48_0
     | ~ equalelemsP(esk48_0)
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_7])).

 cnf(c_0_16,negated_conjecture,
     ( app(X1,cons(X2,nil)) != esk50_0
     | ~ ssList(X1)
     | app(cons(X2,nil),X3) != esk52_0
     | ~ ssList(X3)
     | ~ ssItem(X2) ),
     inference(split_conjunct,[status(thm)],[[13]c_0_7])).

 cnf(c_0_17,negated_conjecture,
     ( app(esk55_1(X1),cons(esk53_1(X1),nil)) = esk48_0
     | esk48_0 = nil
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[14]c_0_10,c_0_11])])).

 cnf(c_0_18,negated_conjecture,
     ( app(esk48_0,esk52_0) = esk49_0 ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[15]c_0_12,c_0_9]),[16]c_0_13])).

 cnf(c_0_19,negated_conjecture,
     ( ssList(esk52_0) ),
     inference(split_conjunct,[status(thm)],[[17]c_0_7])).

 cnf(c_0_20,negated_conjecture,
     ( esk48_0 = nil
     | ssItem(esk53_1(X1))
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[18]c_0_14,c_0_11])])).

 cnf(c_0_21,negated_conjecture,
     ( esk48_0 = nil
     | ssList(esk55_1(X1))
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[19]c_0_15,c_0_11])])).

 cnf(c_0_22,negated_conjecture,
     ( app(cons(esk53_1(X1),nil),esk54_1(X1)) = X1
     | nil = esk48_0
     | ~ equalelemsP(esk48_0)
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[20]c_0_7])).

 cnf(c_0_23,negated_conjecture,
     ( ssList(esk54_1(X1))
     | nil = esk48_0
     | ~ equalelemsP(esk48_0)
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[21]c_0_7])).

 cnf(c_0_24,negated_conjecture,
     ( app(X1,cons(X2,nil)) != esk48_0
     | app(cons(X2,nil),X3) != esk52_0
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(rw,[status(thm)],[[22]c_0_16,[23]c_0_9])).

 cnf(c_0_25,negated_conjecture,
     ( app(esk55_1(esk52_0),cons(esk53_1(esk52_0),nil)) = esk48_0
     | esk48_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[24]c_0_17,c_0_18]),c_0_19])])).

 cnf(c_0_26,negated_conjecture,
     ( esk48_0 = nil
     | ssItem(esk53_1(esk52_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_20,c_0_18]),c_0_19])])).

 cnf(c_0_27,negated_conjecture,
     ( esk48_0 = nil
     | ssList(esk55_1(esk52_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_21,c_0_18]),c_0_19])])).

 cnf(c_0_28,negated_conjecture,
     ( app(cons(esk53_1(X1),nil),esk54_1(X1)) = X1
     | esk48_0 = nil
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[27]c_0_22,c_0_11])])).

 cnf(c_0_29,negated_conjecture,
     ( esk48_0 = nil
     | ssList(esk54_1(X1))
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[28]c_0_23,c_0_11])])).

 cnf(c_0_30,negated_conjecture,
     ( app(esk55_1(X1),cons(esk53_1(X1),nil)) = esk48_0
     | ~ equalelemsP(esk48_0)
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1)
     | nil != esk49_0 ),
     inference(split_conjunct,[status(thm)],[[29]c_0_7])).

 fof(c_0_31,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(X2,nil) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).

 cnf(c_0_32,negated_conjecture,
     ( nil = esk51_0
     | nil != esk50_0 ),
     inference(split_conjunct,[status(thm)],[[30]c_0_7])).

 cnf(c_0_33,negated_conjecture,
     ( esk48_0 = nil
     | app(cons(esk53_1(esk52_0),nil),X1) != esk52_0
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[31]c_0_24,c_0_25]),c_0_26]),[32]c_0_27])).

 cnf(c_0_34,negated_conjecture,
     ( app(cons(esk53_1(esk52_0),nil),esk54_1(esk52_0)) = esk52_0
     | esk48_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[33]c_0_28,c_0_18]),c_0_19])])).

 cnf(c_0_35,negated_conjecture,
     ( esk48_0 = nil
     | ssList(esk54_1(esk52_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[34]c_0_29,c_0_18]),c_0_19])])).

 cnf(c_0_36,negated_conjecture,
     ( ssList(esk55_1(X1))
     | ~ equalelemsP(esk48_0)
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1)
     | nil != esk49_0 ),
     inference(split_conjunct,[status(thm)],[[35]c_0_7])).

 cnf(c_0_37,negated_conjecture,
     ( app(esk55_1(X1),cons(esk53_1(X1),nil)) = esk48_0
     | app(esk48_0,X1) != esk49_0
     | esk49_0 != nil
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[36]c_0_30,c_0_11])])).

 cnf(c_0_38,plain,
     ( app(X1,nil) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[37]c_0_31])).

 cnf(c_0_39,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[38]c_0_2])).

 cnf(c_0_40,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[39]c_0_7])).

 cnf(c_0_41,negated_conjecture,
     ( esk49_0 = nil
     | esk48_0 != nil ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[40]c_0_32,c_0_13]),[41]c_0_9])).

 cnf(c_0_42,negated_conjecture,
     ( esk48_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[42]c_0_33,c_0_34]),[43]c_0_35])).

 cnf(c_0_43,negated_conjecture,
     ( ssList(esk55_1(X1))
     | app(esk48_0,X1) != esk49_0
     | esk49_0 != nil
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[44]c_0_36,c_0_11])])).

 cnf(c_0_44,negated_conjecture,
     ( ssItem(esk53_1(X1))
     | ~ equalelemsP(esk48_0)
     | app(esk48_0,X1) != esk49_0
     | ~ ssList(X1)
     | nil != esk49_0 ),
     inference(split_conjunct,[status(thm)],[[45]c_0_7])).

 fof(c_0_45,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).

 cnf(c_0_46,negated_conjecture,
     ( app(esk55_1(nil),cons(esk53_1(nil),nil)) = esk48_0
     | esk49_0 != esk48_0
     | esk49_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[46]c_0_37,c_0_38]),c_0_39]),c_0_40])])).

 cnf(c_0_47,negated_conjecture,
     ( esk49_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[47]c_0_41,c_0_42])])).

 cnf(c_0_48,negated_conjecture,
     ( ssList(esk55_1(nil))
     | esk49_0 != esk48_0
     | esk49_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[48]c_0_43,c_0_38]),c_0_39]),c_0_40])])).

 cnf(c_0_49,negated_conjecture,
     ( ssItem(esk53_1(X1))
     | app(esk48_0,X1) != esk49_0
     | esk49_0 != nil
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[49]c_0_44,c_0_11])])).

 cnf(c_0_50,plain,
     ( nil = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[50]c_0_45])).

 cnf(c_0_51,negated_conjecture,
     ( app(esk55_1(nil),cons(esk53_1(nil),nil)) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[51]c_0_46,c_0_42]),c_0_42])]),c_0_47])])).

 cnf(c_0_52,negated_conjecture,
     ( ssList(esk55_1(nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[[52]c_0_48,c_0_42])]),c_0_47])])).

 fof(c_0_53,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

 cnf(c_0_54,negated_conjecture,
     ( ssItem(esk53_1(nil))
     | esk49_0 != esk48_0
     | esk49_0 != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[53]c_0_49,c_0_38]),c_0_39]),c_0_40])])).

 cnf(c_0_55,negated_conjecture,
     ( esk55_1(nil) = nil
     | ~ ssList(cons(esk53_1(nil),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[54]c_0_50,c_0_51]),c_0_52])])).

 cnf(c_0_56,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[55]c_0_53])).

 cnf(c_0_57,negated_conjecture,
     ( ssItem(esk53_1(nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[[56]c_0_54,c_0_42])]),c_0_47])])).

 cnf(c_0_58,negated_conjecture,
     ( esk55_1(nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[57]c_0_55,c_0_56]),c_0_39]),c_0_57])])).

 cnf(c_0_59,plain,
     ( nil = X2
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[58]c_0_45])).

 cnf(c_0_60,negated_conjecture,
     ( app(nil,cons(esk53_1(nil),nil)) = nil ),
     inference(rw,[status(thm)],[[59]c_0_51,[60]c_0_58])).

 fof(c_0_61,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) != X3 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

 cnf(c_0_62,negated_conjecture,
     ( cons(esk53_1(nil),nil) = nil
     | ~ ssList(cons(esk53_1(nil),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[61]c_0_59,c_0_60]),c_0_39])])).

 cnf(c_0_63,plain,
     ( cons(X1,X2) != X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[62]c_0_61])).

 cnf(c_0_64,negated_conjecture,
     ( cons(esk53_1(nil),nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[63]c_0_62,c_0_56]),c_0_39]),c_0_57])])).

 cnf(c_0_65,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[64]c_0_63,c_0_64]),c_0_39]),c_0_57])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC327+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.23  % Computer   : n128.star.cs.uiowa.edu
 % 0.02/0.23  % Model      : x86_64 x86_64
 % 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.23  % Memory     : 32218.75MB
 % 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.23  % CPULimit   : 300
 % 0.02/0.23  % DateTime   : Tue May 10 20:41:22 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.02/0.27  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.02/0.27  # and selection function SelectCQPrecWNTNp.
 % 0.02/0.27  #
 % 0.02/0.27  # Preprocessing time       : 0.014 s
 % 0.02/0.27  # Presaturation interreduction done
 % 0.02/0.27
 % 0.02/0.27  # Proof found!
 % 0.02/0.27  # SZS status Theorem
 % 0.02/0.27  # SZS output start CNFRefutation
 % 0.02/0.27  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>((app(X3,X5)!=X4|~(equalelemsP(X3)))|?[X6]:(ssItem(X6)&?[X7]:((ssList(X7)&app(cons(X6,nil),X7)=X5)&?[X8]:(ssList(X8)&app(X8,cons(X6,nil))=X3))))))|(nil!=X4&nil=X3))|(?[X9]:(((ssList(X9)&app(X1,X9)=X2)&![X10]:(ssItem(X10)=>![X11]:(ssList(X11)=>(app(cons(X10,nil),X11)!=X9|![X12]:(ssList(X12)=>app(X12,cons(X10,nil))!=X1)))))&equalelemsP(X1))&(nil!=X1|nil=X2)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.02/0.27  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
 % 0.02/0.27  fof(c_0_2, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.02/0.27  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 0.02/0.27  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.02/0.27  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax18)).
 % 0.02/0.27  fof(c_0_6, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>((app(X3,X5)!=X4|~equalelemsP(X3))|?[X6]:(ssItem(X6)&?[X7]:((ssList(X7)&app(cons(X6,nil),X7)=X5)&?[X8]:(ssList(X8)&app(X8,cons(X6,nil))=X3))))))|(nil!=X4&nil=X3))|(?[X9]:(((ssList(X9)&app(X1,X9)=X2)&![X10]:(ssItem(X10)=>![X11]:(ssList(X11)=>(app(cons(X10,nil),X11)!=X9|![X12]:(ssList(X12)=>app(X12,cons(X10,nil))!=X1)))))&equalelemsP(X1))&(nil!=X1|nil=X2))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.02/0.27  fof(c_0_7, negated_conjecture,
 % (![X18]:![X19]:![X20]:![X21]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(ssList(esk52_0)&((app(esk50_0,esk52_0)=esk51_0&equalelemsP(esk50_0))&(~ssItem(X18)|((~ssList(X19)|app(cons(X18,nil),X19)!=esk52_0)|(~ssList(X20)|app(X20,cons(X18,nil))!=esk50_0))))))&(nil=esk51_0|nil!=esk50_0))&(((nil=esk48_0|((ssItem(esk53_1(X21))|(~ssList(X21)|app(esk48_0,X21)!=esk49_0))|~equalelemsP(esk48_0)))&(nil!=esk49_0|((ssItem(esk53_1(X21))|(~ssList(X21)|app(esk48_0,X21)!=esk49_0))|~equalelemsP(esk48_0))))&(((nil=esk48_0|((ssList(esk54_1(X21))|(~ssList(X21)|app(esk48_0,X21)!=esk49_0))|~equalelemsP(esk48_0)))&(nil!=esk49_0|((ssList(esk54_1(X21))|(~ssList(X21)|app(esk48_0,X21)!=esk49_0))|~equalelemsP(esk48_0))))&(((nil=esk48_0|((app(cons(esk53_1(X21),nil),esk54_1(X21))=X21|(~ssList(X21)|app(esk48_0,X21)!=esk49_0))|~equalelemsP(esk48_0)))&(nil!=esk49_0|((app(cons(esk53_1(X21),nil),esk54_1(X21))=X21|(~ssList(X21)|app(esk48_0,X21)!=esk49_0))|~equalelemsP(esk48_0))))&(((nil
 % =esk48_0|((ssList(esk55_1(X21))|(~ssList(X21)|app(esk48_0,X21)!=esk49_0))|~equalelemsP(esk48_0)))&(nil!=esk49_0|((ssList(esk55_1(X21))|(~ssList(X21)|app(esk48_0,X21)!=esk49_0))|~equalelemsP(esk48_0))))&((nil=esk48_0|((app(esk55_1(X21),cons(esk53_1(X21),nil))=esk48_0|(~ssList(X21)|app(esk48_0,X21)!=esk49_0))|~equalelemsP(esk48_0)))&(nil!=esk49_0|((app(esk55_1(X21),cons(esk53_1(X21),nil))=esk48_0|(~ssList(X21)|app(esk48_0,X21)!=esk49_0))|~equalelemsP(esk48_0)))))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])])).
 % 0.02/0.27  cnf(c_0_8, negated_conjecture, (equalelemsP(esk50_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_9, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_10, negated_conjecture, (app(esk55_1(X1),cons(esk53_1(X1),nil))=esk48_0|nil=esk48_0|~equalelemsP(esk48_0)|app(esk48_0,X1)!=esk49_0|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_11, negated_conjecture, (equalelemsP(esk48_0)), inference(rw,[status(thm)],[c_0_8, c_0_9])).
 % 0.02/0.27  cnf(c_0_12, negated_conjecture, (app(esk50_0,esk52_0)=esk51_0), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_13, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_14, negated_conjecture, (ssItem(esk53_1(X1))|nil=esk48_0|~equalelemsP(esk48_0)|app(esk48_0,X1)!=esk49_0|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_15, negated_conjecture, (ssList(esk55_1(X1))|nil=esk48_0|~equalelemsP(esk48_0)|app(esk48_0,X1)!=esk49_0|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_16, negated_conjecture, (app(X1,cons(X2,nil))!=esk50_0|~ssList(X1)|app(cons(X2,nil),X3)!=esk52_0|~ssList(X3)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_17, negated_conjecture, (app(esk55_1(X1),cons(esk53_1(X1),nil))=esk48_0|esk48_0=nil|app(esk48_0,X1)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_11])])).
 % 0.02/0.27  cnf(c_0_18, negated_conjecture, (app(esk48_0,esk52_0)=esk49_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_9]), c_0_13])).
 % 0.02/0.27  cnf(c_0_19, negated_conjecture, (ssList(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_20, negated_conjecture, (esk48_0=nil|ssItem(esk53_1(X1))|app(esk48_0,X1)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_11])])).
 % 0.02/0.27  cnf(c_0_21, negated_conjecture, (esk48_0=nil|ssList(esk55_1(X1))|app(esk48_0,X1)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_11])])).
 % 0.02/0.27  cnf(c_0_22, negated_conjecture, (app(cons(esk53_1(X1),nil),esk54_1(X1))=X1|nil=esk48_0|~equalelemsP(esk48_0)|app(esk48_0,X1)!=esk49_0|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_23, negated_conjecture, (ssList(esk54_1(X1))|nil=esk48_0|~equalelemsP(esk48_0)|app(esk48_0,X1)!=esk49_0|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_24, negated_conjecture, (app(X1,cons(X2,nil))!=esk48_0|app(cons(X2,nil),X3)!=esk52_0|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(rw,[status(thm)],[c_0_16, c_0_9])).
 % 0.02/0.27  cnf(c_0_25, negated_conjecture, (app(esk55_1(esk52_0),cons(esk53_1(esk52_0),nil))=esk48_0|esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
 % 0.02/0.27  cnf(c_0_26, negated_conjecture, (esk48_0=nil|ssItem(esk53_1(esk52_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_18]), c_0_19])])).
 % 0.02/0.27  cnf(c_0_27, negated_conjecture, (esk48_0=nil|ssList(esk55_1(esk52_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_18]), c_0_19])])).
 % 0.02/0.27  cnf(c_0_28, negated_conjecture, (app(cons(esk53_1(X1),nil),esk54_1(X1))=X1|esk48_0=nil|app(esk48_0,X1)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_11])])).
 % 0.02/0.27  cnf(c_0_29, negated_conjecture, (esk48_0=nil|ssList(esk54_1(X1))|app(esk48_0,X1)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_11])])).
 % 0.02/0.27  cnf(c_0_30, negated_conjecture, (app(esk55_1(X1),cons(esk53_1(X1),nil))=esk48_0|~equalelemsP(esk48_0)|app(esk48_0,X1)!=esk49_0|~ssList(X1)|nil!=esk49_0), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  fof(c_0_31, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
 % 0.02/0.27  cnf(c_0_32, negated_conjecture, (nil=esk51_0|nil!=esk50_0), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_33, negated_conjecture, (esk48_0=nil|app(cons(esk53_1(esk52_0),nil),X1)!=esk52_0|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27])).
 % 0.02/0.27  cnf(c_0_34, negated_conjecture, (app(cons(esk53_1(esk52_0),nil),esk54_1(esk52_0))=esk52_0|esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_18]), c_0_19])])).
 % 0.02/0.27  cnf(c_0_35, negated_conjecture, (esk48_0=nil|ssList(esk54_1(esk52_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_18]), c_0_19])])).
 % 0.02/0.27  cnf(c_0_36, negated_conjecture, (ssList(esk55_1(X1))|~equalelemsP(esk48_0)|app(esk48_0,X1)!=esk49_0|~ssList(X1)|nil!=esk49_0), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_37, negated_conjecture, (app(esk55_1(X1),cons(esk53_1(X1),nil))=esk48_0|app(esk48_0,X1)!=esk49_0|esk49_0!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_11])])).
 % 0.02/0.27  cnf(c_0_38, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
 % 0.02/0.27  cnf(c_0_39, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.02/0.27  cnf(c_0_40, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  cnf(c_0_41, negated_conjecture, (esk49_0=nil|esk48_0!=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_13]), c_0_9])).
 % 0.02/0.27  cnf(c_0_42, negated_conjecture, (esk48_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])).
 % 0.02/0.27  cnf(c_0_43, negated_conjecture, (ssList(esk55_1(X1))|app(esk48_0,X1)!=esk49_0|esk49_0!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_11])])).
 % 0.02/0.27  cnf(c_0_44, negated_conjecture, (ssItem(esk53_1(X1))|~equalelemsP(esk48_0)|app(esk48_0,X1)!=esk49_0|~ssList(X1)|nil!=esk49_0), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 0.02/0.27  fof(c_0_45, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
 % 0.02/0.27  cnf(c_0_46, negated_conjecture, (app(esk55_1(nil),cons(esk53_1(nil),nil))=esk48_0|esk49_0!=esk48_0|esk49_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_40])])).
 % 0.02/0.27  cnf(c_0_47, negated_conjecture, (esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_42])])).
 % 0.02/0.27  cnf(c_0_48, negated_conjecture, (ssList(esk55_1(nil))|esk49_0!=esk48_0|esk49_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_38]), c_0_39]), c_0_40])])).
 % 0.02/0.27  cnf(c_0_49, negated_conjecture, (ssItem(esk53_1(X1))|app(esk48_0,X1)!=esk49_0|esk49_0!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_11])])).
 % 0.02/0.27  cnf(c_0_50, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
 % 0.02/0.27  cnf(c_0_51, negated_conjecture, (app(esk55_1(nil),cons(esk53_1(nil),nil))=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_42]), c_0_42])]), c_0_47])])).
 % 0.02/0.27  cnf(c_0_52, negated_conjecture, (ssList(esk55_1(nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_42])]), c_0_47])])).
 % 0.02/0.27  fof(c_0_53, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 0.02/0.27  cnf(c_0_54, negated_conjecture, (ssItem(esk53_1(nil))|esk49_0!=esk48_0|esk49_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_38]), c_0_39]), c_0_40])])).
 % 0.02/0.27  cnf(c_0_55, negated_conjecture, (esk55_1(nil)=nil|~ssList(cons(esk53_1(nil),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])])).
 % 0.02/0.27  cnf(c_0_56, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
 % 0.02/0.27  cnf(c_0_57, negated_conjecture, (ssItem(esk53_1(nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_42])]), c_0_47])])).
 % 0.02/0.27  cnf(c_0_58, negated_conjecture, (esk55_1(nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_39]), c_0_57])])).
 % 0.02/0.27  cnf(c_0_59, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
 % 0.02/0.27  cnf(c_0_60, negated_conjecture, (app(nil,cons(esk53_1(nil),nil))=nil), inference(rw,[status(thm)],[c_0_51, c_0_58])).
 % 0.02/0.27  fof(c_0_61, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
 % 0.02/0.27  cnf(c_0_62, negated_conjecture, (cons(esk53_1(nil),nil)=nil|~ssList(cons(esk53_1(nil),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_39])])).
 % 0.02/0.27  cnf(c_0_63, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
 % 0.02/0.27  cnf(c_0_64, negated_conjecture, (cons(esk53_1(nil),nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_56]), c_0_39]), c_0_57])])).
 % 0.02/0.27  cnf(c_0_65, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_39]), c_0_57])]), ['proof']).
 % 0.02/0.27  # SZS output end CNFRefutation
 % 0.02/0.27  # Training examples: 0 positive, 0 negative
 % 0.02/0.27
 % 0.02/0.27  # -------------------------------------------------
 % 0.02/0.27  # User time                : 0.032 s
 % 0.02/0.27  # System time              : 0.004 s
 % 0.02/0.27  # Total time               : 0.036 s
 % 0.02/0.27  # Maximum resident set size: 1880 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
