   --------------------------------------------------------------------------

                   TSTP Solution File: SWC395+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC395+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n131.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:46 EDT 2016

 % Result     : Theorem 1.22s
 % Output     : CNFRefutation 1.22s
 % Verified   :
 % Statistics : Number of formulae       :  137 (9147 expanded)
 %              Number of clauses        :   94 (3720 expanded)
 %              Number of leaves         :   21 (2321 expanded)
 %              Depth                    :   30
 %              Number of atoms          :  529 (91266 expanded)
 %              Number of equality atoms :  181 (34748 expanded)
 %              Maximal formula depth    :   29 (   5 average)
 %              Maximal clause size      :   23 (   3 average)
 %              Maximal term depth       :    6 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ~ ssList(X4)
                   | X2 != X4
                   | X1 != X3
                   | ! [X5] :
                       ( ~ ssItem(X5)
                       | ~ memberP(X1,X5)
                       | memberP(X2,X5) )
                   | ( nil != X3
                     & nil = X4 )
                   | ( ? [X6] :
                         ( ssList(X6)
                         & X3 != X6
                         & ? [X7] :
                             ( ssList(X7)
                             & ? [X8] :
                                 ( ssList(X8)
                                 & tl(X4) = X7
                                 & app(X7,X8) = X6
                                 & ? [X9] :
                                     ( ssItem(X9)
                                     & cons(X9,nil) = X8
                                     & hd(X4) = X9
                                     & neq(nil,X4) )
                                 & neq(nil,X4) ) ) )
                     & neq(X4,nil) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssList(tl(X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_4,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssItem(hd(X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax22)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => cons(hd(X1),tl(X1)) = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax78)).

 fof(c_0_8,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil != X1
            => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax86)).

 fof(c_0_9,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil = app(X1,X2)
           <=> ( nil = X2
               & nil = X1 ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

 fof(c_0_10,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( nil != X1
            => hd(app(X1,X2)) = hd(X1) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax85)).

 fof(c_0_11,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => nil != cons(X2,X1) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax21)).

 fof(c_0_12,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

 fof(c_0_13,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(nil,X1) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

 fof(c_0_14,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ( memberP(X1,X2)
           <=> ? [X3] :
                 ( ssList(X3)
                 & ? [X4] :
                     ( ssList(X4)
                     & app(X3,cons(X2,X4)) = X1 ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax3)).

 fof(c_0_15,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssItem(X3)
              => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

 fof(c_0_16,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => hd(cons(X2,X1)) = X2 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax23)).

 fof(c_0_17,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => tl(cons(X2,X1)) = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

 fof(c_0_18,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( memberP(cons(X2,X3),X1)
               <=> ( X1 = X2
                   | memberP(X3,X1) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax37)).

 fof(c_0_19,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ( memberP(app(X2,X3),X1)
               <=> ( memberP(X2,X1)
                   | memberP(X3,X1) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax36)).

 fof(c_0_20,axiom,(
     ! [X1] :
       ( ssItem(X1)
      => ~ memberP(nil,X1) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax38)).

 fof(c_0_21,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ~ ssList(X4)
                     | X2 != X4
                     | X1 != X3
                     | ! [X5] :
                         ( ~ ssItem(X5)
                         | ~ memberP(X1,X5)
                         | memberP(X2,X5) )
                     | ( nil != X3
                       & nil = X4 )
                     | ( ? [X6] :
                           ( ssList(X6)
                           & X3 != X6
                           & ? [X7] :
                               ( ssList(X7)
                               & ? [X8] :
                                   ( ssList(X8)
                                   & tl(X4) = X7
                                   & app(X7,X8) = X6
                                   & ? [X9] :
                                       ( ssItem(X9)
                                       & cons(X9,nil) = X8
                                       & hd(X4) = X9
                                       & neq(nil,X4) )
                                   & neq(nil,X4) ) ) )
                       & neq(X4,nil) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_22,negated_conjecture,(
     ! [X15,X16,X17,X18] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ssItem(esk52_0)
       & memberP(esk48_0,esk52_0)
       & ~ memberP(esk49_0,esk52_0)
       & ( nil = esk50_0
         | nil != esk51_0 )
       & ( ~ ssList(X15)
         | esk50_0 = X15
         | ~ ssList(X16)
         | ~ ssList(X17)
         | tl(esk51_0) != X16
         | app(X16,X17) != X15
         | ~ ssItem(X18)
         | cons(X18,nil) != X17
         | hd(esk51_0) != X18
         | ~ neq(nil,esk51_0)
         | ~ neq(nil,esk51_0)
         | ~ neq(esk51_0,nil) ) ) ),
     inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).

 cnf(c_0_23,negated_conjecture,
     ( esk50_0 = X4
     | ~ neq(esk51_0,nil)
     | ~ neq(nil,esk51_0)
     | ~ neq(nil,esk51_0)
     | hd(esk51_0) != X1
     | cons(X1,nil) != X2
     | ~ ssItem(X1)
     | app(X3,X2) != X4
     | tl(esk51_0) != X3
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ ssList(X4) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_22])).

 fof(c_0_24,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

 cnf(c_0_25,negated_conjecture,
     ( esk50_0 = X4
     | hd(esk51_0) != X1
     | tl(esk51_0) != X3
     | app(X3,X2) != X4
     | cons(X1,nil) != X2
     | ~ ssItem(X1)
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ neq(nil,esk51_0)
     | ~ neq(esk51_0,nil) ),
     inference(cn,[status(thm)],[[5]c_0_23])).

 cnf(c_0_26,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[6]c_0_24])).

 fof(c_0_27,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssList(tl(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).

 cnf(c_0_28,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[7]c_0_22])).

 cnf(c_0_29,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[8]c_0_22])).

 cnf(c_0_30,negated_conjecture,
     ( app(tl(esk51_0),cons(hd(esk51_0),nil)) = esk50_0
     | ~ ssList(cons(hd(esk51_0),nil))
     | ~ ssList(tl(esk51_0))
     | ~ neq(nil,esk51_0)
     | ~ neq(esk51_0,nil)
     | ~ ssItem(hd(esk51_0)) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_25])])])]),[9]c_0_26])).

 cnf(c_0_31,plain,
     ( ssList(tl(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[10]c_0_27])).

 cnf(c_0_32,negated_conjecture,
     ( ssList(esk51_0) ),
     inference(rw,[status(thm)],[[11]c_0_28,[12]c_0_29])).

 fof(c_0_33,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

 cnf(c_0_34,negated_conjecture,
     ( app(tl(esk51_0),cons(hd(esk51_0),nil)) = esk50_0
     | esk51_0 = nil
     | ~ ssList(cons(hd(esk51_0),nil))
     | ~ neq(nil,esk51_0)
     | ~ neq(esk51_0,nil)
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[13]c_0_30,c_0_31]),c_0_32])])).

 cnf(c_0_35,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_33])).

 cnf(c_0_36,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[15]c_0_4])).

 fof(c_0_37,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssItem(hd(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).

 cnf(c_0_38,negated_conjecture,
     ( app(tl(esk51_0),cons(hd(esk51_0),nil)) = esk50_0
     | esk51_0 = nil
     | ~ neq(nil,esk51_0)
     | ~ neq(esk51_0,nil)
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[16]c_0_34,c_0_35]),c_0_36])])).

 cnf(c_0_39,plain,
     ( ssItem(hd(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[17]c_0_37])).

 fof(c_0_40,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).

 fof(c_0_41,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | cons(hd(X2),tl(X2)) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

 fof(c_0_42,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | nil = X3
       | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

 fof(c_0_43,plain,(
     ! [X3,X4] :
       ( ( nil = X4
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil = X3
         | nil != app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( nil != X4
         | nil != X3
         | nil = app(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).

 cnf(c_0_44,negated_conjecture,
     ( app(tl(esk51_0),cons(hd(esk51_0),nil)) = esk50_0
     | esk51_0 = nil
     | ~ neq(nil,esk51_0)
     | ~ neq(esk51_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_38,c_0_39]),c_0_32])])).

 cnf(c_0_45,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_40])).

 cnf(c_0_46,plain,
     ( cons(hd(X1),tl(X1)) = X1
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[20]c_0_41])).

 cnf(c_0_47,plain,
     ( tl(app(X1,X2)) = app(tl(X1),X2)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[21]c_0_42])).

 cnf(c_0_48,plain,
     ( nil = X1
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != app(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[22]c_0_43])).

 cnf(c_0_49,negated_conjecture,
     ( app(tl(esk51_0),cons(hd(esk51_0),nil)) = esk50_0
     | esk51_0 = nil
     | ~ neq(esk51_0,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[23]c_0_44,c_0_45]),c_0_32]),c_0_36])])).

 cnf(c_0_50,plain,
     ( cons(hd(app(X1,X2)),app(tl(X1),X2)) = app(X1,X2)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[24]c_0_46,c_0_47]),c_0_26]),[25]c_0_48])).

 cnf(c_0_51,negated_conjecture,
     ( app(tl(esk51_0),cons(hd(esk51_0),nil)) = esk50_0
     | esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_49,c_0_45]),c_0_36]),c_0_32])])).

 cnf(c_0_52,negated_conjecture,
     ( cons(hd(app(esk51_0,cons(hd(esk51_0),nil))),esk50_0) = app(esk51_0,cons(hd(esk51_0),nil))
     | esk51_0 = nil
     | ~ ssList(cons(hd(esk51_0),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_50,c_0_51]),c_0_32])])).

 cnf(c_0_53,negated_conjecture,
     ( cons(hd(app(esk51_0,cons(hd(esk51_0),nil))),esk50_0) = app(esk51_0,cons(hd(esk51_0),nil))
     | esk51_0 = nil
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[28]c_0_52,c_0_35]),c_0_36])])).

 fof(c_0_54,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | nil = X3
       | hd(app(X3,X4)) = hd(X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).

 cnf(c_0_55,negated_conjecture,
     ( cons(hd(app(esk51_0,cons(hd(esk51_0),nil))),esk50_0) = app(esk51_0,cons(hd(esk51_0),nil))
     | esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[29]c_0_53,c_0_39]),c_0_32])])).

 cnf(c_0_56,plain,
     ( hd(app(X1,X2)) = hd(X1)
     | nil = X1
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[30]c_0_54])).

 fof(c_0_57,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | nil != cons(X4,X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).

 cnf(c_0_58,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[31]c_0_22])).

 cnf(c_0_59,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[32]c_0_22])).

 cnf(c_0_60,negated_conjecture,
     ( app(esk51_0,cons(hd(esk51_0),nil)) = cons(hd(esk51_0),esk50_0)
     | esk51_0 = nil
     | ~ ssList(cons(hd(esk51_0),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[33]c_0_55,c_0_56]),c_0_32])])).

 cnf(c_0_61,plain,
     ( nil != cons(X1,X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[34]c_0_57])).

 cnf(c_0_62,negated_conjecture,
     ( ssList(esk50_0) ),
     inference(rw,[status(thm)],[[35]c_0_58,[36]c_0_59])).

 fof(c_0_63,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssList(X6)
       | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).

 fof(c_0_64,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(nil,X2) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).

 fof(c_0_65,plain,(
     ! [X5,X6,X9,X10] :
       ( ( ssList(esk3_2(X5,X6))
         | ~ memberP(X5,X6)
         | ~ ssItem(X6)
         | ~ ssList(X5) )
       & ( ssList(esk4_2(X5,X6))
         | ~ memberP(X5,X6)
         | ~ ssItem(X6)
         | ~ ssList(X5) )
       & ( app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6))) = X5
         | ~ memberP(X5,X6)
         | ~ ssItem(X6)
         | ~ ssList(X5) )
       & ( ~ ssList(X9)
         | ~ ssList(X10)
         | app(X9,cons(X6,X10)) != X5
         | memberP(X5,X6)
         | ~ ssItem(X6)
         | ~ ssList(X5) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])])).

 cnf(c_0_66,negated_conjecture,
     ( memberP(esk48_0,esk52_0) ),
     inference(split_conjunct,[status(thm)],[[37]c_0_22])).

 cnf(c_0_67,negated_conjecture,
     ( app(esk51_0,cons(hd(esk51_0),nil)) = cons(hd(esk51_0),esk50_0)
     | esk51_0 = nil
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[38]c_0_60,c_0_35]),c_0_36])])).

 cnf(c_0_68,negated_conjecture,
     ( esk51_0 = nil
     | app(esk51_0,cons(hd(esk51_0),nil)) != nil
     | ~ ssItem(hd(app(esk51_0,cons(hd(esk51_0),nil)))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[39]c_0_61,c_0_55]),c_0_62])])).

 cnf(c_0_69,plain,
     ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[40]c_0_63])).

 cnf(c_0_70,plain,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[41]c_0_64])).

 cnf(c_0_71,plain,
     ( app(esk3_2(X1,X2),cons(X2,esk4_2(X1,X2))) = X1
     | ~ ssList(X1)
     | ~ ssItem(X2)
     | ~ memberP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[42]c_0_65])).

 cnf(c_0_72,negated_conjecture,
     ( memberP(esk50_0,esk52_0) ),
     inference(rw,[status(thm)],[[43]c_0_66,[44]c_0_59])).

 cnf(c_0_73,negated_conjecture,
     ( ssItem(esk52_0) ),
     inference(split_conjunct,[status(thm)],[[45]c_0_22])).

 cnf(c_0_74,plain,
     ( ssList(esk3_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssItem(X2)
     | ~ memberP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[46]c_0_65])).

 cnf(c_0_75,negated_conjecture,
     ( app(esk51_0,cons(hd(esk51_0),nil)) = cons(hd(esk51_0),esk50_0)
     | esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[47]c_0_67,c_0_39]),c_0_32])])).

 cnf(c_0_76,negated_conjecture,
     ( esk51_0 = nil
     | app(esk51_0,cons(hd(esk51_0),nil)) != nil
     | ~ ssList(cons(hd(esk51_0),nil))
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[48]c_0_68,c_0_56]),c_0_32])])).

 fof(c_0_77,plain,(
     ! [X4,X5,X6] :
       ( ~ ssList(X4)
       | ~ ssList(X5)
       | ~ ssItem(X6)
       | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).

 cnf(c_0_78,plain,
     ( app(nil,app(X1,X2)) = app(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[49]c_0_69,c_0_70]),c_0_36])])).

 cnf(c_0_79,negated_conjecture,
     ( app(esk3_2(esk50_0,esk52_0),cons(esk52_0,esk4_2(esk50_0,esk52_0))) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[50]c_0_71,c_0_72]),c_0_62]),c_0_73])])).

 cnf(c_0_80,negated_conjecture,
     ( ssList(esk3_2(esk50_0,esk52_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[51]c_0_74,c_0_72]),c_0_62]),c_0_73])])).

 cnf(c_0_81,plain,
     ( ssList(esk4_2(X1,X2))
     | ~ ssList(X1)
     | ~ ssItem(X2)
     | ~ memberP(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[52]c_0_65])).

 cnf(c_0_82,negated_conjecture,
     ( hd(cons(hd(esk51_0),esk50_0)) = hd(esk51_0)
     | esk51_0 = nil
     | ~ ssList(cons(hd(esk51_0),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[53]c_0_56,c_0_75]),c_0_32])])).

 cnf(c_0_83,negated_conjecture,
     ( esk51_0 = nil
     | ssList(cons(hd(esk51_0),esk50_0))
     | ~ ssList(cons(hd(esk51_0),nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[54]c_0_26,c_0_75]),c_0_32])])).

 cnf(c_0_84,negated_conjecture,
     ( esk51_0 = nil
     | app(esk51_0,cons(hd(esk51_0),nil)) != nil
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[55]c_0_76,c_0_35]),c_0_36])])).

 fof(c_0_85,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | hd(cons(X4,X3)) = X4 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).

 fof(c_0_86,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | tl(cons(X4,X3)) = X3 ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).

 cnf(c_0_87,plain,
     ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     inference(split_conjunct,[status(thm)],[[56]c_0_77])).

 cnf(c_0_88,negated_conjecture,
     ( app(nil,esk50_0) = esk50_0
     | ~ ssList(cons(esk52_0,esk4_2(esk50_0,esk52_0))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[57]c_0_78,c_0_79]),c_0_80])])).

 cnf(c_0_89,negated_conjecture,
     ( ssList(esk4_2(esk50_0,esk52_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[58]c_0_81,c_0_72]),c_0_62]),c_0_73])])).

 cnf(c_0_90,negated_conjecture,
     ( hd(cons(hd(esk51_0),esk50_0)) = hd(esk51_0)
     | esk51_0 = nil
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[59]c_0_82,c_0_35]),c_0_36])])).

 cnf(c_0_91,negated_conjecture,
     ( esk51_0 = nil
     | ssList(cons(hd(esk51_0),esk50_0))
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[60]c_0_83,c_0_35]),c_0_36])])).

 cnf(c_0_92,negated_conjecture,
     ( esk51_0 = nil
     | app(esk51_0,cons(hd(esk51_0),nil)) != nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[61]c_0_84,c_0_39]),c_0_32])])).

 cnf(c_0_93,plain,
     ( hd(cons(X1,X2)) = X1
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[62]c_0_85])).

 cnf(c_0_94,plain,
     ( tl(cons(X1,X2)) = X2
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[63]c_0_86])).

 cnf(c_0_95,plain,
     ( hd(cons(X1,app(X2,X3))) = hd(cons(X1,X2))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[64]c_0_56,c_0_87]),c_0_35]),[65]c_0_61])).

 cnf(c_0_96,negated_conjecture,
     ( app(nil,esk50_0) = esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[66]c_0_88,c_0_35]),c_0_89]),c_0_73])])).

 cnf(c_0_97,negated_conjecture,
     ( hd(cons(hd(esk51_0),esk50_0)) = hd(esk51_0)
     | esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[67]c_0_90,c_0_39]),c_0_32])])).

 cnf(c_0_98,negated_conjecture,
     ( esk51_0 = nil
     | ssList(cons(hd(esk51_0),esk50_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[68]c_0_91,c_0_39]),c_0_32])])).

 cnf(c_0_99,negated_conjecture,
     ( esk51_0 = nil
     | cons(hd(esk51_0),esk50_0) != nil ),
     inference(spm,[status(thm)],[[69]c_0_92,[70]c_0_75])).

 fof(c_0_100,plain,(
     ! [X4,X5,X6] :
       ( ( ~ memberP(cons(X5,X6),X4)
         | X4 = X5
         | memberP(X6,X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) )
       & ( X4 != X5
         | memberP(cons(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) )
       & ( ~ memberP(X6,X4)
         | memberP(cons(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssItem(X5)
         | ~ ssItem(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])).

 cnf(c_0_101,plain,
     ( cons(X1,tl(cons(X1,X2))) = cons(X1,X2)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[71]c_0_46,c_0_93]),c_0_35]),[72]c_0_61])).

 cnf(c_0_102,plain,
     ( cons(hd(cons(X1,X2)),X2) = cons(X1,X2)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[73]c_0_46,c_0_94]),c_0_35]),[74]c_0_61])).

 cnf(c_0_103,negated_conjecture,
     ( hd(cons(X1,esk50_0)) = hd(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[75]c_0_95,c_0_96]),c_0_62]),c_0_36])])).

 cnf(c_0_104,negated_conjecture,
     ( esk51_0 = nil
     | ssItem(hd(esk51_0)) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[76]c_0_39,c_0_97]),c_0_98]),[77]c_0_99])).

 cnf(c_0_105,plain,
     ( memberP(cons(X2,X3),X1)
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | ~ memberP(X3,X1) ),
     inference(split_conjunct,[status(thm)],[[78]c_0_100])).

 cnf(c_0_106,plain,
     ( nil = app(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | nil != X1
     | nil != X2 ),
     inference(split_conjunct,[status(thm)],[[79]c_0_43])).

 cnf(c_0_107,plain,
     ( cons(hd(cons(X1,X2)),tl(cons(X1,X2))) = cons(X1,X2)
     | ~ ssList(X2)
     | ~ ssItem(hd(cons(X1,X2)))
     | ~ ssItem(X1) ),
     inference(spm,[status(thm)],[[80]c_0_101,[81]c_0_102])).

 cnf(c_0_108,negated_conjecture,
     ( hd(cons(hd(esk51_0),nil)) = hd(esk51_0)
     | esk51_0 = nil ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[82]c_0_97,c_0_103]),[83]c_0_104])).

 cnf(c_0_109,negated_conjecture,
     ( memberP(cons(X1,esk50_0),esk52_0)
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[84]c_0_105,c_0_72]),c_0_62]),c_0_73])])).

 cnf(c_0_110,plain,
     ( ssList(cons(X1,app(X2,X3)))
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[85]c_0_26,c_0_87]),[86]c_0_35])).

 cnf(c_0_111,plain,
     ( app(nil,nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_106])]),c_0_36])])).

 cnf(c_0_112,negated_conjecture,
     ( cons(hd(esk51_0),tl(cons(hd(esk51_0),nil))) = cons(hd(esk51_0),nil)
     | esk51_0 = nil ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[87]c_0_107,c_0_108]),c_0_36])]),[88]c_0_104])).

 fof(c_0_113,plain,(
     ! [X4,X5,X6] :
       ( ( ~ memberP(app(X5,X6),X4)
         | memberP(X5,X4)
         | memberP(X6,X4)
         | ~ ssList(X6)
         | ~ ssList(X5)
         | ~ ssItem(X4) )
       & ( ~ memberP(X5,X4)
         | memberP(app(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssList(X5)
         | ~ ssItem(X4) )
       & ( ~ memberP(X6,X4)
         | memberP(app(X5,X6),X4)
         | ~ ssList(X6)
         | ~ ssList(X5)
         | ~ ssItem(X4) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).

 cnf(c_0_114,negated_conjecture,
     ( esk51_0 = nil
     | memberP(app(esk51_0,cons(hd(esk51_0),nil)),esk52_0)
     | ~ ssItem(hd(app(esk51_0,cons(hd(esk51_0),nil)))) ),
     inference(spm,[status(thm)],[[89]c_0_109,[90]c_0_55])).

 cnf(c_0_115,plain,
     ( ssList(cons(X1,nil))
     | ~ ssItem(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[91]c_0_110,c_0_111]),c_0_36])])).

 cnf(c_0_116,negated_conjecture,
     ( ~ memberP(esk49_0,esk52_0) ),
     inference(split_conjunct,[status(thm)],[[92]c_0_22])).

 cnf(c_0_117,negated_conjecture,
     ( esk51_0 = nil
     | ssList(cons(hd(esk51_0),nil))
     | ~ ssList(tl(cons(hd(esk51_0),nil))) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[93]c_0_35,c_0_112]),[94]c_0_104])).

 cnf(c_0_118,plain,
     ( memberP(cons(X2,X3),X1)
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | X1 != X2 ),
     inference(split_conjunct,[status(thm)],[[95]c_0_100])).

 cnf(c_0_119,plain,
     ( memberP(X3,X1)
     | memberP(X2,X1)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ memberP(app(X2,X3),X1) ),
     inference(split_conjunct,[status(thm)],[[96]c_0_113])).

 cnf(c_0_120,negated_conjecture,
     ( esk51_0 = nil
     | memberP(app(esk51_0,cons(hd(esk51_0),nil)),esk52_0) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[97]c_0_114,c_0_56]),c_0_32])]),c_0_115]),[98]c_0_104])).

 cnf(c_0_121,negated_conjecture,
     ( ~ memberP(esk51_0,esk52_0) ),
     inference(rw,[status(thm)],[[99]c_0_116,[100]c_0_29])).

 cnf(c_0_122,negated_conjecture,
     ( esk51_0 = nil
     | ssList(cons(hd(esk51_0),nil)) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[101]c_0_117,c_0_94]),c_0_36])]),[102]c_0_104])).

 fof(c_0_123,plain,(
     ! [X1] :
       ( ssItem(X1)
      => ~ memberP(nil,X1) ) ),
     inference(fof_simplification,[status(thm)],[[103]c_0_20])).

 cnf(c_0_124,plain,
     ( memberP(cons(X1,X2),X1)
     | ~ ssList(X2)
     | ~ ssItem(X1) ),
     inference(er,[status(thm)],[[104]c_0_118])).

 cnf(c_0_125,plain,
     ( memberP(X3,X1)
     | X1 = X2
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | ~ memberP(cons(X2,X3),X1) ),
     inference(split_conjunct,[status(thm)],[[105]c_0_100])).

 cnf(c_0_126,negated_conjecture,
     ( esk51_0 = nil
     | memberP(cons(hd(esk51_0),nil),esk52_0) ),
     inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[106]c_0_119,c_0_120]),c_0_32]),c_0_73])]),c_0_121]),[107]c_0_122])).

 fof(c_0_127,plain,(
     ! [X2] :
       ( ~ ssItem(X2)
       | ~ memberP(nil,X2) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_123])])).

 cnf(c_0_128,plain,
     ( nil = X1
     | memberP(X1,hd(X1))
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[108]c_0_124,c_0_46]),c_0_39]),[109]c_0_31])).

 cnf(c_0_129,negated_conjecture,
     ( hd(esk51_0) = esk52_0
     | esk51_0 = nil
     | memberP(nil,esk52_0) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[110]c_0_125,c_0_126]),c_0_36]),c_0_73])]),[111]c_0_104])).

 cnf(c_0_130,plain,
     ( ~ memberP(nil,X1)
     | ~ ssItem(X1) ),
     inference(split_conjunct,[status(thm)],[[112]c_0_127])).

 cnf(c_0_131,negated_conjecture,
     ( esk51_0 = nil
     | memberP(nil,esk52_0) ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[113]c_0_128,c_0_129]),c_0_32])]),[114]c_0_121])).

 cnf(c_0_132,negated_conjecture,
     ( nil = esk50_0
     | nil != esk51_0 ),
     inference(split_conjunct,[status(thm)],[[115]c_0_22])).

 cnf(c_0_133,negated_conjecture,
     ( esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[116]c_0_130,c_0_131]),c_0_73])])).

 cnf(c_0_134,negated_conjecture,
     ( esk50_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[117]c_0_132,c_0_133])])).

 cnf(c_0_135,negated_conjecture,
     ( ~ memberP(nil,esk52_0) ),
     inference(rw,[status(thm)],[[118]c_0_121,[119]c_0_133])).

 cnf(c_0_136,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(rw,[status(thm)],[[120]c_0_72,c_0_134]),[121]c_0_135]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC395+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.22  % Computer   : n131.star.cs.uiowa.edu
 % 0.03/0.22  % Model      : x86_64 x86_64
 % 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.22  % Memory     : 32218.75MB
 % 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.22  % CPULimit   : 300
 % 0.03/0.22  % DateTime   : Tue May 10 20:43:08 CDT 2016
 % 0.03/0.22  % CPUTime    :
 % 0.03/0.22  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 1.22/1.42  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 1.22/1.42  # and selection function SelectCQPrecWNTNp.
 % 1.22/1.42  #
 % 1.22/1.42  # Preprocessing time       : 0.014 s
 % 1.22/1.42  # Presaturation interreduction done
 % 1.22/1.42
 % 1.22/1.42  # Proof found!
 % 1.22/1.42  # SZS status Theorem
 % 1.22/1.42  # SZS output start CNFRefutation
 % 1.22/1.42  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~(ssList(X4))|X2!=X4)|X1!=X3)|![X5]:((~(ssItem(X5))|~(memberP(X1,X5)))|memberP(X2,X5)))|(nil!=X3&nil=X4))|(?[X6]:((ssList(X6)&X3!=X6)&?[X7]:(ssList(X7)&?[X8]:((((ssList(X8)&tl(X4)=X7)&app(X7,X8)=X6)&?[X9]:(((ssItem(X9)&cons(X9,nil)=X8)&hd(X4)=X9)&neq(nil,X4)))&neq(nil,X4))))&neq(X4,nil))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 1.22/1.42  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 1.22/1.42  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
 % 1.22/1.42  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 1.22/1.42  fof(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 1.22/1.42  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax22)).
 % 1.22/1.42  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 1.22/1.42  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax78)).
 % 1.22/1.42  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax86)).
 % 1.22/1.42  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)<=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
 % 1.22/1.42  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>hd(app(X1,X2))=hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax85)).
 % 1.22/1.42  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>nil!=cons(X2,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax21)).
 % 1.22/1.42  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
 % 1.22/1.42  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 1.22/1.42  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>(memberP(X1,X2)<=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(X3,cons(X2,X4))=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax3)).
 % 1.22/1.42  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
 % 1.22/1.42  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax23)).
 % 1.22/1.42  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
 % 1.22/1.42  fof(c_0_18, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)<=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax37)).
 % 1.22/1.42  fof(c_0_19, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(memberP(app(X2,X3),X1)<=>(memberP(X2,X1)|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax36)).
 % 1.22/1.42  fof(c_0_20, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax38)).
 % 1.22/1.42  fof(c_0_21, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((((~ssList(X4)|X2!=X4)|X1!=X3)|![X5]:((~ssItem(X5)|~memberP(X1,X5))|memberP(X2,X5)))|(nil!=X3&nil=X4))|(?[X6]:((ssList(X6)&X3!=X6)&?[X7]:(ssList(X7)&?[X8]:((((ssList(X8)&tl(X4)=X7)&app(X7,X8)=X6)&?[X9]:(((ssItem(X9)&cons(X9,nil)=X8)&hd(X4)=X9)&neq(nil,X4)))&neq(nil,X4))))&neq(X4,nil)))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 1.22/1.42  fof(c_0_22, negated_conjecture, (![X15]:![X16]:![X17]:![X18]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&((ssItem(esk52_0)&memberP(esk48_0,esk52_0))&~memberP(esk49_0,esk52_0)))&(nil=esk50_0|nil!=esk51_0))&(((~ssList(X15)|esk50_0=X15)|(~ssList(X16)|((((~ssList(X17)|tl(esk51_0)!=X16)|app(X16,X17)!=X15)|(((~ssItem(X18)|cons(X18,nil)!=X17)|hd(esk51_0)!=X18)|~neq(nil,esk51_0)))|~neq(nil,esk51_0))))|~neq(esk51_0,nil))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])])).
 % 1.22/1.42  cnf(c_0_23, negated_conjecture, (esk50_0=X4|~neq(esk51_0,nil)|~neq(nil,esk51_0)|~neq(nil,esk51_0)|hd(esk51_0)!=X1|cons(X1,nil)!=X2|~ssItem(X1)|app(X3,X2)!=X4|tl(esk51_0)!=X3|~ssList(X2)|~ssList(X3)|~ssList(X4)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 1.22/1.42  fof(c_0_24, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
 % 1.22/1.42  cnf(c_0_25, negated_conjecture, (esk50_0=X4|hd(esk51_0)!=X1|tl(esk51_0)!=X3|app(X3,X2)!=X4|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X4)|~ssList(X3)|~ssList(X2)|~neq(nil,esk51_0)|~neq(esk51_0,nil)), inference(cn,[status(thm)],[c_0_23])).
 % 1.22/1.42  cnf(c_0_26, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
 % 1.22/1.42  fof(c_0_27, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])).
 % 1.22/1.42  cnf(c_0_28, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 1.22/1.42  cnf(c_0_29, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 1.22/1.42  cnf(c_0_30, negated_conjecture, (app(tl(esk51_0),cons(hd(esk51_0),nil))=esk50_0|~ssList(cons(hd(esk51_0),nil))|~ssList(tl(esk51_0))|~neq(nil,esk51_0)|~neq(esk51_0,nil)|~ssItem(hd(esk51_0))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_25])])])]), c_0_26])).
 % 1.22/1.42  cnf(c_0_31, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
 % 1.22/1.42  cnf(c_0_32, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
 % 1.22/1.42  fof(c_0_33, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
 % 1.22/1.42  cnf(c_0_34, negated_conjecture, (app(tl(esk51_0),cons(hd(esk51_0),nil))=esk50_0|esk51_0=nil|~ssList(cons(hd(esk51_0),nil))|~neq(nil,esk51_0)|~neq(esk51_0,nil)|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_35, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
 % 1.22/1.42  cnf(c_0_36, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
 % 1.22/1.42  fof(c_0_37, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
 % 1.22/1.42  cnf(c_0_38, negated_conjecture, (app(tl(esk51_0),cons(hd(esk51_0),nil))=esk50_0|esk51_0=nil|~neq(nil,esk51_0)|~neq(esk51_0,nil)|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
 % 1.22/1.42  cnf(c_0_39, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
 % 1.22/1.42  fof(c_0_40, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
 % 1.22/1.42  fof(c_0_41, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
 % 1.22/1.42  fof(c_0_42, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
 % 1.22/1.42  fof(c_0_43, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
 % 1.22/1.42  cnf(c_0_44, negated_conjecture, (app(tl(esk51_0),cons(hd(esk51_0),nil))=esk50_0|esk51_0=nil|~neq(nil,esk51_0)|~neq(esk51_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_45, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
 % 1.22/1.42  cnf(c_0_46, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 1.22/1.42  cnf(c_0_47, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
 % 1.22/1.42  cnf(c_0_48, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
 % 1.22/1.42  cnf(c_0_49, negated_conjecture, (app(tl(esk51_0),cons(hd(esk51_0),nil))=esk50_0|esk51_0=nil|~neq(esk51_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_32]), c_0_36])])).
 % 1.22/1.42  cnf(c_0_50, plain, (cons(hd(app(X1,X2)),app(tl(X1),X2))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_26]), c_0_48])).
 % 1.22/1.42  cnf(c_0_51, negated_conjecture, (app(tl(esk51_0),cons(hd(esk51_0),nil))=esk50_0|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_45]), c_0_36]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_52, negated_conjecture, (cons(hd(app(esk51_0,cons(hd(esk51_0),nil))),esk50_0)=app(esk51_0,cons(hd(esk51_0),nil))|esk51_0=nil|~ssList(cons(hd(esk51_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_53, negated_conjecture, (cons(hd(app(esk51_0,cons(hd(esk51_0),nil))),esk50_0)=app(esk51_0,cons(hd(esk51_0),nil))|esk51_0=nil|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_35]), c_0_36])])).
 % 1.22/1.42  fof(c_0_54, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|hd(app(X3,X4))=hd(X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
 % 1.22/1.42  cnf(c_0_55, negated_conjecture, (cons(hd(app(esk51_0,cons(hd(esk51_0),nil))),esk50_0)=app(esk51_0,cons(hd(esk51_0),nil))|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_39]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_56, plain, (hd(app(X1,X2))=hd(X1)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
 % 1.22/1.42  fof(c_0_57, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|nil!=cons(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
 % 1.22/1.42  cnf(c_0_58, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 1.22/1.42  cnf(c_0_59, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 1.22/1.42  cnf(c_0_60, negated_conjecture, (app(esk51_0,cons(hd(esk51_0),nil))=cons(hd(esk51_0),esk50_0)|esk51_0=nil|~ssList(cons(hd(esk51_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_61, plain, (nil!=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
 % 1.22/1.42  cnf(c_0_62, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_58, c_0_59])).
 % 1.22/1.42  fof(c_0_63, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
 % 1.22/1.42  fof(c_0_64, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])).
 % 1.22/1.42  fof(c_0_65, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk3_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&((((ssList(esk4_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&(((app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6)))=X5|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(X9,cons(X6,X10))!=X5))|memberP(X5,X6))|~ssItem(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])])])).
 % 1.22/1.42  cnf(c_0_66, negated_conjecture, (memberP(esk48_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 1.22/1.42  cnf(c_0_67, negated_conjecture, (app(esk51_0,cons(hd(esk51_0),nil))=cons(hd(esk51_0),esk50_0)|esk51_0=nil|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_35]), c_0_36])])).
 % 1.22/1.42  cnf(c_0_68, negated_conjecture, (esk51_0=nil|app(esk51_0,cons(hd(esk51_0),nil))!=nil|~ssItem(hd(app(esk51_0,cons(hd(esk51_0),nil))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_55]), c_0_62])])).
 % 1.22/1.42  cnf(c_0_69, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
 % 1.22/1.42  cnf(c_0_70, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
 % 1.22/1.42  cnf(c_0_71, plain, (app(esk3_2(X1,X2),cons(X2,esk4_2(X1,X2)))=X1|~ssList(X1)|~ssItem(X2)|~memberP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
 % 1.22/1.42  cnf(c_0_72, negated_conjecture, (memberP(esk50_0,esk52_0)), inference(rw,[status(thm)],[c_0_66, c_0_59])).
 % 1.22/1.42  cnf(c_0_73, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 1.22/1.42  cnf(c_0_74, plain, (ssList(esk3_2(X1,X2))|~ssList(X1)|~ssItem(X2)|~memberP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
 % 1.22/1.42  cnf(c_0_75, negated_conjecture, (app(esk51_0,cons(hd(esk51_0),nil))=cons(hd(esk51_0),esk50_0)|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_39]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_76, negated_conjecture, (esk51_0=nil|app(esk51_0,cons(hd(esk51_0),nil))!=nil|~ssList(cons(hd(esk51_0),nil))|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_56]), c_0_32])])).
 % 1.22/1.42  fof(c_0_77, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
 % 1.22/1.42  cnf(c_0_78, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_36])])).
 % 1.22/1.42  cnf(c_0_79, negated_conjecture, (app(esk3_2(esk50_0,esk52_0),cons(esk52_0,esk4_2(esk50_0,esk52_0)))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_62]), c_0_73])])).
 % 1.22/1.42  cnf(c_0_80, negated_conjecture, (ssList(esk3_2(esk50_0,esk52_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_72]), c_0_62]), c_0_73])])).
 % 1.22/1.42  cnf(c_0_81, plain, (ssList(esk4_2(X1,X2))|~ssList(X1)|~ssItem(X2)|~memberP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
 % 1.22/1.42  cnf(c_0_82, negated_conjecture, (hd(cons(hd(esk51_0),esk50_0))=hd(esk51_0)|esk51_0=nil|~ssList(cons(hd(esk51_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_75]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_83, negated_conjecture, (esk51_0=nil|ssList(cons(hd(esk51_0),esk50_0))|~ssList(cons(hd(esk51_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_75]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_84, negated_conjecture, (esk51_0=nil|app(esk51_0,cons(hd(esk51_0),nil))!=nil|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_35]), c_0_36])])).
 % 1.22/1.42  fof(c_0_85, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
 % 1.22/1.42  fof(c_0_86, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
 % 1.22/1.42  cnf(c_0_87, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
 % 1.22/1.42  cnf(c_0_88, negated_conjecture, (app(nil,esk50_0)=esk50_0|~ssList(cons(esk52_0,esk4_2(esk50_0,esk52_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80])])).
 % 1.22/1.42  cnf(c_0_89, negated_conjecture, (ssList(esk4_2(esk50_0,esk52_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_72]), c_0_62]), c_0_73])])).
 % 1.22/1.42  cnf(c_0_90, negated_conjecture, (hd(cons(hd(esk51_0),esk50_0))=hd(esk51_0)|esk51_0=nil|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_35]), c_0_36])])).
 % 1.22/1.42  cnf(c_0_91, negated_conjecture, (esk51_0=nil|ssList(cons(hd(esk51_0),esk50_0))|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_35]), c_0_36])])).
 % 1.22/1.42  cnf(c_0_92, negated_conjecture, (esk51_0=nil|app(esk51_0,cons(hd(esk51_0),nil))!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_39]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_93, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
 % 1.22/1.42  cnf(c_0_94, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_86])).
 % 1.22/1.42  cnf(c_0_95, plain, (hd(cons(X1,app(X2,X3)))=hd(cons(X1,X2))|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_87]), c_0_35]), c_0_61])).
 % 1.22/1.42  cnf(c_0_96, negated_conjecture, (app(nil,esk50_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_35]), c_0_89]), c_0_73])])).
 % 1.22/1.42  cnf(c_0_97, negated_conjecture, (hd(cons(hd(esk51_0),esk50_0))=hd(esk51_0)|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_39]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_98, negated_conjecture, (esk51_0=nil|ssList(cons(hd(esk51_0),esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_39]), c_0_32])])).
 % 1.22/1.42  cnf(c_0_99, negated_conjecture, (esk51_0=nil|cons(hd(esk51_0),esk50_0)!=nil), inference(spm,[status(thm)],[c_0_92, c_0_75])).
 % 1.22/1.42  fof(c_0_100, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])).
 % 1.22/1.42  cnf(c_0_101, plain, (cons(X1,tl(cons(X1,X2)))=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_93]), c_0_35]), c_0_61])).
 % 1.22/1.42  cnf(c_0_102, plain, (cons(hd(cons(X1,X2)),X2)=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_94]), c_0_35]), c_0_61])).
 % 1.22/1.42  cnf(c_0_103, negated_conjecture, (hd(cons(X1,esk50_0))=hd(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_62]), c_0_36])])).
 % 1.22/1.42  cnf(c_0_104, negated_conjecture, (esk51_0=nil|ssItem(hd(esk51_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_97]), c_0_98]), c_0_99])).
 % 1.22/1.42  cnf(c_0_105, plain, (memberP(cons(X2,X3),X1)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~memberP(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
 % 1.22/1.42  cnf(c_0_106, plain, (nil=app(X1,X2)|~ssList(X1)|~ssList(X2)|nil!=X1|nil!=X2), inference(split_conjunct,[status(thm)],[c_0_43])).
 % 1.22/1.42  cnf(c_0_107, plain, (cons(hd(cons(X1,X2)),tl(cons(X1,X2)))=cons(X1,X2)|~ssList(X2)|~ssItem(hd(cons(X1,X2)))|~ssItem(X1)), inference(spm,[status(thm)],[c_0_101, c_0_102])).
 % 1.22/1.42  cnf(c_0_108, negated_conjecture, (hd(cons(hd(esk51_0),nil))=hd(esk51_0)|esk51_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_103]), c_0_104])).
 % 1.22/1.42  cnf(c_0_109, negated_conjecture, (memberP(cons(X1,esk50_0),esk52_0)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_72]), c_0_62]), c_0_73])])).
 % 1.22/1.42  cnf(c_0_110, plain, (ssList(cons(X1,app(X2,X3)))|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_87]), c_0_35])).
 % 1.22/1.42  cnf(c_0_111, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_106])]), c_0_36])])).
 % 1.22/1.42  cnf(c_0_112, negated_conjecture, (cons(hd(esk51_0),tl(cons(hd(esk51_0),nil)))=cons(hd(esk51_0),nil)|esk51_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_36])]), c_0_104])).
 % 1.22/1.42  fof(c_0_113, plain, (![X4]:![X5]:![X6]:(((((~memberP(app(X5,X6),X4)|(memberP(X5,X4)|memberP(X6,X4)))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&(((((~memberP(X5,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
 % 1.22/1.42  cnf(c_0_114, negated_conjecture, (esk51_0=nil|memberP(app(esk51_0,cons(hd(esk51_0),nil)),esk52_0)|~ssItem(hd(app(esk51_0,cons(hd(esk51_0),nil))))), inference(spm,[status(thm)],[c_0_109, c_0_55])).
 % 1.22/1.42  cnf(c_0_115, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_36])])).
 % 1.22/1.42  cnf(c_0_116, negated_conjecture, (~memberP(esk49_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 1.22/1.42  cnf(c_0_117, negated_conjecture, (esk51_0=nil|ssList(cons(hd(esk51_0),nil))|~ssList(tl(cons(hd(esk51_0),nil)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_112]), c_0_104])).
 % 1.22/1.42  cnf(c_0_118, plain, (memberP(cons(X2,X3),X1)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_100])).
 % 1.22/1.42  cnf(c_0_119, plain, (memberP(X3,X1)|memberP(X2,X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(app(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_113])).
 % 1.22/1.42  cnf(c_0_120, negated_conjecture, (esk51_0=nil|memberP(app(esk51_0,cons(hd(esk51_0),nil)),esk52_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_56]), c_0_32])]), c_0_115]), c_0_104])).
 % 1.22/1.42  cnf(c_0_121, negated_conjecture, (~memberP(esk51_0,esk52_0)), inference(rw,[status(thm)],[c_0_116, c_0_29])).
 % 1.22/1.42  cnf(c_0_122, negated_conjecture, (esk51_0=nil|ssList(cons(hd(esk51_0),nil))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_94]), c_0_36])]), c_0_104])).
 % 1.22/1.42  fof(c_0_123, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_20])).
 % 1.22/1.42  cnf(c_0_124, plain, (memberP(cons(X1,X2),X1)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_118])).
 % 1.22/1.42  cnf(c_0_125, plain, (memberP(X3,X1)|X1=X2|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~memberP(cons(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
 % 1.22/1.42  cnf(c_0_126, negated_conjecture, (esk51_0=nil|memberP(cons(hd(esk51_0),nil),esk52_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_32]), c_0_73])]), c_0_121]), c_0_122])).
 % 1.22/1.42  fof(c_0_127, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_123])])).
 % 1.22/1.42  cnf(c_0_128, plain, (nil=X1|memberP(X1,hd(X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_46]), c_0_39]), c_0_31])).
 % 1.22/1.42  cnf(c_0_129, negated_conjecture, (hd(esk51_0)=esk52_0|esk51_0=nil|memberP(nil,esk52_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_36]), c_0_73])]), c_0_104])).
 % 1.22/1.42  cnf(c_0_130, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
 % 1.22/1.42  cnf(c_0_131, negated_conjecture, (esk51_0=nil|memberP(nil,esk52_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_32])]), c_0_121])).
 % 1.22/1.42  cnf(c_0_132, negated_conjecture, (nil=esk50_0|nil!=esk51_0), inference(split_conjunct,[status(thm)],[c_0_22])).
 % 1.22/1.42  cnf(c_0_133, negated_conjecture, (esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131]), c_0_73])])).
 % 1.22/1.42  cnf(c_0_134, negated_conjecture, (esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_132, c_0_133])])).
 % 1.22/1.42  cnf(c_0_135, negated_conjecture, (~memberP(nil,esk52_0)), inference(rw,[status(thm)],[c_0_121, c_0_133])).
 % 1.22/1.42  cnf(c_0_136, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_134]), c_0_135]), ['proof']).
 % 1.22/1.42  # SZS output end CNFRefutation
 % 1.22/1.42  # Training examples: 0 positive, 0 negative
 % 1.22/1.43
 % 1.22/1.43  # -------------------------------------------------
 % 1.22/1.43  # User time                : 1.185 s
 % 1.22/1.43  # System time              : 0.018 s
 % 1.22/1.43  # Total time               : 1.204 s
 % 1.22/1.43  # Maximum resident set size: 1872 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
