   --------------------------------------------------------------------------

                   TSTP Solution File: SWC414+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC414+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n138.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:51 EDT 2016

 % Result     : Theorem 249.82s
 % Output     : CNFRefutation 249.82s
 % Verified   :
 % Statistics : -

 % Comments   :
 %------------------------------------------------------------------------------
 %----WARNING: Could not form TPTP format derivation
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC414+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.22  % Computer   : n138.star.cs.uiowa.edu
 % 0.02/0.22  % Model      : x86_64 x86_64
 % 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.22  % Memory     : 32218.75MB
 % 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.22  % CPULimit   : 300
 % 0.02/0.22  % DateTime   : Tue May 10 20:43:38 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 151.12/151.33  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 151.12/151.33  # and selection function SelectCQPrecWNTNp.
 % 151.12/151.33  #
 % 151.12/151.33  # Preprocessing time       : 0.015 s
 % 151.12/151.33  # Presaturation interreduction done
 % 151.42/151.64  # No success with AutoSched0
 % 151.42/151.64  # Trying AutoSched1 for 73 seconds
 % 224.41/224.69  # AutoSched1-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S00
 % 224.41/224.69  # and selection function NoSelection.
 % 224.41/224.69  #
 % 224.41/224.69  # Preprocessing time       : 0.015 s
 % 224.41/224.69  # Presaturation interreduction done
 % 224.62/224.88  # No success with AutoSched1
 % 224.62/224.88  # Trying AutoSched2 for 23 seconds
 % 247.63/247.88  # AutoSched2-Mode selected heuristic G_E___107_C37_F1_PI_AE_Q4_CS_SP_PS_S0Y
 % 247.63/247.88  # and selection function SelectMaxLComplexAvoidPosPred.
 % 247.63/247.88  #
 % 247.63/247.88  # Preprocessing time       : 0.015 s
 % 247.63/247.88  # Presaturation interreduction done
 % 247.73/247.92  # No success with AutoSched2
 % 247.73/247.92  # Trying AutoSched3 for 17 seconds
 % 249.82/250.01  # AutoSched3-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S0S
 % 249.82/250.01  # and selection function SelectComplexG.
 % 249.82/250.01  #
 % 249.82/250.01  # Preprocessing time       : 0.015 s
 % 249.82/250.01
 % 249.82/250.01  # Proof found!
 % 249.82/250.01  # SZS status Theorem
 % 249.82/250.01  # SZS output start CNFRefutation
 % 249.82/250.01  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((~(ssList(X4))|X2!=X4)|X1!=X3)|(((?[X5]:(ssItem(X5)&?[X6]:(ssItem(X6)&?[X7]:(ssList(X7)&?[X8]:(ssList(X8)&?[X9]:((ssList(X9)&app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9)=X2)&app(app(app(app(X7,cons(X6,nil)),X8),cons(X5,nil)),X9)=X1)))))|![X10]:(ssItem(X10)=>![X11]:(ssItem(X11)=>![X12]:(ssList(X12)=>app(app(cons(X10,nil),cons(X11,nil)),X12)!=X2))))|![X13]:(ssItem(X13)=>![X14]:(ssItem(X14)=>![X15]:((~(ssList(X15))|app(app(cons(X13,nil),cons(X14,nil)),X15)!=X4)|app(app(cons(X14,nil),cons(X13,nil)),X15)!=X3))))&(?[X16]:(ssItem(X16)&?[X17]:(ssItem(X17)&?[X18]:(ssList(X18)&app(app(cons(X16,nil),cons(X17,nil)),X18)=X4)))|![X19]:(ssItem(X19)=>![X20]:(ssItem(X20)=>![X21]:(ssList(X21)=>app(app(cons(X19,nil),cons(X20,nil)),X21)!=X2)))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
 % 249.82/250.01  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 249.82/250.01  fof(c_0_2, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 249.82/250.01  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 249.82/250.01  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001+0.ax', ax84)).
 % 249.82/250.01  fof(c_0_5, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(((~ssList(X4)|X2!=X4)|X1!=X3)|(((?[X5]:(ssItem(X5)&?[X6]:(ssItem(X6)&?[X7]:(ssList(X7)&?[X8]:(ssList(X8)&?[X9]:((ssList(X9)&app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9)=X2)&app(app(app(app(X7,cons(X6,nil)),X8),cons(X5,nil)),X9)=X1)))))|![X10]:(ssItem(X10)=>![X11]:(ssItem(X11)=>![X12]:(ssList(X12)=>app(app(cons(X10,nil),cons(X11,nil)),X12)!=X2))))|![X13]:(ssItem(X13)=>![X14]:(ssItem(X14)=>![X15]:((~ssList(X15)|app(app(cons(X13,nil),cons(X14,nil)),X15)!=X4)|app(app(cons(X14,nil),cons(X13,nil)),X15)!=X3))))&(?[X16]:(ssItem(X16)&?[X17]:(ssItem(X17)&?[X18]:(ssList(X18)&app(app(cons(X16,nil),cons(X17,nil)),X18)=X4)))|![X19]:(ssItem(X19)=>![X20]:(ssItem(X20)=>![X21]:(ssList(X21)=>app(app(cons(X19,nil),cons(X20,nil)),X21)!=X2))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 249.82/250.01  fof(c_0_6, negated_conjecture,
 % (![X26]:![X27]:![X28]:![X29]:![X30]:![X37]:![X38]:![X39]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(((((~ssItem(X37)|(~ssItem(X38)|(~ssList(X39)|app(app(cons(X37,nil),cons(X38,nil)),X39)!=esk51_0)))|(~ssItem(X26)|(~ssItem(X27)|(~ssList(X28)|(~ssList(X29)|((~ssList(X30)|app(app(app(app(X28,cons(X26,nil)),X29),cons(X27,nil)),X30)!=esk49_0)|app(app(app(app(X28,cons(X27,nil)),X29),cons(X26,nil)),X30)!=esk48_0))))))&((ssItem(esk58_0)|(~ssItem(X26)|(~ssItem(X27)|(~ssList(X28)|(~ssList(X29)|((~ssList(X30)|app(app(app(app(X28,cons(X26,nil)),X29),cons(X27,nil)),X30)!=esk49_0)|app(app(app(app(X28,cons(X27,nil)),X29),cons(X26,nil)),X30)!=esk48_0))))))&((ssItem(esk59_0)|(~ssItem(X26)|(~ssItem(X27)|(~ssList(X28)|(~ssList(X29)|((~ssList(X30)|app(app(app(app(X28,cons(X26,nil)),X29),cons(X27,nil)),X30)!=esk49_0)|app(app(app(app(X28,cons(X27,nil)),X29),cons(X26,nil)),X30)!=esk48_0))))))&((ssList(esk60_0)|(~ssItem(X26)|(~ssItem(X27)|(~ssList(X28)|(~ssList(X29)|((~ss
 % List(X30)|app(app(app(app(X28,cons(X26,nil)),X29),cons(X27,nil)),X30)!=esk49_0)|app(app(app(app(X28,cons(X27,nil)),X29),cons(X26,nil)),X30)!=esk48_0))))))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|(~ssItem(X26)|(~ssItem(X27)|(~ssList(X28)|(~ssList(X29)|((~ssList(X30)|app(app(app(app(X28,cons(X26,nil)),X29),cons(X27,nil)),X30)!=esk49_0)|app(app(app(app(X28,cons(X27,nil)),X29),cons(X26,nil)),X30)!=esk48_0))))))))))&((((~ssItem(X37)|(~ssItem(X38)|(~ssList(X39)|app(app(cons(X37,nil),cons(X38,nil)),X39)!=esk51_0)))|ssItem(esk52_0))&((ssItem(esk58_0)|ssItem(esk52_0))&((ssItem(esk59_0)|ssItem(esk52_0))&((ssList(esk60_0)|ssItem(esk52_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|ssItem(esk52_0))))))&((((~ssItem(X37)|(~ssItem(X38)|(~ssList(X39)|app(app(cons(X37,nil),cons(X38,nil)),X39)!=esk51_0)))|ssItem(esk53_0))&((ssItem(esk58_0)|ssItem(esk53_0))&((ssItem(esk59_0)|ssItem(esk53_0))&((ssList(esk60_0)|ssItem(esk53_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49
 % _0|ssItem(esk53_0))))))&((((~ssItem(X37)|(~ssItem(X38)|(~ssList(X39)|app(app(cons(X37,nil),cons(X38,nil)),X39)!=esk51_0)))|ssList(esk54_0))&((ssItem(esk58_0)|ssList(esk54_0))&((ssItem(esk59_0)|ssList(esk54_0))&((ssList(esk60_0)|ssList(esk54_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|ssList(esk54_0))))))&(((~ssItem(X37)|(~ssItem(X38)|(~ssList(X39)|app(app(cons(X37,nil),cons(X38,nil)),X39)!=esk51_0)))|app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0)&((ssItem(esk58_0)|app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0)&((ssItem(esk59_0)|app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0)&((ssList(esk60_0)|app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0)&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0)))))))))&((((~ssItem(X37)|(~ssItem(X38)|(~ssList(X39)|app(app(cons(X37,nil),cons(X38,nil)),X39)!=esk51_0)))|ssItem(esk55_0))&((ssItem(esk58_0)|ssItem(esk55_0))&((ssIt
 % em(esk59_0)|ssItem(esk55_0))&((ssList(esk60_0)|ssItem(esk55_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|ssItem(esk55_0))))))&((((~ssItem(X37)|(~ssItem(X38)|(~ssList(X39)|app(app(cons(X37,nil),cons(X38,nil)),X39)!=esk51_0)))|ssItem(esk56_0))&((ssItem(esk58_0)|ssItem(esk56_0))&((ssItem(esk59_0)|ssItem(esk56_0))&((ssList(esk60_0)|ssItem(esk56_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|ssItem(esk56_0))))))&(((((~ssItem(X37)|(~ssItem(X38)|(~ssList(X39)|app(app(cons(X37,nil),cons(X38,nil)),X39)!=esk51_0)))|ssList(esk57_0))&((ssItem(esk58_0)|ssList(esk57_0))&((ssItem(esk59_0)|ssList(esk57_0))&((ssList(esk60_0)|ssList(esk57_0))&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|ssList(esk57_0))))))&(((~ssItem(X37)|(~ssItem(X38)|(~ssList(X39)|app(app(cons(X37,nil),cons(X38,nil)),X39)!=esk51_0)))|app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0)&((ssItem(esk58_0)|app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0)&((ssItem(esk59_0)|app(a
 % pp(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0)&((ssList(esk60_0)|app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0)&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0))))))&(((~ssItem(X37)|(~ssItem(X38)|(~ssList(X39)|app(app(cons(X37,nil),cons(X38,nil)),X39)!=esk51_0)))|app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0)&((ssItem(esk58_0)|app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0)&((ssItem(esk59_0)|app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0)&((ssList(esk60_0)|app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0)&(app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0|app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0)))))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])).
 % 249.82/250.01  fof(c_0_7, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
 % 249.82/250.01  cnf(c_0_8, negated_conjecture, (ssItem(esk55_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_9, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_10, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
 % 249.82/250.01  cnf(c_0_11, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 249.82/250.01  cnf(c_0_12, negated_conjecture, (ssItem(esk55_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_13, negated_conjecture, (app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk51_0|ssItem(esk55_0)), inference(rw,[status(thm)],[c_0_8, c_0_9])).
 % 249.82/250.01  cnf(c_0_14, negated_conjecture, (ssItem(esk55_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_15, negated_conjecture, (ssItem(esk55_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_16, negated_conjecture, (ssItem(esk55_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_17, negated_conjecture, (ssItem(esk56_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  fof(c_0_18, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
 % 249.82/250.01  cnf(c_0_19, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
 % 249.82/250.01  cnf(c_0_20, negated_conjecture, (ssItem(esk55_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15]), c_0_16])).
 % 249.82/250.01  cnf(c_0_21, negated_conjecture, (ssItem(esk56_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_22, negated_conjecture, (app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk51_0|ssItem(esk56_0)), inference(rw,[status(thm)],[c_0_17, c_0_9])).
 % 249.82/250.01  cnf(c_0_23, negated_conjecture, (ssItem(esk56_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_24, negated_conjecture, (ssItem(esk56_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_25, negated_conjecture, (ssItem(esk56_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_26, negated_conjecture, (app(app(app(app(X1,cons(X2,nil)),X3),cons(X4,nil)),X5)!=esk48_0|app(app(app(app(X1,cons(X4,nil)),X3),cons(X2,nil)),X5)!=esk49_0|~ssList(X5)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~ssItem(X4)|app(app(cons(X6,nil),cons(X7,nil)),X8)!=esk51_0|~ssList(X8)|~ssItem(X7)|~ssItem(X6)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_27, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
 % 249.82/250.01  cnf(c_0_28, negated_conjecture, (ssList(cons(esk55_0,nil))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
 % 249.82/250.01  cnf(c_0_29, negated_conjecture, (ssItem(esk56_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24]), c_0_25])).
 % 249.82/250.01  cnf(c_0_30, negated_conjecture, (app(app(app(app(X1,cons(X2,nil)),X3),cons(X4,nil)),X5)!=esk51_0|app(app(app(app(X1,cons(X4,nil)),X3),cons(X2,nil)),X5)!=esk48_0|app(app(cons(X6,nil),cons(X7,nil)),X8)!=esk51_0|~ssList(X8)|~ssList(X5)|~ssList(X3)|~ssList(X1)|~ssItem(X7)|~ssItem(X6)|~ssItem(X2)|~ssItem(X4)), inference(rw,[status(thm)],[c_0_26, c_0_9])).
 % 249.82/250.01  cnf(c_0_31, negated_conjecture, (app(nil,cons(esk55_0,nil))=cons(esk55_0,nil)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
 % 249.82/250.01  cnf(c_0_32, negated_conjecture, (ssList(cons(esk56_0,nil))), inference(spm,[status(thm)],[c_0_19, c_0_29])).
 % 249.82/250.01  fof(c_0_33, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])).
 % 249.82/250.01  cnf(c_0_34, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_35, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_36, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_37, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_38, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk49_0|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_39, negated_conjecture, (ssList(esk54_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_40, negated_conjecture, (ssItem(esk53_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_41, negated_conjecture, (ssItem(esk52_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_42, negated_conjecture, (app(app(app(app(nil,cons(X1,nil)),X2),cons(esk55_0,nil)),X3)!=esk48_0|app(app(app(cons(esk55_0,nil),X2),cons(X1,nil)),X3)!=esk51_0|app(app(cons(X4,nil),cons(X5,nil)),X6)!=esk51_0|~ssList(X6)|~ssList(X3)|~ssList(X2)|~ssItem(X5)|~ssItem(X4)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_11]), c_0_20])])).
 % 249.82/250.01  cnf(c_0_43, negated_conjecture, (app(nil,cons(esk56_0,nil))=cons(esk56_0,nil)), inference(spm,[status(thm)],[c_0_27, c_0_32])).
 % 249.82/250.01  cnf(c_0_44, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
 % 249.82/250.01  cnf(c_0_45, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk50_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_46, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_47, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk51_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_34, c_0_9])).
 % 249.82/250.01  cnf(c_0_48, negated_conjecture, (app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk51_0|app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk51_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_9]), c_0_9])).
 % 249.82/250.01  cnf(c_0_49, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk51_0|ssItem(esk58_0)), inference(rw,[status(thm)],[c_0_36, c_0_9])).
 % 249.82/250.01  cnf(c_0_50, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk51_0|ssItem(esk59_0)), inference(rw,[status(thm)],[c_0_37, c_0_9])).
 % 249.82/250.01  cnf(c_0_51, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk51_0|ssList(esk60_0)), inference(rw,[status(thm)],[c_0_38, c_0_9])).
 % 249.82/250.01  cnf(c_0_52, negated_conjecture, (ssList(esk54_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_53, negated_conjecture, (app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk51_0|ssList(esk54_0)), inference(rw,[status(thm)],[c_0_39, c_0_9])).
 % 249.82/250.01  cnf(c_0_54, negated_conjecture, (ssList(esk54_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_55, negated_conjecture, (ssList(esk54_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_56, negated_conjecture, (ssList(esk54_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_57, negated_conjecture, (ssItem(esk53_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_58, negated_conjecture, (app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk51_0|ssItem(esk53_0)), inference(rw,[status(thm)],[c_0_40, c_0_9])).
 % 249.82/250.01  cnf(c_0_59, negated_conjecture, (ssItem(esk53_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_60, negated_conjecture, (ssItem(esk53_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_61, negated_conjecture, (ssItem(esk53_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_62, negated_conjecture, (ssItem(esk52_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_63, negated_conjecture, (app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk51_0|ssItem(esk52_0)), inference(rw,[status(thm)],[c_0_41, c_0_9])).
 % 249.82/250.01  cnf(c_0_64, negated_conjecture, (ssItem(esk52_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_65, negated_conjecture, (ssItem(esk52_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_66, negated_conjecture, (ssItem(esk52_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_67, negated_conjecture, (ssList(esk57_0)|app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk49_0), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_68, negated_conjecture, (app(app(app(cons(esk56_0,nil),X1),cons(esk55_0,nil)),X2)!=esk48_0|app(app(app(cons(esk55_0,nil),X1),cons(esk56_0,nil)),X2)!=esk51_0|app(app(cons(X3,nil),cons(X4,nil)),X5)!=esk51_0|~ssList(X5)|~ssList(X2)|~ssList(X1)|~ssItem(X4)|~ssItem(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_29])])).
 % 249.82/250.01  cnf(c_0_69, negated_conjecture, (app(cons(esk56_0,nil),nil)=cons(esk56_0,nil)), inference(spm,[status(thm)],[c_0_44, c_0_32])).
 % 249.82/250.01  cnf(c_0_70, negated_conjecture, (app(cons(esk55_0,nil),nil)=cons(esk55_0,nil)), inference(spm,[status(thm)],[c_0_44, c_0_28])).
 % 249.82/250.01  cnf(c_0_71, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk48_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(rw,[status(thm)],[c_0_45, c_0_46])).
 % 249.82/250.01  cnf(c_0_72, negated_conjecture, (app(app(cons(esk52_0,nil),cons(esk53_0,nil)),esk54_0)=esk51_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50]), c_0_51])).
 % 249.82/250.01  cnf(c_0_73, negated_conjecture, (ssList(esk54_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_55]), c_0_56])).
 % 249.82/250.01  cnf(c_0_74, negated_conjecture, (ssItem(esk53_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_60]), c_0_61])).
 % 249.82/250.01  cnf(c_0_75, negated_conjecture, (ssItem(esk52_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_65]), c_0_66])).
 % 249.82/250.01  cnf(c_0_76, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_77, negated_conjecture, (ssList(esk57_0)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_78, negated_conjecture, (app(app(cons(esk58_0,nil),cons(esk59_0,nil)),esk60_0)=esk51_0|ssList(esk57_0)), inference(rw,[status(thm)],[c_0_67, c_0_9])).
 % 249.82/250.01  cnf(c_0_79, negated_conjecture, (ssList(esk57_0)|ssItem(esk58_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_80, negated_conjecture, (ssList(esk57_0)|ssItem(esk59_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_81, negated_conjecture, (ssList(esk57_0)|ssList(esk60_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
 % 249.82/250.01  cnf(c_0_82, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),X1)!=esk48_0|app(app(cons(esk55_0,nil),cons(esk56_0,nil)),X1)!=esk51_0|app(app(cons(X2,nil),cons(X3,nil)),X4)!=esk51_0|~ssList(X4)|~ssList(X1)|~ssItem(X3)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), c_0_11])])).
 % 249.82/250.01  cnf(c_0_83, negated_conjecture, (app(app(cons(esk56_0,nil),cons(esk55_0,nil)),esk57_0)=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), c_0_74]), c_0_75])])).
 % 249.82/250.01  cnf(c_0_84, negated_conjecture, (app(app(cons(esk55_0,nil),cons(esk56_0,nil)),esk57_0)=esk51_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_72]), c_0_73]), c_0_74]), c_0_75])])).
 % 249.82/250.01  cnf(c_0_85, negated_conjecture, (ssList(esk57_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79]), c_0_80]), c_0_81])).
 % 249.82/250.01  cnf(c_0_86, negated_conjecture, (app(app(cons(X1,nil),cons(X2,nil)),X3)!=esk51_0|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84]), c_0_85])])).
 % 249.82/250.01  cnf(c_0_87, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_84]), c_0_85]), c_0_29]), c_0_20])]), ['proof']).
 % 249.82/250.01  # SZS output end CNFRefutation
 % 249.82/250.01  # Training examples: 0 positive, 0 negative
 % 249.82/250.02
 % 249.82/250.02  # -------------------------------------------------
 % 249.82/250.02  # User time                : 244.517 s
 % 249.82/250.02  # System time              : 5.295 s
 % 249.82/250.02  # Total time               : 249.813 s
 % 249.82/250.02  # Maximum resident set size: 1896 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
