   --------------------------------------------------------------------------

                   TSTP Solution File: SWC412-1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC412-1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n149.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:51 EDT 2016

 % Result     : Unsatisfiable 225.76s
 % Output     : CNFRefutation 225.76s
 % Verified   :
 % Statistics : Number of clauses        :  101 (1997 expanded)
 %              Number of leaves         :   52 (1392 expanded)
 %              Depth                    :   20
 %              Number of atoms          :  278 (5174 expanded)
 %              Number of equality atoms :   79 (1139 expanded)
 %              Maximal clause size      :    9 (   3 average)
 %              Maximal term depth       :    4 (   2 average)

 % Comments   :
 %------------------------------------------------------------------------------
 cnf(c_0_0,negated_conjecture,
     ( app(app(cons(sk11,nil),sk13),cons(sk12,nil)) = sk2
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_15)).

 cnf(c_0_1,negated_conjecture,
     ( sk2 = sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

 cnf(c_0_2,negated_conjecture,
     ( ssItem(sk11)
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_12)).

 cnf(c_0_3,negated_conjecture,
     ( ssList(sk13)
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_14)).

 cnf(c_0_4,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2
     | app(app(cons(sk11,nil),sk13),cons(sk12,nil)) = sk2 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_55)).

 cnf(c_0_5,negated_conjecture,
     ( ssList(sk7)
     | app(app(cons(sk11,nil),sk13),cons(sk12,nil)) = sk2 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_51)).

 cnf(c_0_6,negated_conjecture,
     ( ssItem(sk5)
     | app(app(cons(sk11,nil),sk13),cons(sk12,nil)) = sk2 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_28)).

 cnf(c_0_7,negated_conjecture,
     ( ssItem(sk6)
     | app(app(cons(sk11,nil),sk13),cons(sk12,nil)) = sk2 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_45)).

 cnf(c_0_8,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2
     | ssItem(sk11) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_48)).

 cnf(c_0_9,negated_conjecture,
     ( ssList(sk7)
     | ssItem(sk11) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_47)).

 cnf(c_0_10,negated_conjecture,
     ( ssItem(sk5)
     | ssItem(sk11) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_25)).

 cnf(c_0_11,negated_conjecture,
     ( ssItem(sk6)
     | ssItem(sk11) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_42)).

 cnf(c_0_12,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2
     | ssList(sk13) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_54)).

 cnf(c_0_13,negated_conjecture,
     ( ssList(sk7)
     | ssList(sk13) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_50)).

 cnf(c_0_14,negated_conjecture,
     ( ssItem(sk5)
     | ssList(sk13) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_27)).

 cnf(c_0_15,negated_conjecture,
     ( ssItem(sk6)
     | ssList(sk13) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_44)).

 cnf(c_0_16,negated_conjecture,
     ( ssItem(sk12)
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

 cnf(c_0_17,axiom,
     ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2)
     | ~ duplicatefreeP(X5)
     | ~ ssList(X5) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

 cnf(c_0_18,axiom,
     ( app(cons(X1,nil),X2) = cons(X1,X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

 cnf(c_0_19,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2
     | ssItem(sk12) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_53)).

 cnf(c_0_20,negated_conjecture,
     ( ssList(sk7)
     | ssItem(sk12) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_49)).

 cnf(c_0_21,negated_conjecture,
     ( ssItem(sk5)
     | ssItem(sk12) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_26)).

 cnf(c_0_22,negated_conjecture,
     ( ssItem(sk6)
     | ssItem(sk12) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_43)).

 cnf(c_0_23,axiom,
     ( ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

 cnf(c_0_24,axiom,
     ( duplicatefreeP(X1)
     | ssItem(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

 cnf(c_0_25,axiom,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

 cnf(c_0_26,axiom,
     ( ssList(app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

 cnf(c_0_27,negated_conjecture,
     ( ssList(sk2) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

 cnf(c_0_28,negated_conjecture,
     ( ssItem(sk8)
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

 cnf(c_0_29,negated_conjecture,
     ( app(app(cons(X5,nil),X6),cons(X4,nil)) = sk3
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk4
     | ~ ssItem(X4)
     | ~ ssItem(X5)
     | ~ ssList(X6)
     | app(app(cons(X4,nil),X6),cons(X5,nil)) != sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

 cnf(c_0_30,negated_conjecture,
     ( sk1 = sk3 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

 cnf(c_0_31,axiom,
     ( ssList(skaf42(X1,X2)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause53)).

 cnf(c_0_32,negated_conjecture,
     ( ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk4
     | app(app(cons(sk12,nil),sk13),cons(sk11,nil)) != sk1 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_16)).

 cnf(c_0_33,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2
     | ssItem(sk8) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_35)).

 cnf(c_0_34,negated_conjecture,
     ( ssList(sk7)
     | ssItem(sk8) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_34)).

 cnf(c_0_35,negated_conjecture,
     ( ssItem(sk5)
     | ssItem(sk8) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_21)).

 cnf(c_0_36,negated_conjecture,
     ( ssItem(sk6)
     | ssItem(sk8) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_30)).

 cnf(c_0_37,negated_conjecture,
     ( ssItem(sk9)
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

 cnf(c_0_38,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2
     | app(app(cons(sk8,nil),cons(sk9,nil)),sk10) = sk2 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_41)).

 cnf(c_0_39,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2
     | ssItem(sk9) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_39)).

 cnf(c_0_40,negated_conjecture,
     ( ssList(sk7)
     | ssItem(sk9) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_36)).

 cnf(c_0_41,negated_conjecture,
     ( ssItem(sk5)
     | ssItem(sk9) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_22)).

 cnf(c_0_42,negated_conjecture,
     ( ssItem(sk6)
     | ssItem(sk9) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_31)).

 cnf(c_0_43,negated_conjecture,
     ( ssList(sk7)
     | app(app(cons(X2,nil),X3),cons(X1,nil)) = sk3
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | app(app(cons(X1,nil),X3),cons(X2,nil)) != sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_19)).

 cnf(c_0_44,axiom,
     ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
     | ~ ssItem(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause157)).

 cnf(c_0_45,negated_conjecture,
     ( ssList(sk7)
     | app(app(cons(sk12,nil),sk13),cons(sk11,nil)) != sk1 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_52)).

 cnf(c_0_46,negated_conjecture,
     ( ssList(sk10)
     | ~ ssItem(X1)
     | ~ ssItem(X2)
     | ~ ssList(X3)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_10)).

 cnf(c_0_47,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2
     | ssList(sk10) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_40)).

 cnf(c_0_48,negated_conjecture,
     ( ssList(sk7)
     | ssList(sk10) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_37)).

 cnf(c_0_49,negated_conjecture,
     ( ssItem(sk5)
     | ssList(sk10) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_23)).

 cnf(c_0_50,negated_conjecture,
     ( ssItem(sk6)
     | ssList(sk10) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_32)).

 cnf(c_0_51,axiom,
     ( hd(cons(X1,X2)) = X1
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause97)).

 cnf(c_0_52,negated_conjecture,
     ( app(app(cons(sk11,nil),sk13),cons(sk12,nil)) = sk2
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(X3) ),
     inference(rw,[status(thm)],[[4]c_0_0,[5]c_0_1])).

 cnf(c_0_53,negated_conjecture,
     ( ssItem(sk11)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(X3) ),
     inference(rw,[status(thm)],[[6]c_0_2,[7]c_0_1])).

 cnf(c_0_54,negated_conjecture,
     ( ssList(sk13)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(X3) ),
     inference(rw,[status(thm)],[[8]c_0_3,[9]c_0_1])).

 cnf(c_0_55,negated_conjecture,
     ( app(app(cons(sk11,nil),sk13),cons(sk12,nil)) = sk2 ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[10]c_0_52,c_0_4]),c_0_5]),c_0_6]),[11]c_0_7])).

 cnf(c_0_56,negated_conjecture,
     ( ssItem(sk11) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[12]c_0_53,c_0_8]),c_0_9]),c_0_10]),[13]c_0_11])).

 cnf(c_0_57,negated_conjecture,
     ( ssList(sk13) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[14]c_0_54,c_0_12]),c_0_13]),c_0_14]),[15]c_0_15])).

 cnf(c_0_58,negated_conjecture,
     ( ssItem(sk12)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(X3) ),
     inference(rw,[status(thm)],[[16]c_0_16,[17]c_0_1])).

 cnf(c_0_59,plain,
     ( ~ ssItem(X1)
     | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
     inference(er,[status(thm)],[[18]c_0_17])).

 cnf(c_0_60,negated_conjecture,
     ( app(cons(sk11,sk13),cons(sk12,nil)) = sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[19]c_0_55,c_0_18]),c_0_56]),c_0_57])])).

 cnf(c_0_61,negated_conjecture,
     ( ssItem(sk12) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[20]c_0_58,c_0_19]),c_0_20]),c_0_21]),[21]c_0_22])).

 cnf(c_0_62,negated_conjecture,
     ( ~ ssList(app(sk2,cons(sk12,X1)))
     | ~ ssList(cons(sk11,sk13))
     | ~ ssList(X1)
     | ~ duplicatefreeP(app(sk2,cons(sk12,X1))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[22]c_0_59,c_0_60]),c_0_61]),c_0_23])])).

 cnf(c_0_63,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(app(sk2,cons(sk12,X2)))
     | ~ ssList(cons(sk11,sk13))
     | ~ ssList(X2) ),
     inference(spm,[status(thm)],[[23]c_0_62,[24]c_0_24])).

 cnf(c_0_64,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(app(sk2,cons(sk12,X2)))
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_63,c_0_25]),c_0_56]),c_0_57])])).

 cnf(c_0_65,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(cons(sk12,X2))
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_64,c_0_26]),c_0_27])])).

 cnf(c_0_66,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_65,c_0_25]),c_0_61])])).

 cnf(c_0_67,negated_conjecture,
     ( ssItem(sk8)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(X3) ),
     inference(rw,[status(thm)],[[28]c_0_28,[29]c_0_1])).

 cnf(c_0_68,negated_conjecture,
     ( app(app(cons(X1,nil),X2),cons(X3,nil)) = sk1
     | app(app(cons(X3,nil),X2),cons(X1,nil)) != sk2
     | app(app(cons(X4,nil),cons(X5,nil)),X6) != sk2
     | ~ ssItem(X1)
     | ~ ssItem(X3)
     | ~ ssItem(X5)
     | ~ ssItem(X4)
     | ~ ssList(X2)
     | ~ ssList(X6) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[30]c_0_29,c_0_30]),c_0_1]),[31]c_0_1])).

 cnf(c_0_69,negated_conjecture,
     ( ssItem(X1) ),
     inference(spm,[status(thm)],[[32]c_0_66,[33]c_0_31])).

 cnf(c_0_70,negated_conjecture,
     ( app(app(cons(sk12,nil),sk13),cons(sk11,nil)) != sk1
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(X3) ),
     inference(rw,[status(thm)],[[34]c_0_32,[35]c_0_1])).

 cnf(c_0_71,negated_conjecture,
     ( ssItem(sk8) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[36]c_0_67,c_0_33]),c_0_34]),c_0_35]),[37]c_0_36])).

 cnf(c_0_72,negated_conjecture,
     ( ssItem(sk9)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(X3) ),
     inference(rw,[status(thm)],[[38]c_0_37,[39]c_0_1])).

 cnf(c_0_73,negated_conjecture,
     ( app(app(cons(X1,nil),X2),cons(X3,nil)) = sk1
     | app(app(cons(X3,nil),X2),cons(X1,nil)) != sk2
     | app(app(cons(X4,nil),cons(X5,nil)),X6) != sk2
     | ~ ssList(X2)
     | ~ ssList(X6) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[40]c_0_68,c_0_69]),c_0_69]),c_0_69]),c_0_69])])).

 cnf(c_0_74,negated_conjecture,
     ( app(app(cons(sk12,nil),sk13),cons(sk11,nil)) != sk1
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk2
     | ~ ssList(X3) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[41]c_0_70,c_0_69]),c_0_69])])).

 cnf(c_0_75,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2
     | app(cons(sk8,cons(sk9,nil)),sk10) = sk2
     | ~ ssList(cons(sk9,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[42]c_0_38,c_0_18]),c_0_71])])).

 cnf(c_0_76,negated_conjecture,
     ( ssItem(sk9) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[43]c_0_72,c_0_39]),c_0_40]),c_0_41]),[44]c_0_42])).

 cnf(c_0_77,negated_conjecture,
     ( app(app(cons(X1,nil),X2),cons(X3,nil)) = sk1
     | ssList(sk7)
     | app(app(cons(X3,nil),X2),cons(X1,nil)) != sk2
     | ~ ssItem(X1)
     | ~ ssItem(X3)
     | ~ ssList(X2) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[45]c_0_43,c_0_30]),[46]c_0_1])).

 cnf(c_0_78,negated_conjecture,
     ( cons(sk11,app(sk13,cons(sk12,nil))) = sk2
     | ~ ssList(cons(sk12,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[47]c_0_60,c_0_44]),c_0_56]),c_0_57])])).

 cnf(c_0_79,negated_conjecture,
     ( app(app(cons(X1,nil),cons(X2,nil)),X3) != sk2
     | ~ ssList(X3) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[48]c_0_73,c_0_55]),c_0_57])]),[49]c_0_74])).

 cnf(c_0_80,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2
     | app(cons(sk8,cons(sk9,nil)),sk10) = sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[50]c_0_75,c_0_25]),c_0_76]),c_0_23])])).

 cnf(c_0_81,negated_conjecture,
     ( ssList(sk7) ),
     inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[51]c_0_77,c_0_55]),c_0_61]),c_0_56]),c_0_57])]),[52]c_0_45])).

 cnf(c_0_82,negated_conjecture,
     ( ssList(sk10)
     | app(app(cons(X1,nil),cons(X2,nil)),X3) != sk2
     | ~ ssItem(X2)
     | ~ ssItem(X1)
     | ~ ssList(X3) ),
     inference(rw,[status(thm)],[[53]c_0_46,[54]c_0_1])).

 cnf(c_0_83,negated_conjecture,
     ( cons(sk11,app(sk13,cons(sk12,nil))) = sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[55]c_0_78,c_0_25]),c_0_61]),c_0_23])])).

 cnf(c_0_84,plain,
     ( app(cons(X1,X2),X3) = cons(X1,app(X2,X3))
     | ~ ssList(X3)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[56]c_0_44,c_0_69])])).

 cnf(c_0_85,negated_conjecture,
     ( app(cons(sk8,cons(sk9,nil)),sk10) = sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[57]c_0_79,c_0_80]),c_0_81])])).

 cnf(c_0_86,negated_conjecture,
     ( ssList(sk10) ),
     inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[[58]c_0_82,c_0_47]),c_0_48]),c_0_49]),[59]c_0_50])).

 cnf(c_0_87,negated_conjecture,
     ( hd(sk2) = sk11
     | ~ ssList(app(sk13,cons(sk12,nil))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[60]c_0_51,c_0_83]),c_0_56])])).

 cnf(c_0_88,negated_conjecture,
     ( cons(sk8,app(cons(sk9,nil),sk10)) = sk2
     | ~ ssList(cons(sk9,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[61]c_0_84,c_0_85]),c_0_86])])).

 cnf(c_0_89,plain,
     ( ssList(cons(X1,X2))
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[62]c_0_25,c_0_69])])).

 cnf(c_0_90,negated_conjecture,
     ( hd(sk2) = sk11
     | ~ ssList(cons(sk12,nil)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[63]c_0_87,c_0_26]),c_0_57])])).

 cnf(c_0_91,plain,
     ( hd(cons(X1,X2)) = X1
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[64]c_0_51,c_0_69])])).

 cnf(c_0_92,negated_conjecture,
     ( cons(sk8,app(cons(sk9,nil),sk10)) = sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[65]c_0_88,c_0_89]),c_0_23])])).

 cnf(c_0_93,negated_conjecture,
     ( hd(sk2) = sk11 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[66]c_0_90,c_0_25]),c_0_61]),c_0_23])])).

 cnf(c_0_94,negated_conjecture,
     ( sk8 = sk11
     | ~ ssList(app(cons(sk9,nil),sk10)) ),
     inference(rw,[status(thm)],[inference(spm,[status(thm)],[[67]c_0_91,c_0_92]),[68]c_0_93])).

 cnf(c_0_95,plain,
     ( app(cons(X1,nil),X2) = cons(X1,X2)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[69]c_0_18,c_0_69])])).

 cnf(c_0_96,negated_conjecture,
     ( sk8 = sk11
     | ~ ssList(cons(sk9,sk10)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[70]c_0_94,c_0_95]),c_0_86])])).

 cnf(c_0_97,negated_conjecture,
     ( sk8 = sk11 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[71]c_0_96,c_0_89]),c_0_86])])).

 cnf(c_0_98,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2
     | app(app(cons(sk11,nil),cons(sk9,nil)),sk10) = sk2 ),
     inference(rw,[status(thm)],[[72]c_0_38,[73]c_0_97])).

 cnf(c_0_99,negated_conjecture,
     ( app(app(cons(sk5,nil),cons(sk6,nil)),sk7) = sk2 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[74]c_0_79,c_0_98]),c_0_86])])).

 cnf(c_0_100,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[75]c_0_79,c_0_99]),c_0_81])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC412-1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.22  % Computer   : n149.star.cs.uiowa.edu
 % 0.02/0.22  % Model      : x86_64 x86_64
 % 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.22  % Memory     : 32218.75MB
 % 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.22  % CPULimit   : 300
 % 0.02/0.22  % DateTime   : Tue May 10 20:43:38 CDT 2016
 % 0.02/0.22  % CPUTime    :
 % 0.02/0.22  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 151.12/151.33  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
 % 151.12/151.33  # and selection function SelectCQArNTEqFirst.
 % 151.12/151.33  #
 % 151.12/151.33  # Preprocessing time       : 0.010 s
 % 151.12/151.33  # Presaturation interreduction done
 % 151.42/151.60  # No success with AutoSched0
 % 151.42/151.60  # Trying AutoSched1 for 73 seconds
 % 224.44/224.60  # AutoSched1-Mode selected heuristic G_E___208_C41_F1_AE_CS_SP_PS_S00
 % 224.44/224.60  # and selection function NoSelection.
 % 224.44/224.60  #
 % 224.44/224.60  # Preprocessing time       : 0.011 s
 % 224.44/224.60  # Presaturation interreduction done
 % 224.65/224.80  # No success with AutoSched1
 % 224.65/224.80  # Trying AutoSched2 for 23 seconds
 % 225.76/225.94  # AutoSched2-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2p
 % 225.76/225.94  # and selection function SelectCQIArNXTEqFirst.
 % 225.76/225.94  #
 % 225.76/225.94  # Preprocessing time       : 0.011 s
 % 225.76/225.94  # Presaturation interreduction done
 % 225.76/225.94
 % 225.76/225.94  # Proof found!
 % 225.76/225.94  # SZS status Unsatisfiable
 % 225.76/225.94  # SZS output start CNFRefutation
 % 225.76/225.94  cnf(c_0_0, negated_conjecture, (app(app(cons(sk11,nil),sk13),cons(sk12,nil))=sk2|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_15)).
 % 225.76/225.94  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
 % 225.76/225.94  cnf(c_0_2, negated_conjecture, (ssItem(sk11)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_12)).
 % 225.76/225.94  cnf(c_0_3, negated_conjecture, (ssList(sk13)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_14)).
 % 225.76/225.94  cnf(c_0_4, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2|app(app(cons(sk11,nil),sk13),cons(sk12,nil))=sk2), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_55)).
 % 225.76/225.94  cnf(c_0_5, negated_conjecture, (ssList(sk7)|app(app(cons(sk11,nil),sk13),cons(sk12,nil))=sk2), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_51)).
 % 225.76/225.94  cnf(c_0_6, negated_conjecture, (ssItem(sk5)|app(app(cons(sk11,nil),sk13),cons(sk12,nil))=sk2), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_28)).
 % 225.76/225.94  cnf(c_0_7, negated_conjecture, (ssItem(sk6)|app(app(cons(sk11,nil),sk13),cons(sk12,nil))=sk2), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_45)).
 % 225.76/225.94  cnf(c_0_8, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2|ssItem(sk11)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_48)).
 % 225.76/225.94  cnf(c_0_9, negated_conjecture, (ssList(sk7)|ssItem(sk11)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_47)).
 % 225.76/225.94  cnf(c_0_10, negated_conjecture, (ssItem(sk5)|ssItem(sk11)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_25)).
 % 225.76/225.94  cnf(c_0_11, negated_conjecture, (ssItem(sk6)|ssItem(sk11)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_42)).
 % 225.76/225.94  cnf(c_0_12, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2|ssList(sk13)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_54)).
 % 225.76/225.94  cnf(c_0_13, negated_conjecture, (ssList(sk7)|ssList(sk13)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_50)).
 % 225.76/225.94  cnf(c_0_14, negated_conjecture, (ssItem(sk5)|ssList(sk13)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_27)).
 % 225.76/225.94  cnf(c_0_15, negated_conjecture, (ssItem(sk6)|ssList(sk13)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_44)).
 % 225.76/225.94  cnf(c_0_16, negated_conjecture, (ssItem(sk12)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
 % 225.76/225.94  cnf(c_0_17, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
 % 225.76/225.94  cnf(c_0_18, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
 % 225.76/225.94  cnf(c_0_19, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2|ssItem(sk12)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_53)).
 % 225.76/225.94  cnf(c_0_20, negated_conjecture, (ssList(sk7)|ssItem(sk12)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_49)).
 % 225.76/225.94  cnf(c_0_21, negated_conjecture, (ssItem(sk5)|ssItem(sk12)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_26)).
 % 225.76/225.94  cnf(c_0_22, negated_conjecture, (ssItem(sk6)|ssItem(sk12)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_43)).
 % 225.76/225.94  cnf(c_0_23, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
 % 225.76/225.94  cnf(c_0_24, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
 % 225.76/225.94  cnf(c_0_25, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
 % 225.76/225.94  cnf(c_0_26, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
 % 225.76/225.94  cnf(c_0_27, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
 % 225.76/225.94  cnf(c_0_28, negated_conjecture, (ssItem(sk8)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
 % 225.76/225.94  cnf(c_0_29, negated_conjecture, (app(app(cons(X5,nil),X6),cons(X4,nil))=sk3|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk4|~ssItem(X4)|~ssItem(X5)|~ssList(X6)|app(app(cons(X4,nil),X6),cons(X5,nil))!=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
 % 225.76/225.94  cnf(c_0_30, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
 % 225.76/225.94  cnf(c_0_31, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause53)).
 % 225.76/225.94  cnf(c_0_32, negated_conjecture, (~ssItem(X1)|~ssItem(X2)|~ssList(X3)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk4|app(app(cons(sk12,nil),sk13),cons(sk11,nil))!=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_16)).
 % 225.76/225.94  cnf(c_0_33, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2|ssItem(sk8)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_35)).
 % 225.76/225.94  cnf(c_0_34, negated_conjecture, (ssList(sk7)|ssItem(sk8)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_34)).
 % 225.76/225.94  cnf(c_0_35, negated_conjecture, (ssItem(sk5)|ssItem(sk8)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_21)).
 % 225.76/225.94  cnf(c_0_36, negated_conjecture, (ssItem(sk6)|ssItem(sk8)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_30)).
 % 225.76/225.94  cnf(c_0_37, negated_conjecture, (ssItem(sk9)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
 % 225.76/225.94  cnf(c_0_38, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2|app(app(cons(sk8,nil),cons(sk9,nil)),sk10)=sk2), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_41)).
 % 225.76/225.94  cnf(c_0_39, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2|ssItem(sk9)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_39)).
 % 225.76/225.94  cnf(c_0_40, negated_conjecture, (ssList(sk7)|ssItem(sk9)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_36)).
 % 225.76/225.94  cnf(c_0_41, negated_conjecture, (ssItem(sk5)|ssItem(sk9)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_22)).
 % 225.76/225.94  cnf(c_0_42, negated_conjecture, (ssItem(sk6)|ssItem(sk9)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_31)).
 % 225.76/225.94  cnf(c_0_43, negated_conjecture, (ssList(sk7)|app(app(cons(X2,nil),X3),cons(X1,nil))=sk3|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|app(app(cons(X1,nil),X3),cons(X2,nil))!=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_19)).
 % 225.76/225.94  cnf(c_0_44, axiom, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause157)).
 % 225.76/225.94  cnf(c_0_45, negated_conjecture, (ssList(sk7)|app(app(cons(sk12,nil),sk13),cons(sk11,nil))!=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_52)).
 % 225.76/225.94  cnf(c_0_46, negated_conjecture, (ssList(sk10)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_10)).
 % 225.76/225.94  cnf(c_0_47, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2|ssList(sk10)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_40)).
 % 225.76/225.94  cnf(c_0_48, negated_conjecture, (ssList(sk7)|ssList(sk10)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_37)).
 % 225.76/225.94  cnf(c_0_49, negated_conjecture, (ssItem(sk5)|ssList(sk10)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_23)).
 % 225.76/225.94  cnf(c_0_50, negated_conjecture, (ssItem(sk6)|ssList(sk10)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_32)).
 % 225.76/225.94  cnf(c_0_51, axiom, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause97)).
 % 225.76/225.94  cnf(c_0_52, negated_conjecture, (app(app(cons(sk11,nil),sk13),cons(sk12,nil))=sk2|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(X3)), inference(rw,[status(thm)],[c_0_0, c_0_1])).
 % 225.76/225.94  cnf(c_0_53, negated_conjecture, (ssItem(sk11)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(X3)), inference(rw,[status(thm)],[c_0_2, c_0_1])).
 % 225.76/225.94  cnf(c_0_54, negated_conjecture, (ssList(sk13)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(X3)), inference(rw,[status(thm)],[c_0_3, c_0_1])).
 % 225.76/225.94  cnf(c_0_55, negated_conjecture, (app(app(cons(sk11,nil),sk13),cons(sk12,nil))=sk2), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_4]), c_0_5]), c_0_6]), c_0_7])).
 % 225.76/225.94  cnf(c_0_56, negated_conjecture, (ssItem(sk11)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_8]), c_0_9]), c_0_10]), c_0_11])).
 % 225.76/225.94  cnf(c_0_57, negated_conjecture, (ssList(sk13)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_12]), c_0_13]), c_0_14]), c_0_15])).
 % 225.76/225.94  cnf(c_0_58, negated_conjecture, (ssItem(sk12)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(X3)), inference(rw,[status(thm)],[c_0_16, c_0_1])).
 % 225.76/225.94  cnf(c_0_59, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_17])).
 % 225.76/225.94  cnf(c_0_60, negated_conjecture, (app(cons(sk11,sk13),cons(sk12,nil))=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_18]), c_0_56]), c_0_57])])).
 % 225.76/225.94  cnf(c_0_61, negated_conjecture, (ssItem(sk12)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_19]), c_0_20]), c_0_21]), c_0_22])).
 % 225.76/225.94  cnf(c_0_62, negated_conjecture, (~ssList(app(sk2,cons(sk12,X1)))|~ssList(cons(sk11,sk13))|~ssList(X1)|~duplicatefreeP(app(sk2,cons(sk12,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61]), c_0_23])])).
 % 225.76/225.94  cnf(c_0_63, negated_conjecture, (ssItem(X1)|~ssList(app(sk2,cons(sk12,X2)))|~ssList(cons(sk11,sk13))|~ssList(X2)), inference(spm,[status(thm)],[c_0_62, c_0_24])).
 % 225.76/225.94  cnf(c_0_64, negated_conjecture, (ssItem(X1)|~ssList(app(sk2,cons(sk12,X2)))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_25]), c_0_56]), c_0_57])])).
 % 225.76/225.94  cnf(c_0_65, negated_conjecture, (ssItem(X1)|~ssList(cons(sk12,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_26]), c_0_27])])).
 % 225.76/225.94  cnf(c_0_66, negated_conjecture, (ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_25]), c_0_61])])).
 % 225.76/225.94  cnf(c_0_67, negated_conjecture, (ssItem(sk8)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(X3)), inference(rw,[status(thm)],[c_0_28, c_0_1])).
 % 225.76/225.94  cnf(c_0_68, negated_conjecture, (app(app(cons(X1,nil),X2),cons(X3,nil))=sk1|app(app(cons(X3,nil),X2),cons(X1,nil))!=sk2|app(app(cons(X4,nil),cons(X5,nil)),X6)!=sk2|~ssItem(X1)|~ssItem(X3)|~ssItem(X5)|~ssItem(X4)|~ssList(X2)|~ssList(X6)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30]), c_0_1]), c_0_1])).
 % 225.76/225.94  cnf(c_0_69, negated_conjecture, (ssItem(X1)), inference(spm,[status(thm)],[c_0_66, c_0_31])).
 % 225.76/225.94  cnf(c_0_70, negated_conjecture, (app(app(cons(sk12,nil),sk13),cons(sk11,nil))!=sk1|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(X3)), inference(rw,[status(thm)],[c_0_32, c_0_1])).
 % 225.76/225.94  cnf(c_0_71, negated_conjecture, (ssItem(sk8)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_33]), c_0_34]), c_0_35]), c_0_36])).
 % 225.76/225.94  cnf(c_0_72, negated_conjecture, (ssItem(sk9)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(X3)), inference(rw,[status(thm)],[c_0_37, c_0_1])).
 % 225.76/225.94  cnf(c_0_73, negated_conjecture, (app(app(cons(X1,nil),X2),cons(X3,nil))=sk1|app(app(cons(X3,nil),X2),cons(X1,nil))!=sk2|app(app(cons(X4,nil),cons(X5,nil)),X6)!=sk2|~ssList(X2)|~ssList(X6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_69]), c_0_69]), c_0_69]), c_0_69])])).
 % 225.76/225.94  cnf(c_0_74, negated_conjecture, (app(app(cons(sk12,nil),sk13),cons(sk11,nil))!=sk1|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk2|~ssList(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_69]), c_0_69])])).
 % 225.76/225.94  cnf(c_0_75, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2|app(cons(sk8,cons(sk9,nil)),sk10)=sk2|~ssList(cons(sk9,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_18]), c_0_71])])).
 % 225.76/225.94  cnf(c_0_76, negated_conjecture, (ssItem(sk9)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_39]), c_0_40]), c_0_41]), c_0_42])).
 % 225.76/225.94  cnf(c_0_77, negated_conjecture, (app(app(cons(X1,nil),X2),cons(X3,nil))=sk1|ssList(sk7)|app(app(cons(X3,nil),X2),cons(X1,nil))!=sk2|~ssItem(X1)|~ssItem(X3)|~ssList(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_30]), c_0_1])).
 % 225.76/225.94  cnf(c_0_78, negated_conjecture, (cons(sk11,app(sk13,cons(sk12,nil)))=sk2|~ssList(cons(sk12,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_44]), c_0_56]), c_0_57])])).
 % 225.76/225.94  cnf(c_0_79, negated_conjecture, (app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk2|~ssList(X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_55]), c_0_57])]), c_0_74])).
 % 225.76/225.94  cnf(c_0_80, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2|app(cons(sk8,cons(sk9,nil)),sk10)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_25]), c_0_76]), c_0_23])])).
 % 225.76/225.94  cnf(c_0_81, negated_conjecture, (ssList(sk7)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_55]), c_0_61]), c_0_56]), c_0_57])]), c_0_45])).
 % 225.76/225.94  cnf(c_0_82, negated_conjecture, (ssList(sk10)|app(app(cons(X1,nil),cons(X2,nil)),X3)!=sk2|~ssItem(X2)|~ssItem(X1)|~ssList(X3)), inference(rw,[status(thm)],[c_0_46, c_0_1])).
 % 225.76/225.94  cnf(c_0_83, negated_conjecture, (cons(sk11,app(sk13,cons(sk12,nil)))=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_25]), c_0_61]), c_0_23])])).
 % 225.76/225.94  cnf(c_0_84, plain, (app(cons(X1,X2),X3)=cons(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_69])])).
 % 225.76/225.94  cnf(c_0_85, negated_conjecture, (app(cons(sk8,cons(sk9,nil)),sk10)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81])])).
 % 225.76/225.94  cnf(c_0_86, negated_conjecture, (ssList(sk10)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_47]), c_0_48]), c_0_49]), c_0_50])).
 % 225.76/225.94  cnf(c_0_87, negated_conjecture, (hd(sk2)=sk11|~ssList(app(sk13,cons(sk12,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_83]), c_0_56])])).
 % 225.76/225.94  cnf(c_0_88, negated_conjecture, (cons(sk8,app(cons(sk9,nil),sk10))=sk2|~ssList(cons(sk9,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86])])).
 % 225.76/225.94  cnf(c_0_89, plain, (ssList(cons(X1,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_69])])).
 % 225.76/225.94  cnf(c_0_90, negated_conjecture, (hd(sk2)=sk11|~ssList(cons(sk12,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_26]), c_0_57])])).
 % 225.76/225.94  cnf(c_0_91, plain, (hd(cons(X1,X2))=X1|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_69])])).
 % 225.76/225.94  cnf(c_0_92, negated_conjecture, (cons(sk8,app(cons(sk9,nil),sk10))=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_23])])).
 % 225.76/225.94  cnf(c_0_93, negated_conjecture, (hd(sk2)=sk11), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_25]), c_0_61]), c_0_23])])).
 % 225.76/225.94  cnf(c_0_94, negated_conjecture, (sk8=sk11|~ssList(app(cons(sk9,nil),sk10))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_93])).
 % 225.76/225.94  cnf(c_0_95, plain, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_69])])).
 % 225.76/225.94  cnf(c_0_96, negated_conjecture, (sk8=sk11|~ssList(cons(sk9,sk10))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_86])])).
 % 225.76/225.94  cnf(c_0_97, negated_conjecture, (sk8=sk11), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_89]), c_0_86])])).
 % 225.76/225.94  cnf(c_0_98, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2|app(app(cons(sk11,nil),cons(sk9,nil)),sk10)=sk2), inference(rw,[status(thm)],[c_0_38, c_0_97])).
 % 225.76/225.94  cnf(c_0_99, negated_conjecture, (app(app(cons(sk5,nil),cons(sk6,nil)),sk7)=sk2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_98]), c_0_86])])).
 % 225.76/225.94  cnf(c_0_100, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_99]), c_0_81])]), ['proof']).
 % 225.76/225.94  # SZS output end CNFRefutation
 % 225.76/225.94  # Training examples: 0 positive, 0 negative
 % 225.76/225.95
 % 225.76/225.95  # -------------------------------------------------
 % 225.76/225.95  # User time                : 220.863 s
 % 225.76/225.95  # System time              : 4.905 s
 % 225.76/225.95  # Total time               : 225.769 s
 % 225.76/225.95  # Maximum resident set size: 1888 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
