   --------------------------------------------------------------------------

                   TSTP Solution File: SWC402-1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC402-1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n129.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:48 EDT 2016

 % Result     : Unsatisfiable 0.06s
 % Output     : CNFRefutation 0.06s
 % Verified   :
 % Statistics : Number of clauses        :   96 ( 539 expanded)
 %              Number of leaves         :   37 ( 355 expanded)
 %              Depth                    :   20
 %              Number of atoms          :  213 ( 934 expanded)
 %              Number of equality atoms :   50 ( 225 expanded)
 %              Maximal clause size      :    7 (   2 average)
 %              Maximal term depth       :    4 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 cnf(c_0_0,axiom,
     ( rearsegP(X3,X2)
     | app(X1,X2) != X3
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause143)).

 cnf(c_0_1,axiom,
     ( ssList(app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

 cnf(c_0_2,negated_conjecture,
     ( memberP(sk1,sk7) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_14)).

 cnf(c_0_3,negated_conjecture,
     ( sk1 = sk3 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

 cnf(c_0_4,negated_conjecture,
     ( ssList(sk1) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

 cnf(c_0_5,negated_conjecture,
     ( app(app(sk5,sk3),sk6) = sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

 cnf(c_0_6,negated_conjecture,
     ( ssList(sk6) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

 cnf(c_0_7,axiom,
     ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2)
     | ~ duplicatefreeP(X5)
     | ~ ssList(X5) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

 cnf(c_0_8,axiom,
     ( app(skaf42(X1,X2),cons(X2,skaf43(X2,X1))) = X1
     | ~ memberP(X1,X2)
     | ~ ssItem(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause169)).

 cnf(c_0_9,negated_conjecture,
     ( ssItem(sk7) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

 cnf(c_0_10,negated_conjecture,
     ( ssList(sk5) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

 cnf(c_0_11,negated_conjecture,
     ( ssList(sk2) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

 cnf(c_0_12,negated_conjecture,
     ( sk2 = sk4 ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

 cnf(c_0_13,axiom,
     ( frontsegP(X3,X1)
     | app(X1,X2) != X3
     | ~ ssList(X2)
     | ~ ssList(X1)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause144)).

 cnf(c_0_14,axiom,
     ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

 cnf(c_0_15,axiom,
     ( ssList(skaf43(X1,X2)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause52)).

 cnf(c_0_16,axiom,
     ( ssList(skaf42(X1,X2)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause53)).

 cnf(c_0_17,axiom,
     ( X1 = X3
     | app(X1,X2) != app(X3,X2)
     | ~ ssList(X1)
     | ~ ssList(X2)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause151)).

 cnf(c_0_18,axiom,
     ( app(skaf46(X1,X2),X2) = X1
     | ~ rearsegP(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause131)).

 cnf(c_0_19,axiom,
     ( duplicatefreeP(X1)
     | ssItem(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

 cnf(c_0_20,axiom,
     ( ssList(skaf46(X1,X2)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause50)).

 cnf(c_0_21,axiom,
     ( X2 = X3
     | app(X1,X2) != app(X1,X3)
     | ~ ssList(X2)
     | ~ ssList(X1)
     | ~ ssList(X3) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause150)).

 cnf(c_0_22,axiom,
     ( app(X2,skaf45(X1,X2)) = X1
     | ~ frontsegP(X1,X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause132)).

 cnf(c_0_23,axiom,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

 cnf(c_0_24,axiom,
     ( ssList(skaf45(X1,X2)) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause51)).

 cnf(c_0_25,axiom,
     ( rearsegP(app(X3,X1),X2)
     | ~ rearsegP(X1,X2)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause136)).

 cnf(c_0_26,axiom,
     ( rearsegP(X1,nil)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause58)).

 cnf(c_0_27,axiom,
     ( ssList(nil) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

 cnf(c_0_28,axiom,
     ( frontsegP(X1,X1)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause61)).

 cnf(c_0_29,axiom,
     ( app(X1,nil) = X1
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause73)).

 cnf(c_0_30,axiom,
     ( frontsegP(app(X1,X3),X2)
     | ~ frontsegP(X1,X2)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause137)).

 cnf(c_0_31,axiom,
     ( frontsegP(X1,nil)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause60)).

 cnf(c_0_32,axiom,
     ( app(X1,X2) = nil
     | nil != X1
     | nil != X2
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause145)).

 cnf(c_0_33,axiom,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause74)).

 cnf(c_0_34,axiom,
     ( memberP(app(X1,X3),X2)
     | ~ memberP(X1,X2)
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause140)).

 cnf(c_0_35,axiom,
     ( memberP(app(X3,X1),X2)
     | ~ memberP(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssItem(X2) ),
     file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause141)).

 cnf(c_0_36,negated_conjecture,
     ( ~ memberP(sk2,sk7) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_15)).

 cnf(c_0_37,plain,
     ( rearsegP(app(X1,X2),X2)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_0]),[4]c_0_1])).

 cnf(c_0_38,negated_conjecture,
     ( memberP(sk3,sk7) ),
     inference(rw,[status(thm)],[[5]c_0_2,[6]c_0_3])).

 cnf(c_0_39,negated_conjecture,
     ( ssList(sk3) ),
     inference(rw,[status(thm)],[[7]c_0_4,[8]c_0_3])).

 cnf(c_0_40,negated_conjecture,
     ( rearsegP(sk4,sk6)
     | ~ ssList(app(sk5,sk3)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[9]c_0_37,c_0_5]),c_0_6])])).

 cnf(c_0_41,plain,
     ( ~ ssItem(X1)
     | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
     inference(er,[status(thm)],[[10]c_0_7])).

 cnf(c_0_42,negated_conjecture,
     ( app(skaf42(sk3,sk7),cons(sk7,skaf43(sk7,sk3))) = sk3 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[11]c_0_8,c_0_38]),c_0_9]),c_0_39])])).

 cnf(c_0_43,negated_conjecture,
     ( rearsegP(sk4,sk6) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[12]c_0_40,c_0_1]),c_0_10]),c_0_39])])).

 cnf(c_0_44,negated_conjecture,
     ( ssList(sk4) ),
     inference(rw,[status(thm)],[[13]c_0_11,[14]c_0_12])).

 cnf(c_0_45,plain,
     ( frontsegP(app(X1,X2),X1)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]),[15]c_0_1])).

 cnf(c_0_46,negated_conjecture,
     ( app(sk5,app(sk3,sk6)) = sk4 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[16]c_0_5,c_0_14]),c_0_10]),c_0_39]),c_0_6])])).

 cnf(c_0_47,negated_conjecture,
     ( ~ ssList(app(sk3,cons(sk7,X1)))
     | ~ ssList(X1)
     | ~ duplicatefreeP(app(sk3,cons(sk7,X1))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[17]c_0_41,c_0_42]),c_0_9]),c_0_15]),c_0_16])])).

 cnf(c_0_48,negated_conjecture,
     ( X1 = app(sk5,sk3)
     | app(X1,sk6) != sk4
     | ~ ssList(app(sk5,sk3))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[18]c_0_17,c_0_5]),c_0_6])])).

 cnf(c_0_49,negated_conjecture,
     ( app(skaf46(sk4,sk6),sk6) = sk4 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[19]c_0_18,c_0_43]),c_0_6]),c_0_44])])).

 cnf(c_0_50,negated_conjecture,
     ( frontsegP(sk4,sk5)
     | ~ ssList(app(sk3,sk6)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[20]c_0_45,c_0_46]),c_0_10])])).

 cnf(c_0_51,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(app(sk3,cons(sk7,X2)))
     | ~ ssList(X2) ),
     inference(spm,[status(thm)],[[21]c_0_47,[22]c_0_19])).

 cnf(c_0_52,negated_conjecture,
     ( skaf46(sk4,sk6) = app(sk5,sk3)
     | ~ ssList(app(sk5,sk3)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[23]c_0_48,c_0_49]),c_0_20])])).

 cnf(c_0_53,negated_conjecture,
     ( frontsegP(sk4,sk5) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[24]c_0_50,c_0_1]),c_0_39]),c_0_6])])).

 cnf(c_0_54,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(cons(sk7,X2))
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[25]c_0_51,c_0_1]),c_0_39])])).

 cnf(c_0_55,negated_conjecture,
     ( skaf46(sk4,sk6) = app(sk5,sk3) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[26]c_0_52,c_0_1]),c_0_10]),c_0_39])])).

 cnf(c_0_56,negated_conjecture,
     ( X1 = app(sk3,sk6)
     | app(sk5,X1) != sk4
     | ~ ssList(app(sk3,sk6))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[27]c_0_21,c_0_46]),c_0_10])])).

 cnf(c_0_57,negated_conjecture,
     ( app(sk5,skaf45(sk4,sk5)) = sk4 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[28]c_0_22,c_0_53]),c_0_10]),c_0_44])])).

 cnf(c_0_58,negated_conjecture,
     ( ssItem(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[29]c_0_54,c_0_23]),c_0_9])])).

 cnf(c_0_59,negated_conjecture,
     ( ssList(app(sk5,sk3)) ),
     inference(spm,[status(thm)],[[30]c_0_20,[31]c_0_55])).

 cnf(c_0_60,negated_conjecture,
     ( skaf45(sk4,sk5) = app(sk3,sk6)
     | ~ ssList(app(sk3,sk6)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[32]c_0_56,c_0_57]),c_0_24])])).

 cnf(c_0_61,plain,
     ( rearsegP(app(X1,X2),nil)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[33]c_0_25,c_0_26]),c_0_27])])).

 cnf(c_0_62,negated_conjecture,
     ( ssItem(X1) ),
     inference(spm,[status(thm)],[[34]c_0_58,[35]c_0_59])).

 cnf(c_0_63,negated_conjecture,
     ( skaf45(sk4,sk5) = app(sk3,sk6) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[36]c_0_60,c_0_1]),c_0_39]),c_0_6])])).

 cnf(c_0_64,negated_conjecture,
     ( rearsegP(sk3,nil)
     | ~ ssList(cons(sk7,skaf43(sk7,sk3))) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[37]c_0_61,c_0_42]),c_0_16])])).

 cnf(c_0_65,plain,
     ( ssList(cons(X1,X2))
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[38]c_0_23,c_0_62])])).

 cnf(c_0_66,negated_conjecture,
     ( app(sk5,app(app(sk3,sk6),X1)) = app(sk4,X1)
     | ~ ssList(app(sk3,sk6))
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[39]c_0_14,c_0_46]),c_0_10])])).

 cnf(c_0_67,negated_conjecture,
     ( ssList(app(sk3,sk6)) ),
     inference(spm,[status(thm)],[[40]c_0_24,[41]c_0_63])).

 cnf(c_0_68,negated_conjecture,
     ( rearsegP(sk3,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[42]c_0_64,c_0_65]),c_0_15])])).

 cnf(c_0_69,plain,
     ( app(skaf46(X1,nil),nil) = X1
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[43]c_0_18,c_0_26]),c_0_27])])).

 cnf(c_0_70,negated_conjecture,
     ( app(sk5,app(app(sk3,sk6),X1)) = app(sk4,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[44]c_0_66,c_0_67])])).

 cnf(c_0_71,plain,
     ( app(X1,skaf45(X1,X1)) = X1
     | ~ ssList(X1) ),
     inference(spm,[status(thm)],[[45]c_0_22,[46]c_0_28])).

 cnf(c_0_72,negated_conjecture,
     ( app(skaf46(sk3,nil),nil) = sk3 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[47]c_0_18,c_0_68]),c_0_27]),c_0_39])])).

 cnf(c_0_73,plain,
     ( skaf46(X1,nil) = X1
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[48]c_0_29,c_0_69]),c_0_20])])).

 cnf(c_0_74,plain,
     ( X1 = nil
     | app(X2,X1) != X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[49]c_0_21,c_0_29]),c_0_27])])).

 cnf(c_0_75,negated_conjecture,
     ( app(sk4,skaf45(app(sk3,sk6),app(sk3,sk6))) = sk4 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[50]c_0_70,c_0_71]),c_0_46]),c_0_24]),c_0_67])])).

 cnf(c_0_76,plain,
     ( frontsegP(app(X1,X2),nil)
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[51]c_0_30,c_0_31]),c_0_27])])).

 cnf(c_0_77,negated_conjecture,
     ( app(sk3,nil) = sk3 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[52]c_0_72,c_0_73]),c_0_39])])).

 cnf(c_0_78,negated_conjecture,
     ( skaf45(app(sk3,sk6),app(sk3,sk6)) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[53]c_0_74,c_0_75]),c_0_24]),c_0_44])])).

 cnf(c_0_79,plain,
     ( app(nil,nil) = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_32])]),c_0_27])])).

 cnf(c_0_80,negated_conjecture,
     ( frontsegP(sk3,nil) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[54]c_0_76,c_0_77]),c_0_27]),c_0_39])])).

 cnf(c_0_81,plain,
     ( app(nil,app(X1,X2)) = app(X1,X2)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[55]c_0_14,c_0_33]),c_0_27])])).

 cnf(c_0_82,negated_conjecture,
     ( app(app(sk3,sk6),nil) = app(sk3,sk6) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[56]c_0_71,c_0_78]),c_0_67])])).

 cnf(c_0_83,plain,
     ( app(nil,app(nil,X1)) = app(nil,X1)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[57]c_0_14,c_0_79]),c_0_27])])).

 cnf(c_0_84,negated_conjecture,
     ( app(nil,skaf45(sk3,nil)) = sk3 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[58]c_0_22,c_0_80]),c_0_27]),c_0_39])])).

 cnf(c_0_85,plain,
     ( frontsegP(app(X1,app(X2,X3)),app(X1,X2))
     | ~ ssList(X3)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(csr,[status(thm)],[inference(spm,[status(thm)],[[59]c_0_45,c_0_14]),[60]c_0_1])).

 cnf(c_0_86,negated_conjecture,
     ( app(nil,app(sk3,sk6)) = app(sk3,sk6) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[61]c_0_81,c_0_82]),c_0_67]),c_0_27])])).

 cnf(c_0_87,negated_conjecture,
     ( app(nil,sk3) = sk3 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[62]c_0_83,c_0_84]),c_0_24])])).

 cnf(c_0_88,negated_conjecture,
     ( frontsegP(app(sk3,sk6),sk3) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[63]c_0_85,c_0_86]),c_0_87]),c_0_6]),c_0_27]),c_0_39])])).

 cnf(c_0_89,negated_conjecture,
     ( memberP(app(sk3,X1),sk7)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[64]c_0_34,c_0_38]),c_0_9]),c_0_39])])).

 cnf(c_0_90,negated_conjecture,
     ( app(sk3,skaf45(app(sk3,sk6),sk3)) = app(sk3,sk6) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[65]c_0_22,c_0_88]),c_0_39]),c_0_67])])).

 cnf(c_0_91,plain,
     ( memberP(app(X1,X2),X3)
     | ~ memberP(X2,X3)
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[[66]c_0_35,c_0_62])])).

 cnf(c_0_92,negated_conjecture,
     ( memberP(app(sk3,sk6),sk7) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[67]c_0_89,c_0_90]),c_0_24])])).

 cnf(c_0_93,negated_conjecture,
     ( memberP(app(X1,app(sk3,sk6)),sk7)
     | ~ ssList(X1) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[68]c_0_91,c_0_92]),c_0_67])])).

 cnf(c_0_94,negated_conjecture,
     ( ~ memberP(sk4,sk7) ),
     inference(rw,[status(thm)],[[69]c_0_36,[70]c_0_12])).

 cnf(c_0_95,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[71]c_0_93,c_0_46]),c_0_10])]),[72]c_0_94]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC402-1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.22  % Computer   : n129.star.cs.uiowa.edu
 % 0.02/0.22  % Model      : x86_64 x86_64
 % 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.22  % Memory     : 32218.75MB
 % 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.22  % CPULimit   : 300
 % 0.02/0.22  % DateTime   : Tue May 10 20:43:53 CDT 2016
 % 0.02/0.22  % CPUTime    :
 % 0.02/0.22  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.06/0.45  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
 % 0.06/0.45  # and selection function SelectCQArNTEqFirst.
 % 0.06/0.45  #
 % 0.06/0.45  # Preprocessing time       : 0.009 s
 % 0.06/0.45  # Presaturation interreduction done
 % 0.06/0.45
 % 0.06/0.45  # Proof found!
 % 0.06/0.45  # SZS status Unsatisfiable
 % 0.06/0.45  # SZS output start CNFRefutation
 % 0.06/0.45  cnf(c_0_0, axiom, (rearsegP(X3,X2)|app(X1,X2)!=X3|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause143)).
 % 0.06/0.45  cnf(c_0_1, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
 % 0.06/0.45  cnf(c_0_2, negated_conjecture, (memberP(sk1,sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_14)).
 % 0.06/0.45  cnf(c_0_3, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
 % 0.06/0.45  cnf(c_0_4, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
 % 0.06/0.45  cnf(c_0_5, negated_conjecture, (app(app(sk5,sk3),sk6)=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
 % 0.06/0.45  cnf(c_0_6, negated_conjecture, (ssList(sk6)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
 % 0.06/0.45  cnf(c_0_7, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
 % 0.06/0.45  cnf(c_0_8, axiom, (app(skaf42(X1,X2),cons(X2,skaf43(X2,X1)))=X1|~memberP(X1,X2)|~ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause169)).
 % 0.06/0.45  cnf(c_0_9, negated_conjecture, (ssItem(sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
 % 0.06/0.45  cnf(c_0_10, negated_conjecture, (ssList(sk5)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
 % 0.06/0.45  cnf(c_0_11, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
 % 0.06/0.45  cnf(c_0_12, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
 % 0.06/0.45  cnf(c_0_13, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause144)).
 % 0.06/0.45  cnf(c_0_14, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
 % 0.06/0.45  cnf(c_0_15, axiom, (ssList(skaf43(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause52)).
 % 0.06/0.45  cnf(c_0_16, axiom, (ssList(skaf42(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause53)).
 % 0.06/0.45  cnf(c_0_17, axiom, (X1=X3|app(X1,X2)!=app(X3,X2)|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause151)).
 % 0.06/0.45  cnf(c_0_18, axiom, (app(skaf46(X1,X2),X2)=X1|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause131)).
 % 0.06/0.45  cnf(c_0_19, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
 % 0.06/0.45  cnf(c_0_20, axiom, (ssList(skaf46(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause50)).
 % 0.06/0.45  cnf(c_0_21, axiom, (X2=X3|app(X1,X2)!=app(X1,X3)|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause150)).
 % 0.06/0.45  cnf(c_0_22, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause132)).
 % 0.06/0.45  cnf(c_0_23, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
 % 0.06/0.45  cnf(c_0_24, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause51)).
 % 0.06/0.45  cnf(c_0_25, axiom, (rearsegP(app(X3,X1),X2)|~rearsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause136)).
 % 0.06/0.45  cnf(c_0_26, axiom, (rearsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause58)).
 % 0.06/0.45  cnf(c_0_27, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
 % 0.06/0.45  cnf(c_0_28, axiom, (frontsegP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause61)).
 % 0.06/0.45  cnf(c_0_29, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause73)).
 % 0.06/0.45  cnf(c_0_30, axiom, (frontsegP(app(X1,X3),X2)|~frontsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause137)).
 % 0.06/0.45  cnf(c_0_31, axiom, (frontsegP(X1,nil)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause60)).
 % 0.06/0.45  cnf(c_0_32, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause145)).
 % 0.06/0.45  cnf(c_0_33, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause74)).
 % 0.06/0.45  cnf(c_0_34, axiom, (memberP(app(X1,X3),X2)|~memberP(X1,X2)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause140)).
 % 0.06/0.45  cnf(c_0_35, axiom, (memberP(app(X3,X1),X2)|~memberP(X1,X2)|~ssList(X1)|~ssList(X3)|~ssItem(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause141)).
 % 0.06/0.45  cnf(c_0_36, negated_conjecture, (~memberP(sk2,sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_15)).
 % 0.06/0.45  cnf(c_0_37, plain, (rearsegP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_0]), c_0_1])).
 % 0.06/0.45  cnf(c_0_38, negated_conjecture, (memberP(sk3,sk7)), inference(rw,[status(thm)],[c_0_2, c_0_3])).
 % 0.06/0.45  cnf(c_0_39, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_4, c_0_3])).
 % 0.06/0.45  cnf(c_0_40, negated_conjecture, (rearsegP(sk4,sk6)|~ssList(app(sk5,sk3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_5]), c_0_6])])).
 % 0.06/0.45  cnf(c_0_41, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_7])).
 % 0.06/0.45  cnf(c_0_42, negated_conjecture, (app(skaf42(sk3,sk7),cons(sk7,skaf43(sk7,sk3)))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_38]), c_0_9]), c_0_39])])).
 % 0.06/0.45  cnf(c_0_43, negated_conjecture, (rearsegP(sk4,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_1]), c_0_10]), c_0_39])])).
 % 0.06/0.45  cnf(c_0_44, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_11, c_0_12])).
 % 0.06/0.45  cnf(c_0_45, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_1])).
 % 0.06/0.45  cnf(c_0_46, negated_conjecture, (app(sk5,app(sk3,sk6))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_14]), c_0_10]), c_0_39]), c_0_6])])).
 % 0.06/0.45  cnf(c_0_47, negated_conjecture, (~ssList(app(sk3,cons(sk7,X1)))|~ssList(X1)|~duplicatefreeP(app(sk3,cons(sk7,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_9]), c_0_15]), c_0_16])])).
 % 0.06/0.45  cnf(c_0_48, negated_conjecture, (X1=app(sk5,sk3)|app(X1,sk6)!=sk4|~ssList(app(sk5,sk3))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_5]), c_0_6])])).
 % 0.06/0.45  cnf(c_0_49, negated_conjecture, (app(skaf46(sk4,sk6),sk6)=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_43]), c_0_6]), c_0_44])])).
 % 0.06/0.45  cnf(c_0_50, negated_conjecture, (frontsegP(sk4,sk5)|~ssList(app(sk3,sk6))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_10])])).
 % 0.06/0.45  cnf(c_0_51, negated_conjecture, (ssItem(X1)|~ssList(app(sk3,cons(sk7,X2)))|~ssList(X2)), inference(spm,[status(thm)],[c_0_47, c_0_19])).
 % 0.06/0.45  cnf(c_0_52, negated_conjecture, (skaf46(sk4,sk6)=app(sk5,sk3)|~ssList(app(sk5,sk3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_20])])).
 % 0.06/0.45  cnf(c_0_53, negated_conjecture, (frontsegP(sk4,sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_1]), c_0_39]), c_0_6])])).
 % 0.06/0.45  cnf(c_0_54, negated_conjecture, (ssItem(X1)|~ssList(cons(sk7,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_1]), c_0_39])])).
 % 0.06/0.45  cnf(c_0_55, negated_conjecture, (skaf46(sk4,sk6)=app(sk5,sk3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_1]), c_0_10]), c_0_39])])).
 % 0.06/0.45  cnf(c_0_56, negated_conjecture, (X1=app(sk3,sk6)|app(sk5,X1)!=sk4|~ssList(app(sk3,sk6))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_46]), c_0_10])])).
 % 0.06/0.45  cnf(c_0_57, negated_conjecture, (app(sk5,skaf45(sk4,sk5))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_53]), c_0_10]), c_0_44])])).
 % 0.06/0.45  cnf(c_0_58, negated_conjecture, (ssItem(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_23]), c_0_9])])).
 % 0.06/0.45  cnf(c_0_59, negated_conjecture, (ssList(app(sk5,sk3))), inference(spm,[status(thm)],[c_0_20, c_0_55])).
 % 0.06/0.45  cnf(c_0_60, negated_conjecture, (skaf45(sk4,sk5)=app(sk3,sk6)|~ssList(app(sk3,sk6))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_24])])).
 % 0.06/0.45  cnf(c_0_61, plain, (rearsegP(app(X1,X2),nil)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
 % 0.06/0.45  cnf(c_0_62, negated_conjecture, (ssItem(X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
 % 0.06/0.45  cnf(c_0_63, negated_conjecture, (skaf45(sk4,sk5)=app(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_1]), c_0_39]), c_0_6])])).
 % 0.06/0.45  cnf(c_0_64, negated_conjecture, (rearsegP(sk3,nil)|~ssList(cons(sk7,skaf43(sk7,sk3)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_42]), c_0_16])])).
 % 0.06/0.45  cnf(c_0_65, plain, (ssList(cons(X1,X2))|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_62])])).
 % 0.06/0.45  cnf(c_0_66, negated_conjecture, (app(sk5,app(app(sk3,sk6),X1))=app(sk4,X1)|~ssList(app(sk3,sk6))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_46]), c_0_10])])).
 % 0.06/0.45  cnf(c_0_67, negated_conjecture, (ssList(app(sk3,sk6))), inference(spm,[status(thm)],[c_0_24, c_0_63])).
 % 0.06/0.45  cnf(c_0_68, negated_conjecture, (rearsegP(sk3,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_15])])).
 % 0.06/0.45  cnf(c_0_69, plain, (app(skaf46(X1,nil),nil)=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_26]), c_0_27])])).
 % 0.06/0.45  cnf(c_0_70, negated_conjecture, (app(sk5,app(app(sk3,sk6),X1))=app(sk4,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_67])])).
 % 0.06/0.45  cnf(c_0_71, plain, (app(X1,skaf45(X1,X1))=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_22, c_0_28])).
 % 0.06/0.45  cnf(c_0_72, negated_conjecture, (app(skaf46(sk3,nil),nil)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_68]), c_0_27]), c_0_39])])).
 % 0.06/0.45  cnf(c_0_73, plain, (skaf46(X1,nil)=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_69]), c_0_20])])).
 % 0.06/0.45  cnf(c_0_74, plain, (X1=nil|app(X2,X1)!=X2|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_29]), c_0_27])])).
 % 0.06/0.45  cnf(c_0_75, negated_conjecture, (app(sk4,skaf45(app(sk3,sk6),app(sk3,sk6)))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_46]), c_0_24]), c_0_67])])).
 % 0.06/0.45  cnf(c_0_76, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_27])])).
 % 0.06/0.45  cnf(c_0_77, negated_conjecture, (app(sk3,nil)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_39])])).
 % 0.06/0.45  cnf(c_0_78, negated_conjecture, (skaf45(app(sk3,sk6),app(sk3,sk6))=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_24]), c_0_44])])).
 % 0.06/0.45  cnf(c_0_79, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_32])]), c_0_27])])).
 % 0.06/0.45  cnf(c_0_80, negated_conjecture, (frontsegP(sk3,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_27]), c_0_39])])).
 % 0.06/0.45  cnf(c_0_81, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_33]), c_0_27])])).
 % 0.06/0.45  cnf(c_0_82, negated_conjecture, (app(app(sk3,sk6),nil)=app(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_78]), c_0_67])])).
 % 0.06/0.45  cnf(c_0_83, plain, (app(nil,app(nil,X1))=app(nil,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_79]), c_0_27])])).
 % 0.06/0.45  cnf(c_0_84, negated_conjecture, (app(nil,skaf45(sk3,nil))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_80]), c_0_27]), c_0_39])])).
 % 0.06/0.45  cnf(c_0_85, plain, (frontsegP(app(X1,app(X2,X3)),app(X1,X2))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_14]), c_0_1])).
 % 0.06/0.45  cnf(c_0_86, negated_conjecture, (app(nil,app(sk3,sk6))=app(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_67]), c_0_27])])).
 % 0.06/0.45  cnf(c_0_87, negated_conjecture, (app(nil,sk3)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_24])])).
 % 0.06/0.45  cnf(c_0_88, negated_conjecture, (frontsegP(app(sk3,sk6),sk3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87]), c_0_6]), c_0_27]), c_0_39])])).
 % 0.06/0.45  cnf(c_0_89, negated_conjecture, (memberP(app(sk3,X1),sk7)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_38]), c_0_9]), c_0_39])])).
 % 0.06/0.45  cnf(c_0_90, negated_conjecture, (app(sk3,skaf45(app(sk3,sk6),sk3))=app(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_88]), c_0_39]), c_0_67])])).
 % 0.06/0.45  cnf(c_0_91, plain, (memberP(app(X1,X2),X3)|~memberP(X2,X3)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_62])])).
 % 0.06/0.45  cnf(c_0_92, negated_conjecture, (memberP(app(sk3,sk6),sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_24])])).
 % 0.06/0.45  cnf(c_0_93, negated_conjecture, (memberP(app(X1,app(sk3,sk6)),sk7)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_67])])).
 % 0.06/0.45  cnf(c_0_94, negated_conjecture, (~memberP(sk4,sk7)), inference(rw,[status(thm)],[c_0_36, c_0_12])).
 % 0.06/0.45  cnf(c_0_95, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_46]), c_0_10])]), c_0_94]), ['proof']).
 % 0.06/0.45  # SZS output end CNFRefutation
 % 0.06/0.45  # Training examples: 0 positive, 0 negative
 % 0.06/0.45
 % 0.06/0.45  # -------------------------------------------------
 % 0.06/0.45  # User time                : 0.220 s
 % 0.06/0.45  # System time              : 0.006 s
 % 0.06/0.45  # Total time               : 0.226 s
 % 0.06/0.45  # Maximum resident set size: 1864 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
