   --------------------------------------------------------------------------

                   TSTP Solution File: SWC415+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC415+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n175.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:51 EDT 2016

 % Result     : Theorem 0.03s
 % Output     : CNFRefutation 0.03s
 % Verified   :
 % Statistics : Number of formulae       :   26 (  71 expanded)
 %              Number of clauses        :   18 (  27 expanded)
 %              Number of leaves         :    4 (  17 expanded)
 %              Depth                    :   10
 %              Number of atoms          :  123 ( 851 expanded)
 %              Number of equality atoms :   39 ( 287 expanded)
 %              Maximal formula depth    :   19 (   5 average)
 %              Maximal clause size      :   30 (   2 average)
 %              Maximal term depth       :    2 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ( ( ~ neq(X2,nil)
                         | ? [X5] :
                             ( ssList(X5)
                             & X3 != X5
                             & tl(X4) = X5
                             & neq(nil,X4) )
                         | ? [X6] :
                             ( ssList(X6)
                             & tl(X2) = X6
                             & X1 = X6
                             & neq(nil,X2) ) )
                       & ( ~ neq(X2,nil)
                         | neq(X4,nil) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssList(tl(X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

 fof(c_0_2,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_3,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_4,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ( ( ~ neq(X2,nil)
                           | ? [X5] :
                               ( ssList(X5)
                               & X3 != X5
                               & tl(X4) = X5
                               & neq(nil,X4) )
                           | ? [X6] :
                               ( ssList(X6)
                               & tl(X2) = X6
                               & X1 = X6
                               & neq(nil,X2) ) )
                         & ( ~ neq(X2,nil)
                           | neq(X4,nil) ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_5,negated_conjecture,(
     ! [X11,X12] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ( neq(esk49_0,nil)
         | neq(esk49_0,nil) )
       & ( ~ neq(esk51_0,nil)
         | neq(esk49_0,nil) )
       & ( neq(esk49_0,nil)
         | ~ ssList(X11)
         | esk50_0 = X11
         | tl(esk51_0) != X11
         | ~ neq(nil,esk51_0) )
       & ( ~ neq(esk51_0,nil)
         | ~ ssList(X11)
         | esk50_0 = X11
         | tl(esk51_0) != X11
         | ~ neq(nil,esk51_0) )
       & ( neq(esk49_0,nil)
         | ~ ssList(X12)
         | tl(esk49_0) != X12
         | esk48_0 != X12
         | ~ neq(nil,esk49_0) )
       & ( ~ neq(esk51_0,nil)
         | ~ ssList(X12)
         | tl(esk49_0) != X12
         | esk48_0 != X12
         | ~ neq(nil,esk49_0) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).

 cnf(c_0_6,negated_conjecture,
     ( neq(esk49_0,nil)
     | neq(esk49_0,nil) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_5])).

 cnf(c_0_7,negated_conjecture,
     ( ~ neq(nil,esk49_0)
     | esk48_0 != X1
     | tl(esk49_0) != X1
     | ~ ssList(X1)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[5]c_0_5])).

 cnf(c_0_8,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_5])).

 cnf(c_0_9,negated_conjecture,
     ( neq(esk49_0,nil) ),
     inference(cn,[status(thm)],[[7]c_0_6])).

 cnf(c_0_10,negated_conjecture,
     ( esk50_0 = X1
     | ~ neq(nil,esk51_0)
     | tl(esk51_0) != X1
     | ~ ssList(X1)
     | ~ neq(esk51_0,nil) ),
     inference(split_conjunct,[status(thm)],[[8]c_0_5])).

 cnf(c_0_11,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[9]c_0_5])).

 cnf(c_0_12,negated_conjecture,
     ( tl(esk49_0) != esk48_0
     | ~ ssList(tl(esk49_0))
     | ~ neq(nil,esk49_0) ),
     inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[10]c_0_7,c_0_8]),c_0_9])])])).

 fof(c_0_13,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssList(tl(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).

 cnf(c_0_14,negated_conjecture,
     ( ~ ssList(tl(esk49_0))
     | ~ neq(nil,esk49_0) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[[11]c_0_10,c_0_11]),c_0_8]),c_0_8]),c_0_8]),c_0_9])])]),[12]c_0_12])).

 cnf(c_0_15,plain,
     ( ssList(tl(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[13]c_0_13])).

 cnf(c_0_16,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_5])).

 fof(c_0_17,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).

 cnf(c_0_18,negated_conjecture,
     ( esk49_0 = nil
     | ~ neq(nil,esk49_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[15]c_0_14,c_0_15]),c_0_16])])).

 cnf(c_0_19,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[16]c_0_17])).

 cnf(c_0_20,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[17]c_0_3])).

 cnf(c_0_21,plain,
     ( ~ ssList(X1)
     | ~ ssList(X2)
     | X1 != X2
     | ~ neq(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[18]c_0_17])).

 cnf(c_0_22,negated_conjecture,
     ( esk49_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[19]c_0_18,c_0_19]),c_0_16]),c_0_20])])).

 cnf(c_0_23,plain,
     ( ~ ssList(X1)
     | ~ neq(X1,X1) ),
     inference(er,[status(thm)],[[20]c_0_21])).

 cnf(c_0_24,negated_conjecture,
     ( neq(nil,nil) ),
     inference(rw,[status(thm)],[[21]c_0_9,[22]c_0_22])).

 cnf(c_0_25,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[23]c_0_23,c_0_24]),c_0_20])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC415+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.03/0.22  % Computer   : n175.star.cs.uiowa.edu
 % 0.03/0.22  % Model      : x86_64 x86_64
 % 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.03/0.22  % Memory     : 32218.75MB
 % 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.03/0.22  % CPULimit   : 300
 % 0.03/0.22  % DateTime   : Tue May 10 20:43:38 CDT 2016
 % 0.03/0.22  % CPUTime    :
 % 0.03/0.23  # Version: 1.9.1pre011
 % 0.03/0.23  # No SInE strategy applied
 % 0.03/0.23  # Trying AutoSched0 for 151 seconds
 % 0.03/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.03/0.25  # and selection function SelectCQPrecWNTNp.
 % 0.03/0.25  #
 % 0.03/0.25  # Preprocessing time       : 0.014 s
 % 0.03/0.25  # Presaturation interreduction done
 % 0.03/0.25
 % 0.03/0.25  # Proof found!
 % 0.03/0.25  # SZS status Theorem
 % 0.03/0.25  # SZS output start CNFRefutation
 % 0.03/0.25  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((~(neq(X2,nil))|?[X5]:(((ssList(X5)&X3!=X5)&tl(X4)=X5)&neq(nil,X4)))|?[X6]:(((ssList(X6)&tl(X2)=X6)&X1=X6)&neq(nil,X2)))&(~(neq(X2,nil))|neq(X4,nil))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.03/0.25  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
 % 0.03/0.25  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 0.03/0.25  fof(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.03/0.25  fof(c_0_4, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((X2!=X4|X1!=X3)|(((~neq(X2,nil)|?[X5]:(((ssList(X5)&X3!=X5)&tl(X4)=X5)&neq(nil,X4)))|?[X6]:(((ssList(X6)&tl(X2)=X6)&X1=X6)&neq(nil,X2)))&(~neq(X2,nil)|neq(X4,nil)))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.03/0.25  fof(c_0_5, negated_conjecture, (![X11]:![X12]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((esk49_0=esk51_0&esk48_0=esk50_0)&((((neq(esk49_0,nil)|neq(esk49_0,nil))&(~neq(esk51_0,nil)|neq(esk49_0,nil)))&((neq(esk49_0,nil)|(((~ssList(X11)|esk50_0=X11)|tl(esk51_0)!=X11)|~neq(nil,esk51_0)))&(~neq(esk51_0,nil)|(((~ssList(X11)|esk50_0=X11)|tl(esk51_0)!=X11)|~neq(nil,esk51_0)))))&((neq(esk49_0,nil)|(((~ssList(X12)|tl(esk49_0)!=X12)|esk48_0!=X12)|~neq(nil,esk49_0)))&(~neq(esk51_0,nil)|(((~ssList(X12)|tl(esk49_0)!=X12)|esk48_0!=X12)|~neq(nil,esk49_0))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])])).
 % 0.03/0.25  cnf(c_0_6, negated_conjecture, (neq(esk49_0,nil)|neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.03/0.25  cnf(c_0_7, negated_conjecture, (~neq(nil,esk49_0)|esk48_0!=X1|tl(esk49_0)!=X1|~ssList(X1)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.03/0.25  cnf(c_0_8, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.03/0.25  cnf(c_0_9, negated_conjecture, (neq(esk49_0,nil)), inference(cn,[status(thm)],[c_0_6])).
 % 0.03/0.25  cnf(c_0_10, negated_conjecture, (esk50_0=X1|~neq(nil,esk51_0)|tl(esk51_0)!=X1|~ssList(X1)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.03/0.25  cnf(c_0_11, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.03/0.25  cnf(c_0_12, negated_conjecture, (tl(esk49_0)!=esk48_0|~ssList(tl(esk49_0))|~neq(nil,esk49_0)), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_8]), c_0_9])])])).
 % 0.03/0.25  fof(c_0_13, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
 % 0.03/0.25  cnf(c_0_14, negated_conjecture, (~ssList(tl(esk49_0))|~neq(nil,esk49_0)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_11]), c_0_8]), c_0_8]), c_0_8]), c_0_9])])]), c_0_12])).
 % 0.03/0.25  cnf(c_0_15, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
 % 0.03/0.25  cnf(c_0_16, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
 % 0.03/0.25  fof(c_0_17, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
 % 0.03/0.25  cnf(c_0_18, negated_conjecture, (esk49_0=nil|~neq(nil,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])])).
 % 0.03/0.25  cnf(c_0_19, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.03/0.25  cnf(c_0_20, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
 % 0.03/0.25  cnf(c_0_21, plain, (~ssList(X1)|~ssList(X2)|X1!=X2|~neq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
 % 0.03/0.25  cnf(c_0_22, negated_conjecture, (esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_16]), c_0_20])])).
 % 0.03/0.25  cnf(c_0_23, plain, (~ssList(X1)|~neq(X1,X1)), inference(er,[status(thm)],[c_0_21])).
 % 0.03/0.25  cnf(c_0_24, negated_conjecture, (neq(nil,nil)), inference(rw,[status(thm)],[c_0_9, c_0_22])).
 % 0.03/0.25  cnf(c_0_25, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_20])]), ['proof']).
 % 0.03/0.25  # SZS output end CNFRefutation
 % 0.03/0.25  # Training examples: 0 positive, 0 negative
 % 0.03/0.25
 % 0.03/0.25  # -------------------------------------------------
 % 0.03/0.25  # User time                : 0.027 s
 % 0.03/0.25  # System time              : 0.002 s
 % 0.03/0.25  # Total time               : 0.029 s
 % 0.03/0.25  # Maximum resident set size: 1868 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
