   --------------------------------------------------------------------------

                   TSTP Solution File: SWC418+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC418+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n115.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:52 EDT 2016

 % Result     : Theorem 0.02s
 % Output     : CNFRefutation 0.02s
 % Verified   :
 % Statistics : Number of formulae       :   61 ( 164 expanded)
 %              Number of clauses        :   39 (  67 expanded)
 %              Number of leaves         :   11 (  40 expanded)
 %              Depth                    :   14
 %              Number of atoms          :  255 (1961 expanded)
 %              Number of equality atoms :   99 ( 810 expanded)
 %              Maximal formula depth    :   30 (   5 average)
 %              Maximal clause size      :   26 (   3 average)
 %              Maximal term depth       :    5 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ~ neq(X2,nil)
                     | ? [X5] :
                         ( ssItem(X5)
                         & ? [X6] :
                             ( ssList(X6)
                             & ? [X7] :
                                 ( ssList(X7)
                                 & app(app(X6,cons(X5,nil)),X7) = X2
                                 & app(app(X7,cons(X5,nil)),X6) = X1 ) ) )
                     | ( nil != X3
                       & nil = X4 )
                     | ( ? [X8] :
                           ( ssList(X8)
                           & X3 != X8
                           & ? [X9] :
                               ( ssList(X9)
                               & ? [X10] :
                                   ( ssList(X10)
                                   & tl(X4) = X9
                                   & app(X9,X10) = X8
                                   & ? [X11] :
                                       ( ssItem(X11)
                                       & cons(X11,nil) = X10
                                       & hd(X4) = X11
                                       & neq(nil,X4) )
                                   & neq(nil,X4) ) ) )
                       & neq(X4,nil) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(nil,X1) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

 fof(c_0_2,axiom,(
     ssList(nil) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

 fof(c_0_3,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => ssList(cons(X2,X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

 fof(c_0_4,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ssList(app(X1,X2)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

 fof(c_0_5,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssItem(X2)
          => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

 fof(c_0_6,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssList(tl(X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

 fof(c_0_7,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => cons(hd(X1),tl(X1)) = X1 ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax78)).

 fof(c_0_8,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ( nil != X1
        => ssItem(hd(X1)) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax22)).

 fof(c_0_9,axiom,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ( neq(X1,X2)
           <=> X1 != X2 ) ) ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

 fof(c_0_10,axiom,(
     ! [X1] :
       ( ssList(X1)
      => app(X1,nil) = X1 ) ),
     file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

 fof(c_0_11,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ~ neq(X2,nil)
                       | ? [X5] :
                           ( ssItem(X5)
                           & ? [X6] :
                               ( ssList(X6)
                               & ? [X7] :
                                   ( ssList(X7)
                                   & app(app(X6,cons(X5,nil)),X7) = X2
                                   & app(app(X7,cons(X5,nil)),X6) = X1 ) ) )
                       | ( nil != X3
                         & nil = X4 )
                       | ( ? [X8] :
                             ( ssList(X8)
                             & X3 != X8
                             & ? [X9] :
                                 ( ssList(X9)
                                 & ? [X10] :
                                     ( ssList(X10)
                                     & tl(X4) = X9
                                     & app(X9,X10) = X8
                                     & ? [X11] :
                                         ( ssItem(X11)
                                         & cons(X11,nil) = X10
                                         & hd(X4) = X11
                                         & neq(nil,X4) )
                                     & neq(nil,X4) ) ) )
                         & neq(X4,nil) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_12,negated_conjecture,(
     ! [X16,X17,X18,X19,X20,X21,X22] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & neq(esk49_0,nil)
       & ( ~ ssItem(X16)
         | ~ ssList(X17)
         | ~ ssList(X18)
         | app(app(X17,cons(X16,nil)),X18) != esk49_0
         | app(app(X18,cons(X16,nil)),X17) != esk48_0 )
       & ( nil = esk50_0
         | nil != esk51_0 )
       & ( ~ ssList(X19)
         | esk50_0 = X19
         | ~ ssList(X20)
         | ~ ssList(X21)
         | tl(esk51_0) != X20
         | app(X20,X21) != X19
         | ~ ssItem(X22)
         | cons(X22,nil) != X21
         | hd(esk51_0) != X22
         | ~ neq(nil,esk51_0)
         | ~ neq(nil,esk51_0)
         | ~ neq(esk51_0,nil) ) ) ),
     inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).

 cnf(c_0_13,negated_conjecture,
     ( app(app(X1,cons(X2,nil)),X3) != esk48_0
     | app(app(X3,cons(X2,nil)),X1) != esk49_0
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssItem(X2) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_12])).

 cnf(c_0_14,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_12])).

 cnf(c_0_15,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[6]c_0_12])).

 fof(c_0_16,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(nil,X2) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).

 cnf(c_0_17,negated_conjecture,
     ( app(app(X1,cons(X2,nil)),X3) != esk51_0
     | app(app(X3,cons(X2,nil)),X1) != esk50_0
     | ~ ssList(X1)
     | ~ ssList(X3)
     | ~ ssItem(X2) ),
     inference(rw,[status(thm)],[inference(rw,[status(thm)],[[7]c_0_13,c_0_14]),[8]c_0_15])).

 cnf(c_0_18,plain,
     ( app(nil,X1) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[9]c_0_16])).

 cnf(c_0_19,plain,
     ( ssList(nil) ),
     inference(split_conjunct,[status(thm)],[[10]c_0_2])).

 fof(c_0_20,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | ssList(cons(X4,X3)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

 cnf(c_0_21,negated_conjecture,
     ( esk50_0 = X4
     | ~ neq(esk51_0,nil)
     | ~ neq(nil,esk51_0)
     | ~ neq(nil,esk51_0)
     | hd(esk51_0) != X1
     | cons(X1,nil) != X2
     | ~ ssItem(X1)
     | app(X3,X2) != X4
     | tl(esk51_0) != X3
     | ~ ssList(X2)
     | ~ ssList(X3)
     | ~ ssList(X4) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_12])).

 cnf(c_0_22,negated_conjecture,
     ( neq(esk49_0,nil) ),
     inference(split_conjunct,[status(thm)],[[12]c_0_12])).

 fof(c_0_23,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssList(X4)
       | ssList(app(X3,X4)) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

 cnf(c_0_24,negated_conjecture,
     ( app(app(X1,cons(X2,nil)),nil) != esk50_0
     | app(cons(X2,nil),X1) != esk51_0
     | ~ ssList(cons(X2,nil))
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[13]c_0_17,c_0_18]),c_0_19])])).

 cnf(c_0_25,plain,
     ( ssList(cons(X1,X2))
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[14]c_0_20])).

 fof(c_0_26,plain,(
     ! [X3,X4] :
       ( ~ ssList(X3)
       | ~ ssItem(X4)
       | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
     inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

 cnf(c_0_27,negated_conjecture,
     ( esk50_0 = X4
     | hd(esk51_0) != X1
     | tl(esk51_0) != X3
     | app(X3,X2) != X4
     | cons(X1,nil) != X2
     | ~ ssItem(X1)
     | ~ ssList(X4)
     | ~ ssList(X3)
     | ~ ssList(X2)
     | ~ neq(nil,esk51_0)
     | ~ neq(esk51_0,nil) ),
     inference(cn,[status(thm)],[[15]c_0_21])).

 cnf(c_0_28,negated_conjecture,
     ( neq(esk51_0,nil) ),
     inference(rw,[status(thm)],[[16]c_0_22,[17]c_0_14])).

 cnf(c_0_29,plain,
     ( ssList(app(X1,X2))
     | ~ ssList(X2)
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[18]c_0_23])).

 fof(c_0_30,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssList(tl(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).

 cnf(c_0_31,negated_conjecture,
     ( ssList(esk49_0) ),
     inference(split_conjunct,[status(thm)],[[19]c_0_12])).

 cnf(c_0_32,negated_conjecture,
     ( app(app(X1,cons(X2,nil)),nil) != esk50_0
     | app(cons(X2,nil),X1) != esk51_0
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[20]c_0_24,c_0_25]),c_0_19])])).

 cnf(c_0_33,plain,
     ( cons(X1,X2) = app(cons(X1,nil),X2)
     | ~ ssItem(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[21]c_0_26])).

 fof(c_0_34,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | cons(hd(X2),tl(X2)) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

 cnf(c_0_35,negated_conjecture,
     ( app(tl(esk51_0),cons(hd(esk51_0),nil)) = esk50_0
     | ~ ssList(cons(hd(esk51_0),nil))
     | ~ ssList(tl(esk51_0))
     | ~ neq(nil,esk51_0)
     | ~ ssItem(hd(esk51_0)) ),
     inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[[22]c_0_27,c_0_28])])])])])]),[23]c_0_29])).

 cnf(c_0_36,plain,
     ( ssList(tl(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[24]c_0_30])).

 cnf(c_0_37,negated_conjecture,
     ( ssList(esk51_0) ),
     inference(rw,[status(thm)],[[25]c_0_31,[26]c_0_14])).

 cnf(c_0_38,negated_conjecture,
     ( app(app(X1,cons(X2,nil)),nil) != esk50_0
     | cons(X2,X1) != esk51_0
     | ~ ssList(X1)
     | ~ ssItem(X2) ),
     inference(spm,[status(thm)],[[27]c_0_32,[28]c_0_33])).

 cnf(c_0_39,plain,
     ( cons(hd(X1),tl(X1)) = X1
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[29]c_0_34])).

 cnf(c_0_40,negated_conjecture,
     ( app(tl(esk51_0),cons(hd(esk51_0),nil)) = esk50_0
     | esk51_0 = nil
     | ~ ssList(cons(hd(esk51_0),nil))
     | ~ neq(nil,esk51_0)
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[30]c_0_35,c_0_36]),c_0_37])])).

 fof(c_0_41,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | nil = X2
       | ssItem(hd(X2)) ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).

 cnf(c_0_42,negated_conjecture,
     ( esk51_0 = nil
     | app(app(tl(esk51_0),cons(hd(esk51_0),nil)),nil) != esk50_0
     | ~ ssList(tl(esk51_0))
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[[31]c_0_38,c_0_39])]),c_0_37])])).

 cnf(c_0_43,negated_conjecture,
     ( app(tl(esk51_0),cons(hd(esk51_0),nil)) = esk50_0
     | esk51_0 = nil
     | ~ neq(nil,esk51_0)
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[32]c_0_40,c_0_25]),c_0_19])])).

 cnf(c_0_44,plain,
     ( ssItem(hd(X1))
     | nil = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[33]c_0_41])).

 fof(c_0_45,plain,(
     ! [X3,X4] :
       ( ( ~ neq(X3,X4)
         | X3 != X4
         | ~ ssList(X4)
         | ~ ssList(X3) )
       & ( X3 = X4
         | neq(X3,X4)
         | ~ ssList(X4)
         | ~ ssList(X3) ) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).

 cnf(c_0_46,negated_conjecture,
     ( esk51_0 = nil
     | app(app(tl(esk51_0),cons(hd(esk51_0),nil)),nil) != esk50_0
     | ~ ssItem(hd(esk51_0)) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[34]c_0_42,c_0_36]),c_0_37])])).

 cnf(c_0_47,negated_conjecture,
     ( app(tl(esk51_0),cons(hd(esk51_0),nil)) = esk50_0
     | esk51_0 = nil
     | ~ neq(nil,esk51_0) ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[35]c_0_43,c_0_44]),c_0_37])])).

 cnf(c_0_48,plain,
     ( neq(X1,X2)
     | X1 = X2
     | ~ ssList(X1)
     | ~ ssList(X2) ),
     inference(split_conjunct,[status(thm)],[[36]c_0_45])).

 cnf(c_0_49,negated_conjecture,
     ( esk51_0 = nil
     | app(app(tl(esk51_0),cons(hd(esk51_0),nil)),nil) != esk50_0 ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[37]c_0_46,c_0_44]),c_0_37])])).

 cnf(c_0_50,negated_conjecture,
     ( app(tl(esk51_0),cons(hd(esk51_0),nil)) = esk50_0
     | esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[38]c_0_47,c_0_48]),c_0_37]),c_0_19])])).

 fof(c_0_51,plain,(
     ! [X2] :
       ( ~ ssList(X2)
       | app(X2,nil) = X2 ) ),
     inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).

 cnf(c_0_52,negated_conjecture,
     ( ssList(esk48_0) ),
     inference(split_conjunct,[status(thm)],[[39]c_0_12])).

 cnf(c_0_53,negated_conjecture,
     ( esk51_0 = nil
     | app(esk50_0,nil) != esk50_0 ),
     inference(spm,[status(thm)],[[40]c_0_49,[41]c_0_50])).

 cnf(c_0_54,plain,
     ( app(X1,nil) = X1
     | ~ ssList(X1) ),
     inference(split_conjunct,[status(thm)],[[42]c_0_51])).

 cnf(c_0_55,negated_conjecture,
     ( ssList(esk50_0) ),
     inference(rw,[status(thm)],[[43]c_0_52,[44]c_0_15])).

 cnf(c_0_56,plain,
     ( ~ ssList(X1)
     | ~ ssList(X2)
     | X1 != X2
     | ~ neq(X1,X2) ),
     inference(split_conjunct,[status(thm)],[[45]c_0_45])).

 cnf(c_0_57,negated_conjecture,
     ( esk51_0 = nil ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[46]c_0_53,c_0_54]),c_0_55])])).

 cnf(c_0_58,plain,
     ( ~ ssList(X1)
     | ~ neq(X1,X1) ),
     inference(er,[status(thm)],[[47]c_0_56])).

 cnf(c_0_59,negated_conjecture,
     ( neq(nil,nil) ),
     inference(rw,[status(thm)],[[48]c_0_28,[49]c_0_57])).

 cnf(c_0_60,negated_conjecture,
     ( $false ),
     inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[50]c_0_58,c_0_59]),c_0_19])]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC418+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.23  % Computer   : n115.star.cs.uiowa.edu
 % 0.02/0.23  % Model      : x86_64 x86_64
 % 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.23  % Memory     : 32218.75MB
 % 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.23  % CPULimit   : 300
 % 0.02/0.23  % DateTime   : Tue May 10 20:43:37 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.24  # Version: 1.9.1pre011
 % 0.02/0.24  # No SInE strategy applied
 % 0.02/0.24  # Trying AutoSched0 for 151 seconds
 % 0.02/0.28  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.02/0.28  # and selection function SelectCQPrecWNTNp.
 % 0.02/0.28  #
 % 0.02/0.28  # Preprocessing time       : 0.014 s
 % 0.02/0.28  # Presaturation interreduction done
 % 0.02/0.28
 % 0.02/0.28  # Proof found!
 % 0.02/0.28  # SZS status Theorem
 % 0.02/0.28  # SZS output start CNFRefutation
 % 0.02/0.28  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~(neq(X2,nil)))|?[X5]:(ssItem(X5)&?[X6]:(ssList(X6)&?[X7]:((ssList(X7)&app(app(X6,cons(X5,nil)),X7)=X2)&app(app(X7,cons(X5,nil)),X6)=X1))))|(nil!=X3&nil=X4))|(?[X8]:((ssList(X8)&X3!=X8)&?[X9]:(ssList(X9)&?[X10]:((((ssList(X10)&tl(X4)=X9)&app(X9,X10)=X8)&?[X11]:(((ssItem(X11)&cons(X11,nil)=X10)&hd(X4)=X11)&neq(nil,X4)))&neq(nil,X4))))&neq(X4,nil)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
 % 0.02/0.28  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
 % 0.02/0.28  fof(c_0_2, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
 % 0.02/0.28  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
 % 0.02/0.28  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
 % 0.02/0.28  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
 % 0.02/0.28  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
 % 0.02/0.28  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax78)).
 % 0.02/0.28  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax22)).
 % 0.02/0.28  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
 % 0.02/0.28  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
 % 0.02/0.28  fof(c_0_11, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~neq(X2,nil))|?[X5]:(ssItem(X5)&?[X6]:(ssList(X6)&?[X7]:((ssList(X7)&app(app(X6,cons(X5,nil)),X7)=X2)&app(app(X7,cons(X5,nil)),X6)=X1))))|(nil!=X3&nil=X4))|(?[X8]:((ssList(X8)&X3!=X8)&?[X9]:(ssList(X9)&?[X10]:((((ssList(X10)&tl(X4)=X9)&app(X9,X10)=X8)&?[X11]:(((ssItem(X11)&cons(X11,nil)=X10)&hd(X4)=X11)&neq(nil,X4)))&neq(nil,X4))))&neq(X4,nil))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.02/0.28  fof(c_0_12, negated_conjecture, (![X16]:![X17]:![X18]:![X19]:![X20]:![X21]:![X22]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((((esk49_0=esk51_0&esk48_0=esk50_0)&neq(esk49_0,nil))&(~ssItem(X16)|(~ssList(X17)|((~ssList(X18)|app(app(X17,cons(X16,nil)),X18)!=esk49_0)|app(app(X18,cons(X16,nil)),X17)!=esk48_0))))&(nil=esk50_0|nil!=esk51_0))&(((~ssList(X19)|esk50_0=X19)|(~ssList(X20)|((((~ssList(X21)|tl(esk51_0)!=X20)|app(X20,X21)!=X19)|(((~ssItem(X22)|cons(X22,nil)!=X21)|hd(esk51_0)!=X22)|~neq(nil,esk51_0)))|~neq(nil,esk51_0))))|~neq(esk51_0,nil)))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
 % 0.02/0.28  cnf(c_0_13, negated_conjecture, (app(app(X1,cons(X2,nil)),X3)!=esk48_0|app(app(X3,cons(X2,nil)),X1)!=esk49_0|~ssList(X1)|~ssList(X3)|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
 % 0.02/0.28  cnf(c_0_14, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_12])).
 % 0.02/0.28  cnf(c_0_15, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_12])).
 % 0.02/0.28  fof(c_0_16, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])).
 % 0.02/0.28  cnf(c_0_17, negated_conjecture, (app(app(X1,cons(X2,nil)),X3)!=esk51_0|app(app(X3,cons(X2,nil)),X1)!=esk50_0|~ssList(X1)|~ssList(X3)|~ssItem(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14]), c_0_15])).
 % 0.02/0.28  cnf(c_0_18, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
 % 0.02/0.28  cnf(c_0_19, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.02/0.28  fof(c_0_20, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
 % 0.02/0.28  cnf(c_0_21, negated_conjecture, (esk50_0=X4|~neq(esk51_0,nil)|~neq(nil,esk51_0)|~neq(nil,esk51_0)|hd(esk51_0)!=X1|cons(X1,nil)!=X2|~ssItem(X1)|app(X3,X2)!=X4|tl(esk51_0)!=X3|~ssList(X2)|~ssList(X3)|~ssList(X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
 % 0.02/0.28  cnf(c_0_22, negated_conjecture, (neq(esk49_0,nil)), inference(split_conjunct,[status(thm)],[c_0_12])).
 % 0.02/0.28  fof(c_0_23, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
 % 0.02/0.28  cnf(c_0_24, negated_conjecture, (app(app(X1,cons(X2,nil)),nil)!=esk50_0|app(cons(X2,nil),X1)!=esk51_0|~ssList(cons(X2,nil))|~ssList(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
 % 0.02/0.28  cnf(c_0_25, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
 % 0.02/0.28  fof(c_0_26, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
 % 0.02/0.28  cnf(c_0_27, negated_conjecture, (esk50_0=X4|hd(esk51_0)!=X1|tl(esk51_0)!=X3|app(X3,X2)!=X4|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X4)|~ssList(X3)|~ssList(X2)|~neq(nil,esk51_0)|~neq(esk51_0,nil)), inference(cn,[status(thm)],[c_0_21])).
 % 0.02/0.28  cnf(c_0_28, negated_conjecture, (neq(esk51_0,nil)), inference(rw,[status(thm)],[c_0_22, c_0_14])).
 % 0.02/0.28  cnf(c_0_29, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
 % 0.02/0.28  fof(c_0_30, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
 % 0.02/0.28  cnf(c_0_31, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
 % 0.02/0.28  cnf(c_0_32, negated_conjecture, (app(app(X1,cons(X2,nil)),nil)!=esk50_0|app(cons(X2,nil),X1)!=esk51_0|~ssList(X1)|~ssItem(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_19])])).
 % 0.02/0.28  cnf(c_0_33, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
 % 0.02/0.28  fof(c_0_34, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
 % 0.02/0.28  cnf(c_0_35, negated_conjecture, (app(tl(esk51_0),cons(hd(esk51_0),nil))=esk50_0|~ssList(cons(hd(esk51_0),nil))|~ssList(tl(esk51_0))|~neq(nil,esk51_0)|~ssItem(hd(esk51_0))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28])])])])])]), c_0_29])).
 % 0.02/0.28  cnf(c_0_36, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
 % 0.02/0.28  cnf(c_0_37, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_31, c_0_14])).
 % 0.02/0.28  cnf(c_0_38, negated_conjecture, (app(app(X1,cons(X2,nil)),nil)!=esk50_0|cons(X2,X1)!=esk51_0|~ssList(X1)|~ssItem(X2)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
 % 0.02/0.28  cnf(c_0_39, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
 % 0.02/0.28  cnf(c_0_40, negated_conjecture, (app(tl(esk51_0),cons(hd(esk51_0),nil))=esk50_0|esk51_0=nil|~ssList(cons(hd(esk51_0),nil))|~neq(nil,esk51_0)|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])])).
 % 0.02/0.28  fof(c_0_41, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
 % 0.02/0.28  cnf(c_0_42, negated_conjecture, (esk51_0=nil|app(app(tl(esk51_0),cons(hd(esk51_0),nil)),nil)!=esk50_0|~ssList(tl(esk51_0))|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39])]), c_0_37])])).
 % 0.02/0.28  cnf(c_0_43, negated_conjecture, (app(tl(esk51_0),cons(hd(esk51_0),nil))=esk50_0|esk51_0=nil|~neq(nil,esk51_0)|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_25]), c_0_19])])).
 % 0.02/0.28  cnf(c_0_44, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
 % 0.02/0.28  fof(c_0_45, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
 % 0.02/0.28  cnf(c_0_46, negated_conjecture, (esk51_0=nil|app(app(tl(esk51_0),cons(hd(esk51_0),nil)),nil)!=esk50_0|~ssItem(hd(esk51_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_36]), c_0_37])])).
 % 0.02/0.28  cnf(c_0_47, negated_conjecture, (app(tl(esk51_0),cons(hd(esk51_0),nil))=esk50_0|esk51_0=nil|~neq(nil,esk51_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_37])])).
 % 0.02/0.28  cnf(c_0_48, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
 % 0.02/0.28  cnf(c_0_49, negated_conjecture, (esk51_0=nil|app(app(tl(esk51_0),cons(hd(esk51_0),nil)),nil)!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_44]), c_0_37])])).
 % 0.02/0.28  cnf(c_0_50, negated_conjecture, (app(tl(esk51_0),cons(hd(esk51_0),nil))=esk50_0|esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_37]), c_0_19])])).
 % 0.02/0.28  fof(c_0_51, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])).
 % 0.02/0.28  cnf(c_0_52, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
 % 0.02/0.28  cnf(c_0_53, negated_conjecture, (esk51_0=nil|app(esk50_0,nil)!=esk50_0), inference(spm,[status(thm)],[c_0_49, c_0_50])).
 % 0.02/0.28  cnf(c_0_54, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
 % 0.02/0.28  cnf(c_0_55, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_52, c_0_15])).
 % 0.02/0.28  cnf(c_0_56, plain, (~ssList(X1)|~ssList(X2)|X1!=X2|~neq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
 % 0.02/0.28  cnf(c_0_57, negated_conjecture, (esk51_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])])).
 % 0.02/0.28  cnf(c_0_58, plain, (~ssList(X1)|~neq(X1,X1)), inference(er,[status(thm)],[c_0_56])).
 % 0.02/0.28  cnf(c_0_59, negated_conjecture, (neq(nil,nil)), inference(rw,[status(thm)],[c_0_28, c_0_57])).
 % 0.02/0.28  cnf(c_0_60, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_19])]), ['proof']).
 % 0.02/0.28  # SZS output end CNFRefutation
 % 0.02/0.28  # Training examples: 0 positive, 0 negative
 % 0.02/0.28
 % 0.02/0.28  # -------------------------------------------------
 % 0.02/0.28  # User time                : 0.043 s
 % 0.02/0.28  # System time              : 0.004 s
 % 0.02/0.28  # Total time               : 0.047 s
 % 0.02/0.28  # Maximum resident set size: 1880 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
