   --------------------------------------------------------------------------

                   TSTP Solution File: SWC406+1 by E---1.9.1

                     [1]View Problem - [2]Process Solution

   --------------------------------------------------------------------------

 %------------------------------------------------------------------------------
 % File       : E---1.9.1
 % Problem    : [3]SWC406+1 : TPTP v6.4.0. Released v2.4.0.
 % Transform  : none
 % Format     : tptp:raw
 % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

 % Computer   : n098.star.cs.uiowa.edu
 % Model      : x86_64 x86_64
 % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
 % Memory     : 32218.75MB
 % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % CPULimit   : 300s
 % DateTime   : Wed May 11 11:17:49 EDT 2016

 % Result     : Theorem 0.02s
 % Output     : CNFRefutation 0.02s
 % Verified   :
 % Statistics : Number of formulae       :   12 (  27 expanded)
 %              Number of clauses        :    9 (   9 expanded)
 %              Number of leaves         :    1 (   6 expanded)
 %              Depth                    :    5
 %              Number of atoms          :   89 ( 479 expanded)
 %              Number of equality atoms :   14 (  74 expanded)
 %              Maximal formula depth    :   23 (   7 average)
 %              Maximal clause size      :   34 (   1 average)
 %              Maximal term depth       :    2 (   1 average)

 % Comments   :
 %------------------------------------------------------------------------------
 fof(c_0_0,conjecture,(
     ! [X1] :
       ( ssList(X1)
      => ! [X2] :
           ( ssList(X2)
          => ! [X3] :
               ( ssList(X3)
              => ! [X4] :
                   ( ssList(X4)
                  => ( X2 != X4
                     | X1 != X3
                     | ? [X5] :
                         ( ssItem(X5)
                         & ( ( ~ memberP(X3,X5)
                             & ! [X6] :
                                 ( ssItem(X6)
                                => ( ~ memberP(X4,X6)
                                   | ~ leq(X6,X5)
                                   | X5 = X6 ) )
                             & memberP(X4,X5) )
                           | ( memberP(X3,X5)
                             & ( ~ memberP(X4,X5)
                               | ? [X6] :
                                   ( ssItem(X6)
                                   & X5 != X6
                                   & memberP(X4,X6)
                                   & leq(X6,X5) ) ) ) ) )
                     | ! [X7] :
                         ( ssItem(X7)
                        => ( ~ memberP(X1,X7)
                           | memberP(X2,X7) ) ) ) ) ) ) ) ),
     file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1)).

 fof(c_0_1,negated_conjecture,(
     ~ ! [X1] :
         ( ssList(X1)
        => ! [X2] :
             ( ssList(X2)
            => ! [X3] :
                 ( ssList(X3)
                => ! [X4] :
                     ( ssList(X4)
                    => ( X2 != X4
                       | X1 != X3
                       | ? [X5] :
                           ( ssItem(X5)
                           & ( ( ~ memberP(X3,X5)
                               & ! [X6] :
                                   ( ssItem(X6)
                                  => ( ~ memberP(X4,X6)
                                     | ~ leq(X6,X5)
                                     | X5 = X6 ) )
                               & memberP(X4,X5) )
                             | ( memberP(X3,X5)
                               & ( ~ memberP(X4,X5)
                                 | ? [X6] :
                                     ( ssItem(X6)
                                     & X5 != X6
                                     & memberP(X4,X6)
                                     & leq(X6,X5) ) ) ) ) )
                       | ! [X7] :
                           ( ssItem(X7)
                          => ( ~ memberP(X1,X7)
                             | memberP(X2,X7) ) ) ) ) ) ) ) ),
     inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

 fof(c_0_2,negated_conjecture,(
     ! [X12,X14] :
       ( ssList(esk48_0)
       & ssList(esk49_0)
       & ssList(esk50_0)
       & ssList(esk51_0)
       & esk49_0 = esk51_0
       & esk48_0 = esk50_0
       & ( ssItem(esk52_1(X12))
         | memberP(esk50_0,X12)
         | ~ memberP(esk51_0,X12)
         | ~ ssItem(X12) )
       & ( memberP(esk51_0,esk52_1(X12))
         | memberP(esk50_0,X12)
         | ~ memberP(esk51_0,X12)
         | ~ ssItem(X12) )
       & ( leq(esk52_1(X12),X12)
         | memberP(esk50_0,X12)
         | ~ memberP(esk51_0,X12)
         | ~ ssItem(X12) )
       & ( X12 != esk52_1(X12)
         | memberP(esk50_0,X12)
         | ~ memberP(esk51_0,X12)
         | ~ ssItem(X12) )
       & ( memberP(esk51_0,X12)
         | ~ memberP(esk50_0,X12)
         | ~ ssItem(X12) )
       & ( ~ ssItem(X14)
         | X12 = X14
         | ~ memberP(esk51_0,X14)
         | ~ leq(X14,X12)
         | ~ memberP(esk50_0,X12)
         | ~ ssItem(X12) )
       & ssItem(esk53_0)
       & memberP(esk48_0,esk53_0)
       & ~ memberP(esk49_0,esk53_0) ) ),
     inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

 cnf(c_0_3,negated_conjecture,
     ( memberP(esk48_0,esk53_0) ),
     inference(split_conjunct,[status(thm)],[[4]c_0_2])).

 cnf(c_0_4,negated_conjecture,
     ( esk48_0 = esk50_0 ),
     inference(split_conjunct,[status(thm)],[[5]c_0_2])).

 cnf(c_0_5,negated_conjecture,
     ( ~ memberP(esk49_0,esk53_0) ),
     inference(split_conjunct,[status(thm)],[[6]c_0_2])).

 cnf(c_0_6,negated_conjecture,
     ( esk49_0 = esk51_0 ),
     inference(split_conjunct,[status(thm)],[[7]c_0_2])).

 cnf(c_0_7,negated_conjecture,
     ( memberP(esk51_0,X1)
     | ~ ssItem(X1)
     | ~ memberP(esk50_0,X1) ),
     inference(split_conjunct,[status(thm)],[[8]c_0_2])).

 cnf(c_0_8,negated_conjecture,
     ( memberP(esk50_0,esk53_0) ),
     inference(rw,[status(thm)],[[9]c_0_3,[10]c_0_4])).

 cnf(c_0_9,negated_conjecture,
     ( ssItem(esk53_0) ),
     inference(split_conjunct,[status(thm)],[[11]c_0_2])).

 cnf(c_0_10,negated_conjecture,
     ( ~ memberP(esk51_0,esk53_0) ),
     inference(rw,[status(thm)],[[12]c_0_5,[13]c_0_6])).

 cnf(c_0_11,negated_conjecture,
     ( $false ),
     inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[[14]c_0_7,c_0_8]),c_0_9])]),[15]c_0_10]),
     [proof]).
 %------------------------------------------------------------------------------
 %----ORIGINAL SYSTEM OUTPUT
 % 0.00/0.03  % Problem    : SWC406+1 : TPTP v6.4.0. Released v2.4.0.
 % 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
 % 0.02/0.23  % Computer   : n098.star.cs.uiowa.edu
 % 0.02/0.23  % Model      : x86_64 x86_64
 % 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
 % 0.02/0.23  % Memory     : 32218.75MB
 % 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
 % 0.02/0.23  % CPULimit   : 300
 % 0.02/0.23  % DateTime   : Tue May 10 20:43:38 CDT 2016
 % 0.02/0.23  % CPUTime    :
 % 0.02/0.23  # Version: 1.9.1pre011
 % 0.02/0.23  # No SInE strategy applied
 % 0.02/0.23  # Trying AutoSched0 for 151 seconds
 % 0.02/0.25  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
 % 0.02/0.25  # and selection function SelectCQPrecWNTNp.
 % 0.02/0.25  #
 % 0.02/0.25  # Preprocessing time       : 0.014 s
 % 0.02/0.25  # Presaturation interreduction done
 % 0.02/0.25
 % 0.02/0.25  # Proof found!
 % 0.02/0.25  # SZS status Theorem
 % 0.02/0.25  # SZS output start CNFRefutation
 % 0.02/0.25  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|?[X5]:(ssItem(X5)&(((~(memberP(X3,X5))&![X6]:(ssItem(X6)=>((~(memberP(X4,X6))|~(leq(X6,X5)))|X5=X6)))&memberP(X4,X5))|(memberP(X3,X5)&(~(memberP(X4,X5))|?[X6]:(((ssItem(X6)&X5!=X6)&memberP(X4,X6))&leq(X6,X5)))))))|![X7]:(ssItem(X7)=>(~(memberP(X1,X7))|memberP(X2,X7))))))))), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1)).
 % 0.02/0.25  fof(c_0_1, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|?[X5]:(ssItem(X5)&(((~memberP(X3,X5)&![X6]:(ssItem(X6)=>((~memberP(X4,X6)|~leq(X6,X5))|X5=X6)))&memberP(X4,X5))|(memberP(X3,X5)&(~memberP(X4,X5)|?[X6]:(((ssItem(X6)&X5!=X6)&memberP(X4,X6))&leq(X6,X5)))))))|![X7]:(ssItem(X7)=>(~memberP(X1,X7)|memberP(X2,X7)))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
 % 0.02/0.25  fof(c_0_2, negated_conjecture, (![X12]:![X14]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((esk49_0=esk51_0&esk48_0=esk50_0)&(((((ssItem(esk52_1(X12))|memberP(esk50_0,X12))|~memberP(esk51_0,X12))|~ssItem(X12))&(((((memberP(esk51_0,esk52_1(X12))|memberP(esk50_0,X12))|~memberP(esk51_0,X12))|~ssItem(X12))&(((leq(esk52_1(X12),X12)|memberP(esk50_0,X12))|~memberP(esk51_0,X12))|~ssItem(X12)))&(((X12!=esk52_1(X12)|memberP(esk50_0,X12))|~memberP(esk51_0,X12))|~ssItem(X12))))&(((memberP(esk51_0,X12)|~memberP(esk50_0,X12))|~ssItem(X12))&(((((~ssItem(X14)|X12=X14)|~memberP(esk51_0,X14))|~leq(X14,X12))|~memberP(esk50_0,X12))|~ssItem(X12)))))&(ssItem(esk53_0)&(memberP(esk48_0,esk53_0)&~memberP(esk49_0,esk53_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
 % 0.02/0.25  cnf(c_0_3, negated_conjecture, (memberP(esk48_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.02/0.25  cnf(c_0_4, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.02/0.25  cnf(c_0_5, negated_conjecture, (~memberP(esk49_0,esk53_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.02/0.25  cnf(c_0_6, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.02/0.25  cnf(c_0_7, negated_conjecture, (memberP(esk51_0,X1)|~ssItem(X1)|~memberP(esk50_0,X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.02/0.25  cnf(c_0_8, negated_conjecture, (memberP(esk50_0,esk53_0)), inference(rw,[status(thm)],[c_0_3, c_0_4])).
 % 0.02/0.25  cnf(c_0_9, negated_conjecture, (ssItem(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
 % 0.02/0.25  cnf(c_0_10, negated_conjecture, (~memberP(esk51_0,esk53_0)), inference(rw,[status(thm)],[c_0_5, c_0_6])).
 % 0.02/0.25  cnf(c_0_11, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9])]), c_0_10]), ['proof']).
 % 0.02/0.25  # SZS output end CNFRefutation
 % 0.02/0.25  # Training examples: 0 positive, 0 negative
 % 0.02/0.25
 % 0.02/0.25  # -------------------------------------------------
 % 0.02/0.25  # User time                : 0.025 s
 % 0.02/0.25  # System time              : 0.003 s
 % 0.02/0.25  # Total time               : 0.027 s
 % 0.02/0.25  # Maximum resident set size: 1876 pages
 %------------------------------------------------------------------------------

   --------------------------------------------------------------------------
