<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC005-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC005-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC005-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC005-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC005-1.p">SWC005-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n117.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:15 EDT 2016

% Result     : Unsatisfiable 0.03s
% Output     : CNFRefutation 0.03s
% Verified   : 
% Statistics : Number of clauses        :   38 (  73 expanded)
%              Number of leaves         :   17 (  43 expanded)
%              Depth                    :   15
%              Number of atoms          :   99 ( 144 expanded)
%              Number of equality atoms :   38 (  51 expanded)
%              Maximal clause size      :    7 (   3 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( neq(sk2,nil)
    | neq(sk2,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | app(app(X1,X2),X3) != sk2
    | app(X1,X3) != sk1
    | ~ neq(X2,nil)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( sk3 = X1
    | ~ ssList(X1)
    | tl(sk4) != X1
    | ~ neq(nil,sk4)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_9"></A>cnf(c_0_9,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause75)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause104)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause76)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause99)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( X1 != X2
    | ~ neq(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause115)).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( neq(sk2,nil) ),
    inference(cn,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( neq(sk4,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( app(app(X1,X2),X3) != sk4
    | app(X1,X3) != sk1
    | ~ neq(X2,nil)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_1]),c_0_18])])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( tl(sk4) = sk1
    | ~ neq(nil,sk4)
    | ~ ssList(tl(sk4)) ),
    inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>,c_0_4]),c_0_18])])])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( app(app(nil,X1),sk1) != sk4
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_6]),c_0_7])])]),c_0_8])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( tl(sk4) = sk1
    | sk4 = nil
    | ~ ssList(tl(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_9]),c_0_7]),c_0_21])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( app(X1,sk1) != sk4
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( tl(sk4) = sk1
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_10]),c_0_21])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( cons(X1,sk1) != sk4
    | ~ neq(cons(X1,nil),nil)
    | ~ ssItem(X1)
    | ~ ssList(cons(X1,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_11]),c_0_8])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( cons(hd(sk4),sk1) = sk4
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,c_0_25]),c_0_21])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( sk4 = nil
    | ~ neq(cons(hd(sk4),nil),nil)
    | ~ ssItem(hd(sk4))
    | ~ ssList(cons(hd(sk4),nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( cons(hd(sk4),nil) = nil
    | sk4 = nil
    | ~ ssItem(hd(sk4))
    | ~ ssList(cons(hd(sk4),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_9]),c_0_7])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( cons(hd(sk4),nil) = nil
    | sk4 = nil
    | ~ ssList(cons(hd(sk4),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_13]),c_0_21])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( cons(hd(sk4),nil) = nil
    | sk4 = nil
    | ~ ssItem(hd(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_14]),c_0_7])])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( cons(hd(sk4),nil) = nil
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_13]),c_0_21])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( sk4 = nil
    | ~ ssItem(hd(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_32]),c_0_7])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,c_0_13]),c_0_21])])).

<A NAME="c_0_35"></A>cnf(c_0_35,plain,
    ( ~ neq(X1,X1)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( neq(nil,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_36]),c_0_7])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC005-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n117.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:34:23 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.03/0.26  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.03/0.26  # and selection function SelectCQArNTEqFirst.
% 0.03/0.26  #
% 0.03/0.26  # Preprocessing time       : 0.009 s
% 0.03/0.26  # Presaturation interreduction done
% 0.03/0.26  
% 0.03/0.26  # Proof found!
% 0.03/0.26  # SZS status Unsatisfiable
% 0.03/0.26  # SZS output start CNFRefutation
% 0.03/0.26  cnf(c_0_0, negated_conjecture, (neq(sk2,nil)|neq(sk2,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 0.03/0.26  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 0.03/0.26  cnf(c_0_2, negated_conjecture, (~ssList(X1)|~ssList(X2)|~ssList(X3)|app(app(X1,X2),X3)!=sk2|app(X1,X3)!=sk1|~neq(X2,nil)|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_11)).
% 0.03/0.26  cnf(c_0_3, negated_conjecture, (sk3=X1|~ssList(X1)|tl(sk4)!=X1|~neq(nil,sk4)|~neq(sk4,nil)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_12)).
% 0.03/0.26  cnf(c_0_4, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.03/0.26  cnf(c_0_5, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 0.03/0.26  cnf(c_0_6, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause74)).
% 0.03/0.26  cnf(c_0_7, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.03/0.26  cnf(c_0_8, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.03/0.26  cnf(c_0_9, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause100)).
% 0.03/0.26  cnf(c_0_10, axiom, (ssList(tl(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause75)).
% 0.03/0.26  cnf(c_0_11, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
% 0.03/0.26  cnf(c_0_12, axiom, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause104)).
% 0.03/0.26  cnf(c_0_13, axiom, (ssItem(hd(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause76)).
% 0.03/0.26  cnf(c_0_14, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.03/0.26  cnf(c_0_15, axiom, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause99)).
% 0.03/0.26  cnf(c_0_16, axiom, (X1!=X2|~neq(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause115)).
% 0.03/0.26  cnf(c_0_17, negated_conjecture, (neq(sk2,nil)), inference(cn,[status(thm)],[c_0_0])).
% 0.03/0.26  cnf(c_0_18, negated_conjecture, (neq(sk4,nil)), inference(rw,[status(thm)],[c_0_17, c_0_1])).
% 0.03/0.26  cnf(c_0_19, negated_conjecture, (app(app(X1,X2),X3)!=sk4|app(X1,X3)!=sk1|~neq(X2,nil)|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_2, c_0_1]), c_0_18])])).
% 0.03/0.26  cnf(c_0_20, negated_conjecture, (tl(sk4)=sk1|~neq(nil,sk4)|~ssList(tl(sk4))), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_4]), c_0_18])])])).
% 0.03/0.26  cnf(c_0_21, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_5, c_0_1])).
% 0.03/0.26  cnf(c_0_22, negated_conjecture, (app(app(nil,X1),sk1)!=sk4|~neq(X1,nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_6]), c_0_7])])]), c_0_8])])).
% 0.03/0.26  cnf(c_0_23, negated_conjecture, (tl(sk4)=sk1|sk4=nil|~ssList(tl(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_9]), c_0_7]), c_0_21])])).
% 0.03/0.26  cnf(c_0_24, negated_conjecture, (app(X1,sk1)!=sk4|~neq(X1,nil)|~ssList(X1)), inference(spm,[status(thm)],[c_0_22, c_0_6])).
% 0.03/0.26  cnf(c_0_25, negated_conjecture, (tl(sk4)=sk1|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_10]), c_0_21])])).
% 0.03/0.26  cnf(c_0_26, negated_conjecture, (cons(X1,sk1)!=sk4|~neq(cons(X1,nil),nil)|~ssItem(X1)|~ssList(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_11]), c_0_8])])).
% 0.03/0.26  cnf(c_0_27, negated_conjecture, (cons(hd(sk4),sk1)=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_25]), c_0_21])])).
% 0.03/0.26  cnf(c_0_28, negated_conjecture, (sk4=nil|~neq(cons(hd(sk4),nil),nil)|~ssItem(hd(sk4))|~ssList(cons(hd(sk4),nil))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
% 0.03/0.26  cnf(c_0_29, negated_conjecture, (cons(hd(sk4),nil)=nil|sk4=nil|~ssItem(hd(sk4))|~ssList(cons(hd(sk4),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_9]), c_0_7])])).
% 0.03/0.26  cnf(c_0_30, negated_conjecture, (cons(hd(sk4),nil)=nil|sk4=nil|~ssList(cons(hd(sk4),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_13]), c_0_21])])).
% 0.03/0.26  cnf(c_0_31, negated_conjecture, (cons(hd(sk4),nil)=nil|sk4=nil|~ssItem(hd(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_14]), c_0_7])])).
% 0.03/0.26  cnf(c_0_32, negated_conjecture, (cons(hd(sk4),nil)=nil|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_13]), c_0_21])])).
% 0.03/0.26  cnf(c_0_33, negated_conjecture, (sk4=nil|~ssItem(hd(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_32]), c_0_7])])).
% 0.03/0.26  cnf(c_0_34, negated_conjecture, (sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_13]), c_0_21])])).
% 0.03/0.26  cnf(c_0_35, plain, (~neq(X1,X1)|~ssList(X1)), inference(er,[status(thm)],[c_0_16])).
% 0.03/0.26  cnf(c_0_36, negated_conjecture, (neq(nil,nil)), inference(rw,[status(thm)],[c_0_18, c_0_34])).
% 0.03/0.26  cnf(c_0_37, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_7])]), ['proof']).
% 0.03/0.26  # SZS output end CNFRefutation
% 0.03/0.26  # Training examples: 0 positive, 0 negative
% 0.03/0.26  
% 0.03/0.26  # -------------------------------------------------
% 0.03/0.26  # User time                : 0.028 s
% 0.03/0.26  # System time              : 0.002 s
% 0.03/0.26  # Total time               : 0.030 s
% 0.03/0.26  # Maximum resident set size: 1864 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
