<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC162+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC162+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC162+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC162%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC162+1.p">SWC162+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n177.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:53 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :   95 ( 910 expanded)
%              Number of clauses        :   65 ( 330 expanded)
%              Number of leaves         :   15 ( 229 expanded)
%              Depth                    :   20
%              Number of atoms          :  400 (10377 expanded)
%              Number of equality atoms :   90 (1514 expanded)
%              Maximal formula depth    :   31 (   5 average)
%              Maximal clause size      :   34 (   2 average)
%              Maximal term depth       :    6 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ~ segmentP(X4,X3)
                    | ! [X5] :
                        ( ssItem(X5)
                       => ! [X6] :
                            ( ssItem(X6)
                           => ! [X7] :
                                ( ssList(X7)
                               => ! [X8] :
                                    ( ssList(X8)
                                   => ! [X9] :
                                        ( ssList(X9)
                                       => ( app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9) != X1
                                          | ~ leq(X6,X5)
                                          | ( ! [X10] :
                                                ( ssItem(X10)
                                               => ( ~ memberP(X8,X10)
                                                  | ( leq(X5,X10)
                                                    & leq(X10,X6) ) ) )
                                            & leq(X5,X6) ) ) ) ) ) ) )
                    | ( ~ singletonP(X3)
                      & neq(X4,nil) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( totalorderedP(X1)
      &lt;=> ! [X2] :
            ( ssItem(X2)
           => ! [X3] :
                ( ssItem(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ! [X5] :
                        ( ssList(X5)
                       => ! [X6] :
                            ( ssList(X6)
                           => ( app(app(X4,cons(X2,X5)),cons(X3,X6)) = X1
                             => leq(X2,X3) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax11)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => totalorderedP(cons(X1,nil)) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax65)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( singletonP(X1)
      &lt;=> ? [X2] :
            ( ssItem(X2)
            & cons(X2,nil) = X1 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax4)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( cyclefreeP(X1)
      &lt;=> ! [X2] :
            ( ssItem(X2)
           => ! [X3] :
                ( ssItem(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ! [X5] :
                        ( ssList(X5)
                       => ! [X6] :
                            ( ssList(X6)
                           => ( app(app(X4,cons(X2,X5)),cons(X3,X6)) = X1
                             => ~ ( leq(X2,X3)
                                  & leq(X3,X2) ) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax8)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => cyclefreeP(cons(X1,nil)) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax59)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( segmentP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax58)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) != X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax18)).

<A NAME="c_0_15"></A>fof(c_0_15,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ~ segmentP(X4,X3)
                      | ! [X5] :
                          ( ssItem(X5)
                         => ! [X6] :
                              ( ssItem(X6)
                             => ! [X7] :
                                  ( ssList(X7)
                                 => ! [X8] :
                                      ( ssList(X8)
                                     => ! [X9] :
                                          ( ssList(X9)
                                         => ( app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9) != X1
                                            | ~ leq(X6,X5)
                                            | ( ! [X10] :
                                                  ( ssItem(X10)
                                                 => ( ~ memberP(X8,X10)
                                                    | ( leq(X5,X10)
                                                      & leq(X10,X6) ) ) )
                                              & leq(X5,X6) ) ) ) ) ) ) )
                      | ( ~ singletonP(X3)
                        & neq(X4,nil) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_16"></A>fof(c_0_16,negated_conjecture,
    ( ssList(esk48_0)
    & ssList(esk49_0)
    & ssList(esk50_0)
    & ssList(esk51_0)
    & esk49_0 = esk51_0
    & esk48_0 = esk50_0
    & segmentP(esk51_0,esk50_0)
    & ssItem(esk52_0)
    & ssItem(esk53_0)
    & ssList(esk54_0)
    & ssList(esk55_0)
    & ssList(esk56_0)
    & app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk48_0
    & leq(esk53_0,esk52_0)
    & ( ssItem(esk57_0)
      | ~ leq(esk52_0,esk53_0) )
    & ( memberP(esk55_0,esk57_0)
      | ~ leq(esk52_0,esk53_0) )
    & ( ~ leq(esk52_0,esk57_0)
      | ~ leq(esk57_0,esk53_0)
      | ~ leq(esk52_0,esk53_0) )
    & ( singletonP(esk50_0)
      | ~ neq(esk51_0,nil) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_19"></A>fof(c_0_19,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( ssList(esk55_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( ssList(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_24"></A>fof(c_0_24,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0) = esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_21]),c_0_22]),c_0_23])])).

<A NAME="c_0_26"></A>cnf(c_0_26,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( ssItem(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_29"></A>fof(c_0_29,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_26]),c_0_27]),c_0_28])])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( app(app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,nil)),esk56_0) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31]),c_0_22]),c_0_28])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( ssList(esk56_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0
    | ~ ssList(app(esk54_0,cons(esk52_0,esk55_0)))
    | ~ ssList(cons(esk53_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_32]),c_0_33])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( ssItem(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_36"></A>fof(c_0_36,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0
    | ~ ssList(app(esk54_0,cons(esk52_0,esk55_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_26]),c_0_27]),c_0_35])])).

<A NAME="c_0_38"></A>cnf(c_0_38,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_39"></A>fof(c_0_39,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).

<A NAME="c_0_40"></A>fof(c_0_40,plain,(
    ! [X7,X8,X9,X10,X11,X12] :
      ( ( ~ totalorderedP(X7)
        | ~ ssItem(X8)
        | ~ ssItem(X9)
        | ~ ssList(X10)
        | ~ ssList(X11)
        | ~ ssList(X12)
        | app(app(X10,cons(X8,X11)),cons(X9,X12)) != X7
        | leq(X8,X9)
        | ~ ssList(X7) )
      & ( ssItem(esk25_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssItem(esk26_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk27_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk28_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk29_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( app(app(esk27_1(X7),cons(esk25_1(X7),esk28_1(X7))),cons(esk26_1(X7),esk29_1(X7))) = X7
        | totalorderedP(X7)
        | ~ ssList(X7) )
      & ( ~ leq(esk25_1(X7),esk26_1(X7))
        | totalorderedP(X7)
        | ~ ssList(X7) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0
    | ~ ssList(cons(esk52_0,esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_38]),c_0_23])])).

<A NAME="c_0_42"></A>fof(c_0_42,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | totalorderedP(cons(X2,nil)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).

<A NAME="c_0_43"></A>fof(c_0_43,plain,(
    ! [X3,X5] :
      ( ( ssItem(esk5_1(X3))
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( cons(esk5_1(X3),nil) = X3
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( ~ ssItem(X5)
        | cons(X5,nil) != X3
        | singletonP(X3)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).

<A NAME="c_0_44"></A>fof(c_0_44,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_47"></A>cnf(c_0_47,plain,
    ( nil = X2
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_48"></A>fof(c_0_48,plain,(
    ! [X7,X8,X9,X10,X11,X12] :
      ( ( ~ cyclefreeP(X7)
        | ~ ssItem(X8)
        | ~ ssItem(X9)
        | ~ ssList(X10)
        | ~ ssList(X11)
        | ~ ssList(X12)
        | app(app(X10,cons(X8,X11)),cons(X9,X12)) != X7
        | ~ leq(X8,X9)
        | ~ leq(X9,X8)
        | ~ ssList(X7) )
      & ( ssItem(esk10_1(X7))
        | cyclefreeP(X7)
        | ~ ssList(X7) )
      & ( ssItem(esk11_1(X7))
        | cyclefreeP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk12_1(X7))
        | cyclefreeP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk13_1(X7))
        | cyclefreeP(X7)
        | ~ ssList(X7) )
      & ( ssList(esk14_1(X7))
        | cyclefreeP(X7)
        | ~ ssList(X7) )
      & ( app(app(esk12_1(X7),cons(esk10_1(X7),esk13_1(X7))),cons(esk11_1(X7),esk14_1(X7))) = X7
        | cyclefreeP(X7)
        | ~ ssList(X7) )
      & ( leq(esk10_1(X7),esk11_1(X7))
        | cyclefreeP(X7)
        | ~ ssList(X7) )
      & ( leq(esk11_1(X7),esk10_1(X7))
        | cyclefreeP(X7)
        | ~ ssList(X7) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).

<A NAME="c_0_49"></A>cnf(c_0_49,plain,
    ( leq(X2,X3)
    | ~ ssList(X1)
    | app(app(X4,cons(X2,X5)),cons(X3,X6)) != X1
    | ~ ssList(X6)
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssItem(X3)
    | ~ ssItem(X2)
    | ~ totalorderedP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>,c_0_26]),c_0_22]),c_0_28])])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,plain,
    ( totalorderedP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( cons(esk5_1(X1),nil) = X1
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>])).

<A NAME="c_0_54"></A>cnf(c_0_54,plain,
    ( ssItem(esk5_1(X1))
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( singletonP(esk50_0)
    | ~ neq(esk51_0,nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( ssList(esk51_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_58"></A>fof(c_0_58,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | cyclefreeP(cons(X2,nil)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( esk56_0 = nil
    | esk50_0 != nil
    | ~ ssList(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,c_0_20]),c_0_33])])).

<A NAME="c_0_60"></A>cnf(c_0_60,plain,
    ( ~ ssList(X1)
    | ~ leq(X2,X3)
    | ~ leq(X3,X2)
    | app(app(X4,cons(X3,X5)),cons(X2,X6)) != X1
    | ~ ssList(X6)
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssItem(X2)
    | ~ ssItem(X3)
    | ~ cyclefreeP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_61"></A>cnf(c_0_61,plain,
    ( leq(X1,X2)
    | ~ totalorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,esk56_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_31]),c_0_33]),c_0_35])])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,<A HREF="#c_0_18">c_0_18</A>])).

<A NAME="c_0_64"></A>cnf(c_0_64,plain,
    ( totalorderedP(X1)
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_52">c_0_52</A>,c_0_53]),<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,negated_conjecture,
    ( esk51_0 = nil
    | singletonP(esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_56]),c_0_27]),c_0_57])])).

<A NAME="c_0_66"></A>cnf(c_0_66,plain,
    ( cyclefreeP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( esk56_0 = nil
    | esk50_0 != nil
    | ~ ssList(app(app(esk54_0,cons(esk52_0,nil)),esk55_0))
    | ~ ssList(cons(esk53_0,nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>,<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_68"></A>cnf(c_0_68,plain,
    ( ~ leq(X1,X2)
    | ~ leq(X2,X1)
    | ~ cyclefreeP(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssItem(X1)
    | ~ ssItem(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>])).

<A NAME="c_0_69"></A>cnf(c_0_69,negated_conjecture,
    ( leq(esk53_0,esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( leq(esk52_0,esk53_0)
    | ~ totalorderedP(esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_62]),c_0_63]),c_0_33]),c_0_22]),c_0_23]),c_0_35]),c_0_28])])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( esk51_0 = nil
    | totalorderedP(esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,c_0_65]),c_0_63])])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( cyclefreeP(X1)
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>,c_0_53]),<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,negated_conjecture,
    ( esk56_0 = nil
    | esk50_0 != nil
    | ~ ssList(app(app(esk54_0,cons(esk52_0,nil)),esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,c_0_26]),c_0_27]),c_0_35])])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( ~ leq(esk52_0,esk53_0)
    | ~ cyclefreeP(esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,c_0_62]),c_0_69]),c_0_63]),c_0_33]),c_0_22]),c_0_23]),c_0_28]),c_0_35])])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( esk51_0 = nil
    | leq(esk52_0,esk53_0) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>,<A HREF="#c_0_71">c_0_71</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,negated_conjecture,
    ( esk51_0 = nil
    | cyclefreeP(esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_65]),c_0_63])])).

<A NAME="c_0_77"></A>cnf(c_0_77,negated_conjecture,
    ( cons(esk53_0,esk56_0) = nil
    | esk50_0 != nil
    | ~ ssList(app(esk54_0,cons(esk52_0,esk55_0)))
    | ~ ssList(cons(esk53_0,esk56_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>,<A HREF="#c_0_62">c_0_62</A>])).

<A NAME="c_0_78"></A>cnf(c_0_78,negated_conjecture,
    ( esk56_0 = nil
    | esk50_0 != nil
    | ~ ssList(app(esk54_0,cons(esk52_0,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>,c_0_38]),c_0_22])])).

<A NAME="c_0_79"></A>fof(c_0_79,plain,(
    ! [X2] :
      ( ( ~ segmentP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | segmentP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( segmentP(esk51_0,esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_81"></A>cnf(c_0_81,negated_conjecture,
    ( esk51_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,c_0_75]),<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( cons(esk53_0,esk56_0) = nil
    | esk50_0 != nil
    | ~ ssList(app(esk54_0,cons(esk52_0,esk55_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_26]),c_0_33]),c_0_35])])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( esk56_0 = nil
    | esk50_0 != nil
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_38]),c_0_23])])).

<A NAME="c_0_84"></A>cnf(c_0_84,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ segmentP(nil,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_79">c_0_79</A>])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( segmentP(nil,esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,<A HREF="#c_0_81">c_0_81</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( cons(esk53_0,esk56_0) = nil
    | esk50_0 != nil
    | ~ ssList(cons(esk52_0,esk55_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,c_0_38]),c_0_23])])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( esk56_0 = nil
    | esk50_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_26]),c_0_27]),c_0_28])])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,c_0_85]),c_0_63])])).

<A NAME="c_0_89"></A>fof(c_0_89,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) != X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( cons(esk53_0,esk56_0) = nil
    | esk50_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>,c_0_26]),c_0_22]),c_0_28])])).

<A NAME="c_0_91"></A>cnf(c_0_91,negated_conjecture,
    ( esk56_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,c_0_88])])).

<A NAME="c_0_92"></A>cnf(c_0_92,plain,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>])).

<A NAME="c_0_93"></A>cnf(c_0_93,negated_conjecture,
    ( cons(esk53_0,nil) = nil ),
    inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,c_0_88])]),<A HREF="#c_0_91">c_0_91</A>])).

<A NAME="c_0_94"></A>cnf(c_0_94,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_93]),c_0_27]),c_0_35])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC162+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n177.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:36:23 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.30  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.30  # and selection function SelectCQPrecWNTNp.
% 0.06/0.30  #
% 0.06/0.30  # Preprocessing time       : 0.015 s
% 0.06/0.30  # Presaturation interreduction done
% 0.06/0.30  
% 0.06/0.30  # Proof found!
% 0.06/0.30  # SZS status Theorem
% 0.06/0.30  # SZS output start CNFRefutation
% 0.06/0.30  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|~(segmentP(X4,X3)))|![X5]:(ssItem(X5)=>![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>![X9]:(ssList(X9)=>((app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9)!=X1|~(leq(X6,X5)))|(![X10]:(ssItem(X10)=>(~(memberP(X8,X10))|(leq(X5,X10)&leq(X10,X6))))&leq(X5,X6)))))))))|(~(singletonP(X3))&neq(X4,nil)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.06/0.30  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.06/0.30  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.06/0.30  fof(c_0_3, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.30  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.06/0.30  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.06/0.30  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.06/0.30  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(totalorderedP(X1)&lt;=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(app(X4,cons(X2,X5)),cons(X3,X6))=X1=>leq(X2,X3)))))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax11)).
% 0.06/0.30  fof(c_0_8, axiom, (![X1]:(ssItem(X1)=>totalorderedP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax65)).
% 0.06/0.30  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)&lt;=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax4)).
% 0.06/0.30  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.06/0.30  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>(cyclefreeP(X1)&lt;=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>![X6]:(ssList(X6)=>(app(app(X4,cons(X2,X5)),cons(X3,X6))=X1=>~((leq(X2,X3)&leq(X3,X2)))))))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax8)).
% 0.06/0.30  fof(c_0_12, axiom, (![X1]:(ssItem(X1)=>cyclefreeP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax59)).
% 0.06/0.30  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>(segmentP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax58)).
% 0.06/0.30  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax18)).
% 0.06/0.30  fof(c_0_15, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|~segmentP(X4,X3))|![X5]:(ssItem(X5)=>![X6]:(ssItem(X6)=>![X7]:(ssList(X7)=>![X8]:(ssList(X8)=>![X9]:(ssList(X9)=>((app(app(app(app(X7,cons(X5,nil)),X8),cons(X6,nil)),X9)!=X1|~leq(X6,X5))|(![X10]:(ssItem(X10)=>(~memberP(X8,X10)|(leq(X5,X10)&leq(X10,X6))))&leq(X5,X6)))))))))|(~singletonP(X3)&neq(X4,nil))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.30  fof(c_0_16, negated_conjecture, ((ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&segmentP(esk51_0,esk50_0))&(ssItem(esk52_0)&(ssItem(esk53_0)&(ssList(esk54_0)&(ssList(esk55_0)&(ssList(esk56_0)&((app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk48_0&leq(esk53_0,esk52_0))&((ssItem(esk57_0)|~leq(esk52_0,esk53_0))&((memberP(esk55_0,esk57_0)|~leq(esk52_0,esk53_0))&((~leq(esk52_0,esk57_0)|~leq(esk57_0,esk53_0))|~leq(esk52_0,esk53_0)))))))))))&(singletonP(esk50_0)|~neq(esk51_0,nil)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])).
% 0.06/0.30  cnf(c_0_17, negated_conjecture, (app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  cnf(c_0_18, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  fof(c_0_19, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.06/0.30  cnf(c_0_20, negated_conjecture, (app(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)),esk56_0)=esk50_0), inference(rw,[status(thm)],[c_0_17, c_0_18])).
% 0.06/0.30  cnf(c_0_21, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
% 0.06/0.30  cnf(c_0_22, negated_conjecture, (ssList(esk55_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  cnf(c_0_23, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  fof(c_0_24, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.06/0.30  cnf(c_0_25, negated_conjecture, (app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0)=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])])).
% 0.06/0.30  cnf(c_0_26, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
% 0.06/0.30  cnf(c_0_27, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_3])).
% 0.06/0.30  cnf(c_0_28, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  fof(c_0_29, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 0.06/0.30  cnf(c_0_30, negated_conjecture, (app(app(app(esk54_0,app(cons(esk52_0,nil),esk55_0)),cons(esk53_0,nil)),esk56_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28])])).
% 0.06/0.30  cnf(c_0_31, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
% 0.06/0.30  cnf(c_0_32, negated_conjecture, (app(app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,nil)),esk56_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_22]), c_0_28])])).
% 0.06/0.30  cnf(c_0_33, negated_conjecture, (ssList(esk56_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  cnf(c_0_34, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0|~ssList(app(esk54_0,cons(esk52_0,esk55_0)))|~ssList(cons(esk53_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_32]), c_0_33])])).
% 0.06/0.30  cnf(c_0_35, negated_conjecture, (ssItem(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  fof(c_0_36, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 0.06/0.30  cnf(c_0_37, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0|~ssList(app(esk54_0,cons(esk52_0,esk55_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_26]), c_0_27]), c_0_35])])).
% 0.06/0.30  cnf(c_0_38, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
% 0.06/0.30  fof(c_0_39, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])).
% 0.06/0.30  fof(c_0_40, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((~totalorderedP(X7)|(~ssItem(X8)|(~ssItem(X9)|(~ssList(X10)|(~ssList(X11)|(~ssList(X12)|(app(app(X10,cons(X8,X11)),cons(X9,X12))!=X7|leq(X8,X9))))))))|~ssList(X7))&(((ssItem(esk25_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssItem(esk26_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk27_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk28_1(X7))|totalorderedP(X7))|~ssList(X7))&(((ssList(esk29_1(X7))|totalorderedP(X7))|~ssList(X7))&(((app(app(esk27_1(X7),cons(esk25_1(X7),esk28_1(X7))),cons(esk26_1(X7),esk29_1(X7)))=X7|totalorderedP(X7))|~ssList(X7))&((~leq(esk25_1(X7),esk26_1(X7))|totalorderedP(X7))|~ssList(X7)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
% 0.06/0.30  cnf(c_0_41, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0|~ssList(cons(esk52_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_23])])).
% 0.06/0.30  fof(c_0_42, plain, (![X2]:(~ssItem(X2)|totalorderedP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
% 0.06/0.30  fof(c_0_43, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])])).
% 0.06/0.30  fof(c_0_44, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).
% 0.06/0.30  cnf(c_0_45, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  cnf(c_0_46, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  cnf(c_0_47, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.30  fof(c_0_48, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((~cyclefreeP(X7)|(~ssItem(X8)|(~ssItem(X9)|(~ssList(X10)|(~ssList(X11)|(~ssList(X12)|(app(app(X10,cons(X8,X11)),cons(X9,X12))!=X7|(~leq(X8,X9)|~leq(X9,X8)))))))))|~ssList(X7))&(((ssItem(esk10_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssItem(esk11_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssList(esk12_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssList(esk13_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((ssList(esk14_1(X7))|cyclefreeP(X7))|~ssList(X7))&(((app(app(esk12_1(X7),cons(esk10_1(X7),esk13_1(X7))),cons(esk11_1(X7),esk14_1(X7)))=X7|cyclefreeP(X7))|~ssList(X7))&(((leq(esk10_1(X7),esk11_1(X7))|cyclefreeP(X7))|~ssList(X7))&((leq(esk11_1(X7),esk10_1(X7))|cyclefreeP(X7))|~ssList(X7))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
% 0.06/0.30  cnf(c_0_49, plain, (leq(X2,X3)|~ssList(X1)|app(app(X4,cons(X2,X5)),cons(X3,X6))!=X1|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssItem(X3)|~ssItem(X2)|~totalorderedP(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
% 0.06/0.30  cnf(c_0_50, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),app(cons(esk53_0,nil),esk56_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_26]), c_0_22]), c_0_28])])).
% 0.06/0.30  cnf(c_0_51, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  cnf(c_0_52, plain, (totalorderedP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
% 0.06/0.30  cnf(c_0_53, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
% 0.06/0.30  cnf(c_0_54, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
% 0.06/0.30  cnf(c_0_55, negated_conjecture, (singletonP(esk50_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  cnf(c_0_56, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
% 0.06/0.30  cnf(c_0_57, negated_conjecture, (ssList(esk51_0)), inference(rw,[status(thm)],[c_0_45, c_0_46])).
% 0.06/0.30  fof(c_0_58, plain, (![X2]:(~ssItem(X2)|cyclefreeP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).
% 0.06/0.30  cnf(c_0_59, negated_conjecture, (esk56_0=nil|esk50_0!=nil|~ssList(app(app(app(esk54_0,cons(esk52_0,nil)),esk55_0),cons(esk53_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_20]), c_0_33])])).
% 0.06/0.30  cnf(c_0_60, plain, (~ssList(X1)|~leq(X2,X3)|~leq(X3,X2)|app(app(X4,cons(X3,X5)),cons(X2,X6))!=X1|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssItem(X2)|~ssItem(X3)|~cyclefreeP(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
% 0.06/0.30  cnf(c_0_61, plain, (leq(X1,X2)|~totalorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_49])).
% 0.06/0.30  cnf(c_0_62, negated_conjecture, (app(app(esk54_0,cons(esk52_0,esk55_0)),cons(esk53_0,esk56_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_31]), c_0_33]), c_0_35])])).
% 0.06/0.30  cnf(c_0_63, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_51, c_0_18])).
% 0.06/0.30  cnf(c_0_64, plain, (totalorderedP(X1)|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
% 0.06/0.30  cnf(c_0_65, negated_conjecture, (esk51_0=nil|singletonP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_27]), c_0_57])])).
% 0.06/0.30  cnf(c_0_66, plain, (cyclefreeP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
% 0.06/0.30  cnf(c_0_67, negated_conjecture, (esk56_0=nil|esk50_0!=nil|~ssList(app(app(esk54_0,cons(esk52_0,nil)),esk55_0))|~ssList(cons(esk53_0,nil))), inference(spm,[status(thm)],[c_0_59, c_0_38])).
% 0.06/0.30  cnf(c_0_68, plain, (~leq(X1,X2)|~leq(X2,X1)|~cyclefreeP(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssItem(X1)|~ssItem(X2)), inference(er,[status(thm)],[c_0_60])).
% 0.06/0.30  cnf(c_0_69, negated_conjecture, (leq(esk53_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  cnf(c_0_70, negated_conjecture, (leq(esk52_0,esk53_0)|~totalorderedP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63]), c_0_33]), c_0_22]), c_0_23]), c_0_35]), c_0_28])])).
% 0.06/0.30  cnf(c_0_71, negated_conjecture, (esk51_0=nil|totalorderedP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_63])])).
% 0.06/0.30  cnf(c_0_72, plain, (cyclefreeP(X1)|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_53]), c_0_54])).
% 0.06/0.30  cnf(c_0_73, negated_conjecture, (esk56_0=nil|esk50_0!=nil|~ssList(app(app(esk54_0,cons(esk52_0,nil)),esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_26]), c_0_27]), c_0_35])])).
% 0.06/0.30  cnf(c_0_74, negated_conjecture, (~leq(esk52_0,esk53_0)|~cyclefreeP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_62]), c_0_69]), c_0_63]), c_0_33]), c_0_22]), c_0_23]), c_0_28]), c_0_35])])).
% 0.06/0.30  cnf(c_0_75, negated_conjecture, (esk51_0=nil|leq(esk52_0,esk53_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
% 0.06/0.30  cnf(c_0_76, negated_conjecture, (esk51_0=nil|cyclefreeP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_65]), c_0_63])])).
% 0.06/0.30  cnf(c_0_77, negated_conjecture, (cons(esk53_0,esk56_0)=nil|esk50_0!=nil|~ssList(app(esk54_0,cons(esk52_0,esk55_0)))|~ssList(cons(esk53_0,esk56_0))), inference(spm,[status(thm)],[c_0_47, c_0_62])).
% 0.06/0.30  cnf(c_0_78, negated_conjecture, (esk56_0=nil|esk50_0!=nil|~ssList(app(esk54_0,cons(esk52_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_38]), c_0_22])])).
% 0.06/0.30  fof(c_0_79, plain, (![X2]:(((~segmentP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|segmentP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
% 0.06/0.30  cnf(c_0_80, negated_conjecture, (segmentP(esk51_0,esk50_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
% 0.06/0.30  cnf(c_0_81, negated_conjecture, (esk51_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])).
% 0.06/0.30  cnf(c_0_82, negated_conjecture, (cons(esk53_0,esk56_0)=nil|esk50_0!=nil|~ssList(app(esk54_0,cons(esk52_0,esk55_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_26]), c_0_33]), c_0_35])])).
% 0.06/0.30  cnf(c_0_83, negated_conjecture, (esk56_0=nil|esk50_0!=nil|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_38]), c_0_23])])).
% 0.06/0.30  cnf(c_0_84, plain, (nil=X1|~ssList(X1)|~segmentP(nil,X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
% 0.06/0.30  cnf(c_0_85, negated_conjecture, (segmentP(nil,esk50_0)), inference(rw,[status(thm)],[c_0_80, c_0_81])).
% 0.06/0.30  cnf(c_0_86, negated_conjecture, (cons(esk53_0,esk56_0)=nil|esk50_0!=nil|~ssList(cons(esk52_0,esk55_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_38]), c_0_23])])).
% 0.06/0.30  cnf(c_0_87, negated_conjecture, (esk56_0=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_26]), c_0_27]), c_0_28])])).
% 0.06/0.30  cnf(c_0_88, negated_conjecture, (esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_63])])).
% 0.06/0.30  fof(c_0_89, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
% 0.06/0.30  cnf(c_0_90, negated_conjecture, (cons(esk53_0,esk56_0)=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_26]), c_0_22]), c_0_28])])).
% 0.06/0.30  cnf(c_0_91, negated_conjecture, (esk56_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_88])])).
% 0.06/0.30  cnf(c_0_92, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
% 0.06/0.30  cnf(c_0_93, negated_conjecture, (cons(esk53_0,nil)=nil), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_88])]), c_0_91])).
% 0.06/0.30  cnf(c_0_94, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_27]), c_0_35])]), ['proof']).
% 0.06/0.30  # SZS output end CNFRefutation
% 0.06/0.30  # Training examples: 0 positive, 0 negative
% 0.06/0.30  
% 0.06/0.30  # -------------------------------------------------
% 0.06/0.30  # User time                : 0.064 s
% 0.06/0.30  # System time              : 0.007 s
% 0.06/0.30  # Total time               : 0.071 s
% 0.06/0.30  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
