<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC152-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC152-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC152-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC152-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC152-1.p">SWC152-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n159.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:51 EDT 2016

% Result     : Unsatisfiable 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of clauses        :   52 ( 308 expanded)
%              Number of leaves         :   24 ( 219 expanded)
%              Depth                    :   18
%              Number of atoms          :  119 ( 567 expanded)
%              Number of equality atoms :   30 ( 135 expanded)
%              Maximal clause size      :    9 (   2 average)
%              Maximal term depth       :    6 (   2 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( app(app(app(app(sk7,cons(sk5,nil)),sk8),cons(sk6,nil)),sk9) = sk1 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( ssList(sk7) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( ssList(sk8) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_5"></A>cnf(c_0_5,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ssItem(sk5) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ssList(sk9) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( ssItem(sk6) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( leq(X2,X4)
    | app(app(X1,cons(X2,X3)),cons(X4,X5)) != X6
    | ~ ssList(X5)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X4)
    | ~ ssItem(X2)
    | ~ totalorderedP(X6)
    | ~ ssList(X6) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause182)).

<A NAME="c_0_13"></A>cnf(c_0_13,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( totalorderedP(cons(X1,nil))
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause67)).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( cons(sk11,nil) = sk3
    | nil = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_22)).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ssItem(sk11)
    | nil = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_18)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( X2 = X1
    | ~ leq(X1,X2)
    | ~ leq(X2,X1)
    | ~ ssItem(X1)
    | ~ ssItem(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause130)).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( leq(sk6,sk5) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( memberP(X4,X2)
    | app(X1,cons(X2,X3)) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause175)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( duplicatefreeP(cons(X1,nil))
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause65)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( totalorderedP(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause4)).

<A NAME="c_0_23"></A>cnf(c_0_23,axiom,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause71)).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( app(app(app(app(sk7,cons(sk5,nil)),sk8),cons(sk6,nil)),sk9) = sk3 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( app(app(app(sk7,app(cons(sk5,nil),sk8)),cons(sk6,nil)),sk9) = sk3
    | ~ ssList(cons(sk5,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_2]),c_0_3]),c_0_4])])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( app(app(app(sk7,app(cons(sk5,nil),sk8)),cons(sk6,nil)),sk9) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_5]),c_0_6]),c_0_7])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( app(app(app(sk7,cons(sk5,sk8)),cons(sk6,nil)),sk9) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>,c_0_8]),c_0_6]),c_0_4])])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( app(app(sk7,cons(sk5,sk8)),app(cons(sk6,nil),sk9)) = sk3
    | ~ ssList(app(sk7,cons(sk5,sk8)))
    | ~ ssList(cons(sk6,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_27]),c_0_9])])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( app(app(sk7,cons(sk5,sk8)),app(cons(sk6,nil),sk9)) = sk3
    | ~ ssList(app(sk7,cons(sk5,sk8))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_5]),c_0_10]),c_0_7])])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( app(app(sk7,cons(sk5,sk8)),app(cons(sk6,nil),sk9)) = sk3
    | ~ ssList(cons(sk5,sk8)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_11]),c_0_3])])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( app(app(sk7,cons(sk5,sk8)),app(cons(sk6,nil),sk9)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_5]),c_0_6]),c_0_4])])).

<A NAME="c_0_32"></A>cnf(c_0_32,plain,
    ( leq(X1,X2)
    | ~ ssItem(X2)
    | ~ ssItem(X1)
    | ~ ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ totalorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( app(app(sk7,cons(sk5,sk8)),cons(sk6,sk9)) = sk3 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,c_0_8]),c_0_10]),c_0_9])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( ssList(sk3) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( leq(sk5,sk6)
    | ~ totalorderedP(sk3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>,c_0_33]),c_0_10]),c_0_6]),c_0_34]),c_0_9]),c_0_4]),c_0_3])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( sk3 = nil
    | totalorderedP(sk3) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_15]),<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( sk3 = nil
    | leq(sk5,sk6) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( sk3 = nil
    | sk6 = sk5 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_37]),c_0_18]),c_0_6]),c_0_10])])).

<A NAME="c_0_39"></A>cnf(c_0_39,plain,
    ( memberP(app(X1,cons(X2,X3)),X2)
    | ~ ssItem(X2)
    | ~ ssList(app(X1,cons(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( app(app(sk7,cons(sk5,sk8)),cons(sk5,sk9)) = sk3
    | sk3 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( sk3 = nil
    | duplicatefreeP(sk3) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_15]),<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( memberP(sk3,sk6)
    | ~ ssList(app(sk7,cons(sk5,sk8))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_33]),c_0_10]),c_0_34]),c_0_9])])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( sk3 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_41]),c_0_6]),c_0_34]),c_0_9]),c_0_4]),c_0_3])]),<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( memberP(sk3,sk6)
    | ~ ssList(cons(sk5,sk8)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_11]),c_0_3])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( sk6 = sk5
    | ~ leq(sk5,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_18]),c_0_10]),c_0_6])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( leq(sk5,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_44]),c_0_22])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( memberP(sk3,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_45">c_0_45</A>,c_0_5]),c_0_6]),c_0_4])])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( sk6 = sk5 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_47])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( memberP(nil,sk5) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>,c_0_44]),<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>,c_0_50]),c_0_6])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC152-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n159.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:36:08 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.29  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.02/0.29  # and selection function SelectCQArNTEqFirst.
% 0.02/0.29  #
% 0.02/0.29  # Preprocessing time       : 0.009 s
% 0.02/0.29  # Presaturation interreduction done
% 0.02/0.29  
% 0.02/0.29  # Proof found!
% 0.02/0.29  # SZS status Unsatisfiable
% 0.02/0.29  # SZS output start CNFRefutation
% 0.02/0.29  cnf(c_0_0, negated_conjecture, (app(app(app(app(sk7,cons(sk5,nil)),sk8),cons(sk6,nil)),sk9)=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_12)).
% 0.02/0.29  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 0.02/0.29  cnf(c_0_2, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause149)).
% 0.02/0.29  cnf(c_0_3, negated_conjecture, (ssList(sk7)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
% 0.02/0.29  cnf(c_0_4, negated_conjecture, (ssList(sk8)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
% 0.02/0.29  cnf(c_0_5, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause86)).
% 0.02/0.29  cnf(c_0_6, negated_conjecture, (ssItem(sk5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 0.02/0.29  cnf(c_0_7, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.02/0.29  cnf(c_0_8, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause120)).
% 0.02/0.29  cnf(c_0_9, negated_conjecture, (ssList(sk9)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_11)).
% 0.02/0.29  cnf(c_0_10, negated_conjecture, (ssItem(sk6)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
% 0.02/0.29  cnf(c_0_11, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause85)).
% 0.02/0.29  cnf(c_0_12, axiom, (leq(X2,X4)|app(app(X1,cons(X2,X3)),cons(X4,X5))!=X6|~ssList(X5)|~ssList(X3)|~ssList(X1)|~ssItem(X4)|~ssItem(X2)|~totalorderedP(X6)|~ssList(X6)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause182)).
% 0.02/0.29  cnf(c_0_13, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_1)).
% 0.02/0.29  cnf(c_0_14, axiom, (totalorderedP(cons(X1,nil))|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause67)).
% 0.02/0.29  cnf(c_0_15, negated_conjecture, (cons(sk11,nil)=sk3|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_22)).
% 0.02/0.29  cnf(c_0_16, negated_conjecture, (ssItem(sk11)|nil=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_18)).
% 0.02/0.29  cnf(c_0_17, axiom, (X2=X1|~leq(X1,X2)|~leq(X2,X1)|~ssItem(X1)|~ssItem(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause130)).
% 0.02/0.29  cnf(c_0_18, negated_conjecture, (leq(sk6,sk5)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_13)).
% 0.02/0.29  cnf(c_0_19, axiom, (memberP(X4,X2)|app(X1,cons(X2,X3))!=X4|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~ssList(X4)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause175)).
% 0.02/0.29  cnf(c_0_20, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause179)).
% 0.02/0.29  cnf(c_0_21, axiom, (duplicatefreeP(cons(X1,nil))|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause65)).
% 0.02/0.29  cnf(c_0_22, axiom, (totalorderedP(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause4)).
% 0.02/0.29  cnf(c_0_23, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause71)).
% 0.02/0.29  cnf(c_0_24, negated_conjecture, (app(app(app(app(sk7,cons(sk5,nil)),sk8),cons(sk6,nil)),sk9)=sk3), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 0.02/0.29  cnf(c_0_25, negated_conjecture, (app(app(app(sk7,app(cons(sk5,nil),sk8)),cons(sk6,nil)),sk9)=sk3|~ssList(cons(sk5,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_2]), c_0_3]), c_0_4])])).
% 0.02/0.29  cnf(c_0_26, negated_conjecture, (app(app(app(sk7,app(cons(sk5,nil),sk8)),cons(sk6,nil)),sk9)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_5]), c_0_6]), c_0_7])])).
% 0.02/0.29  cnf(c_0_27, negated_conjecture, (app(app(app(sk7,cons(sk5,sk8)),cons(sk6,nil)),sk9)=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_8]), c_0_6]), c_0_4])])).
% 0.02/0.29  cnf(c_0_28, negated_conjecture, (app(app(sk7,cons(sk5,sk8)),app(cons(sk6,nil),sk9))=sk3|~ssList(app(sk7,cons(sk5,sk8)))|~ssList(cons(sk6,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_27]), c_0_9])])).
% 0.02/0.29  cnf(c_0_29, negated_conjecture, (app(app(sk7,cons(sk5,sk8)),app(cons(sk6,nil),sk9))=sk3|~ssList(app(sk7,cons(sk5,sk8)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_5]), c_0_10]), c_0_7])])).
% 0.02/0.29  cnf(c_0_30, negated_conjecture, (app(app(sk7,cons(sk5,sk8)),app(cons(sk6,nil),sk9))=sk3|~ssList(cons(sk5,sk8))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_11]), c_0_3])])).
% 0.02/0.29  cnf(c_0_31, negated_conjecture, (app(app(sk7,cons(sk5,sk8)),app(cons(sk6,nil),sk9))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_5]), c_0_6]), c_0_4])])).
% 0.02/0.29  cnf(c_0_32, plain, (leq(X1,X2)|~ssItem(X2)|~ssItem(X1)|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~totalorderedP(app(app(X3,cons(X1,X4)),cons(X2,X5)))), inference(er,[status(thm)],[c_0_12])).
% 0.02/0.29  cnf(c_0_33, negated_conjecture, (app(app(sk7,cons(sk5,sk8)),cons(sk6,sk9))=sk3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_8]), c_0_10]), c_0_9])])).
% 0.02/0.29  cnf(c_0_34, negated_conjecture, (ssList(sk3)), inference(rw,[status(thm)],[c_0_13, c_0_1])).
% 0.02/0.29  cnf(c_0_35, negated_conjecture, (leq(sk5,sk6)|~totalorderedP(sk3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_10]), c_0_6]), c_0_34]), c_0_9]), c_0_4]), c_0_3])])).
% 0.02/0.29  cnf(c_0_36, negated_conjecture, (sk3=nil|totalorderedP(sk3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
% 0.02/0.29  cnf(c_0_37, negated_conjecture, (sk3=nil|leq(sk5,sk6)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
% 0.02/0.29  cnf(c_0_38, negated_conjecture, (sk3=nil|sk6=sk5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_37]), c_0_18]), c_0_6]), c_0_10])])).
% 0.02/0.29  cnf(c_0_39, plain, (memberP(app(X1,cons(X2,X3)),X2)|~ssItem(X2)|~ssList(app(X1,cons(X2,X3)))|~ssList(X3)|~ssList(X1)), inference(er,[status(thm)],[c_0_19])).
% 0.02/0.29  cnf(c_0_40, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_20])).
% 0.02/0.29  cnf(c_0_41, negated_conjecture, (app(app(sk7,cons(sk5,sk8)),cons(sk5,sk9))=sk3|sk3=nil), inference(spm,[status(thm)],[c_0_33, c_0_38])).
% 0.02/0.29  cnf(c_0_42, negated_conjecture, (sk3=nil|duplicatefreeP(sk3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_15]), c_0_16])).
% 0.02/0.29  cnf(c_0_43, negated_conjecture, (memberP(sk3,sk6)|~ssList(app(sk7,cons(sk5,sk8)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_33]), c_0_10]), c_0_34]), c_0_9])])).
% 0.02/0.29  cnf(c_0_44, negated_conjecture, (sk3=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_6]), c_0_34]), c_0_9]), c_0_4]), c_0_3])]), c_0_42])).
% 0.02/0.29  cnf(c_0_45, negated_conjecture, (memberP(sk3,sk6)|~ssList(cons(sk5,sk8))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_11]), c_0_3])])).
% 0.02/0.29  cnf(c_0_46, negated_conjecture, (sk6=sk5|~leq(sk5,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_10]), c_0_6])])).
% 0.02/0.29  cnf(c_0_47, negated_conjecture, (leq(sk5,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_44]), c_0_22])])).
% 0.02/0.29  cnf(c_0_48, negated_conjecture, (memberP(sk3,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_5]), c_0_6]), c_0_4])])).
% 0.02/0.29  cnf(c_0_49, negated_conjecture, (sk6=sk5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47])])).
% 0.02/0.29  cnf(c_0_50, negated_conjecture, (memberP(nil,sk5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_44]), c_0_49])).
% 0.02/0.29  cnf(c_0_51, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_50]), c_0_6])]), ['proof']).
% 0.02/0.29  # SZS output end CNFRefutation
% 0.02/0.29  # Training examples: 0 positive, 0 negative
% 0.02/0.29  
% 0.02/0.29  # -------------------------------------------------
% 0.02/0.29  # User time                : 0.055 s
% 0.02/0.29  # System time              : 0.004 s
% 0.02/0.29  # Total time               : 0.059 s
% 0.02/0.29  # Maximum resident set size: 1868 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
