<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC315-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC315-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC315-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC315-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC315-1.p">SWC315-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n168.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:28 EDT 2016

% Result     : Unsatisfiable 0.07s
% Output     : CNFRefutation 0.07s
% Verified   : 
% Statistics : Number of clauses        :   49 ( 840 expanded)
%              Number of leaves         :   21 ( 546 expanded)
%              Depth                    :   11
%              Number of atoms          :   92 (1490 expanded)
%              Number of equality atoms :   60 ( 721 expanded)
%              Maximal clause size      :    4 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( app(cons(sk7,nil),sk8) = sk4
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_18)).

<A NAME="c_0_3"></A>cnf(c_0_3,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_4"></A>cnf(c_0_4,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( ssList(sk8)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_17)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ssItem(sk7)
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_16)).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( app(cons(sk5,nil),sk6) = sk2
    | nil = sk2 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_11)).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( ssList(sk6)
    | nil = sk2 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_10)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ssItem(sk5)
    | nil = sk2 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_11"></A>cnf(c_0_11,negated_conjecture,
    ( app(sk8,cons(sk7,nil)) = sk3
    | ~ neq(sk4,nil) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_19)).

<A NAME="c_0_12"></A>cnf(c_0_12,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause96)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause97)).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( nil = sk2
    | app(sk6,cons(sk5,nil)) != sk1 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_12)).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( nil = sk3
    | nil != sk4 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( app(cons(sk5,nil),sk6) = sk2
    | nil != sk1 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_14)).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( ssItem(sk5)
    | nil != sk1 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( ssList(sk6)
    | nil != sk1 ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( cons(X1,X2) != nil
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax',clause98)).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( app(cons(sk7,nil),sk8) = sk4
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_2">c_0_2</A>,c_0_3]),c_0_21]),c_0_4])])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( sk4 = nil
    | ssList(sk8) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,c_0_3]),c_0_21]),c_0_4])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( sk4 = nil
    | ssItem(sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>,c_0_3]),c_0_21]),c_0_4])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( app(cons(sk5,nil),sk6) = sk4
    | sk4 = nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_7">c_0_7</A>,c_0_1]),<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( sk4 = nil
    | ssList(sk6) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( sk4 = nil
    | ssItem(sk5) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( cons(sk7,sk8) = sk4
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_10]),c_0_23]),<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,negated_conjecture,
    ( cons(sk5,sk6) = sk4
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_25">c_0_25</A>,c_0_10]),c_0_26]),<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( app(sk8,cons(sk7,nil)) = sk1
    | ~ neq(sk4,nil) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_31"></A>cnf(c_0_31,negated_conjecture,
    ( tl(sk4) = sk8
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_28]),c_0_23]),<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,negated_conjecture,
    ( tl(sk4) = sk6
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_29]),c_0_26]),<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( app(sk8,cons(sk7,nil)) = sk1
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_3]),c_0_21]),c_0_4])])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( sk4 = nil
    | sk8 = sk6 ),
    inference(spm,[status(thm)],[<A HREF="#c_0_31">c_0_31</A>,<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( hd(sk4) = sk7
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_28]),c_0_23]),<A HREF="#c_0_24">c_0_24</A>])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( hd(sk4) = sk5
    | sk4 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_29]),c_0_26]),<A HREF="#c_0_27">c_0_27</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( app(sk6,cons(sk7,nil)) = sk1
    | sk4 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( sk4 = nil
    | sk7 = sk5 ),
    inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,<A HREF="#c_0_36">c_0_36</A>])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( sk4 = nil
    | app(sk6,cons(sk5,nil)) != sk1 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( sk1 = nil
    | sk4 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_16">c_0_16</A>,<A HREF="#c_0_12">c_0_12</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( sk4 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_38]),<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_42"></A>cnf(c_0_42,negated_conjecture,
    ( app(cons(sk5,nil),sk6) = sk4
    | sk1 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( sk1 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_41])])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( app(cons(sk5,nil),sk6) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>,c_0_41]),c_0_43])])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( ssItem(sk5) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_43])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( ssList(sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_43])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( cons(sk5,sk6) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_44]),c_0_45]),c_0_46])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_47]),c_0_45]),c_0_46])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.07  % Problem    : SWC315-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.08  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.07/0.30  % Computer   : n168.star.cs.uiowa.edu
% 0.07/0.30  % Model      : x86_64 x86_64
% 0.07/0.30  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.07/0.30  % Memory     : 32218.75MB
% 0.07/0.30  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.07/0.30  % CPULimit   : 300
% 0.07/0.30  % DateTime   : Tue May 10 20:40:53 CDT 2016
% 0.07/0.30  % CPUTime    : 
% 0.07/0.30  # Version: 1.9.1pre011
% 0.07/0.31  # No SInE strategy applied
% 0.07/0.31  # Trying AutoSched0 for 151 seconds
% 0.07/0.37  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.07/0.37  # and selection function SelectCQArNTEqFirst.
% 0.07/0.37  #
% 0.07/0.37  # Preprocessing time       : 0.019 s
% 0.07/0.37  # Presaturation interreduction done
% 0.07/0.37  
% 0.07/0.37  # Proof found!
% 0.07/0.37  # SZS status Unsatisfiable
% 0.07/0.37  # SZS output start CNFRefutation
% 0.07/0.37  cnf(c_0_0, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_2)).
% 0.07/0.37  cnf(c_0_1, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_5)).
% 0.07/0.37  cnf(c_0_2, negated_conjecture, (app(cons(sk7,nil),sk8)=sk4|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_18)).
% 0.07/0.37  cnf(c_0_3, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause100)).
% 0.07/0.37  cnf(c_0_4, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.07/0.37  cnf(c_0_5, negated_conjecture, (ssList(sk8)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_17)).
% 0.07/0.37  cnf(c_0_6, negated_conjecture, (ssItem(sk7)|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_16)).
% 0.07/0.37  cnf(c_0_7, negated_conjecture, (app(cons(sk5,nil),sk6)=sk2|nil=sk2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_11)).
% 0.07/0.37  cnf(c_0_8, negated_conjecture, (ssList(sk6)|nil=sk2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_10)).
% 0.07/0.37  cnf(c_0_9, negated_conjecture, (ssItem(sk5)|nil=sk2), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_8)).
% 0.07/0.37  cnf(c_0_10, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause120)).
% 0.07/0.37  cnf(c_0_11, negated_conjecture, (app(sk8,cons(sk7,nil))=sk3|~neq(sk4,nil)), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_19)).
% 0.07/0.37  cnf(c_0_12, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_6)).
% 0.07/0.37  cnf(c_0_13, axiom, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause96)).
% 0.07/0.37  cnf(c_0_14, axiom, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause97)).
% 0.07/0.37  cnf(c_0_15, negated_conjecture, (nil=sk2|app(sk6,cons(sk5,nil))!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_12)).
% 0.07/0.37  cnf(c_0_16, negated_conjecture, (nil=sk3|nil!=sk4), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_7)).
% 0.07/0.37  cnf(c_0_17, negated_conjecture, (app(cons(sk5,nil),sk6)=sk2|nil!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_14)).
% 0.07/0.37  cnf(c_0_18, negated_conjecture, (ssItem(sk5)|nil!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_9)).
% 0.07/0.37  cnf(c_0_19, negated_conjecture, (ssList(sk6)|nil!=sk1), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1_13)).
% 0.07/0.37  cnf(c_0_20, axiom, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001-0.ax', clause98)).
% 0.07/0.37  cnf(c_0_21, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 0.07/0.37  cnf(c_0_22, negated_conjecture, (app(cons(sk7,nil),sk8)=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_2, c_0_3]), c_0_21]), c_0_4])])).
% 0.07/0.37  cnf(c_0_23, negated_conjecture, (sk4=nil|ssList(sk8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_3]), c_0_21]), c_0_4])])).
% 0.07/0.37  cnf(c_0_24, negated_conjecture, (sk4=nil|ssItem(sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_3]), c_0_21]), c_0_4])])).
% 0.07/0.37  cnf(c_0_25, negated_conjecture, (app(cons(sk5,nil),sk6)=sk4|sk4=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_1]), c_0_1])).
% 0.07/0.37  cnf(c_0_26, negated_conjecture, (sk4=nil|ssList(sk6)), inference(rw,[status(thm)],[c_0_8, c_0_1])).
% 0.07/0.37  cnf(c_0_27, negated_conjecture, (sk4=nil|ssItem(sk5)), inference(rw,[status(thm)],[c_0_9, c_0_1])).
% 0.07/0.37  cnf(c_0_28, negated_conjecture, (cons(sk7,sk8)=sk4|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_10]), c_0_23]), c_0_24])).
% 0.07/0.37  cnf(c_0_29, negated_conjecture, (cons(sk5,sk6)=sk4|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_10]), c_0_26]), c_0_27])).
% 0.07/0.37  cnf(c_0_30, negated_conjecture, (app(sk8,cons(sk7,nil))=sk1|~neq(sk4,nil)), inference(rw,[status(thm)],[c_0_11, c_0_12])).
% 0.07/0.37  cnf(c_0_31, negated_conjecture, (tl(sk4)=sk8|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_28]), c_0_23]), c_0_24])).
% 0.07/0.37  cnf(c_0_32, negated_conjecture, (tl(sk4)=sk6|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_29]), c_0_26]), c_0_27])).
% 0.07/0.37  cnf(c_0_33, negated_conjecture, (app(sk8,cons(sk7,nil))=sk1|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_3]), c_0_21]), c_0_4])])).
% 0.07/0.37  cnf(c_0_34, negated_conjecture, (sk4=nil|sk8=sk6), inference(spm,[status(thm)],[c_0_31, c_0_32])).
% 0.07/0.37  cnf(c_0_35, negated_conjecture, (hd(sk4)=sk7|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_28]), c_0_23]), c_0_24])).
% 0.07/0.37  cnf(c_0_36, negated_conjecture, (hd(sk4)=sk5|sk4=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_29]), c_0_26]), c_0_27])).
% 0.07/0.37  cnf(c_0_37, negated_conjecture, (app(sk6,cons(sk7,nil))=sk1|sk4=nil), inference(spm,[status(thm)],[c_0_33, c_0_34])).
% 0.07/0.37  cnf(c_0_38, negated_conjecture, (sk4=nil|sk7=sk5), inference(spm,[status(thm)],[c_0_35, c_0_36])).
% 0.07/0.37  cnf(c_0_39, negated_conjecture, (sk4=nil|app(sk6,cons(sk5,nil))!=sk1), inference(rw,[status(thm)],[c_0_15, c_0_1])).
% 0.07/0.37  cnf(c_0_40, negated_conjecture, (sk1=nil|sk4!=nil), inference(rw,[status(thm)],[c_0_16, c_0_12])).
% 0.07/0.37  cnf(c_0_41, negated_conjecture, (sk4=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
% 0.07/0.37  cnf(c_0_42, negated_conjecture, (app(cons(sk5,nil),sk6)=sk4|sk1!=nil), inference(rw,[status(thm)],[c_0_17, c_0_1])).
% 0.07/0.37  cnf(c_0_43, negated_conjecture, (sk1=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41])])).
% 0.07/0.37  cnf(c_0_44, negated_conjecture, (app(cons(sk5,nil),sk6)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_41]), c_0_43])])).
% 0.07/0.37  cnf(c_0_45, negated_conjecture, (ssItem(sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_43])])).
% 0.07/0.37  cnf(c_0_46, negated_conjecture, (ssList(sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_43])])).
% 0.07/0.37  cnf(c_0_47, negated_conjecture, (cons(sk5,sk6)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_44]), c_0_45]), c_0_46])])).
% 0.07/0.37  cnf(c_0_48, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_47]), c_0_45]), c_0_46])]), ['proof']).
% 0.07/0.37  # SZS output end CNFRefutation
% 0.07/0.37  # Training examples: 0 positive, 0 negative
% 0.07/0.37  
% 0.07/0.37  # -------------------------------------------------
% 0.07/0.37  # User time                : 0.063 s
% 0.07/0.37  # System time              : 0.004 s
% 0.07/0.37  # Total time               : 0.067 s
% 0.07/0.37  # Maximum resident set size: 1864 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
