<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC066+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC066+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC066+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC066%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC066+1.p">SWC066+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n128.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:30 EDT 2016

% Result     : Theorem 49.35s
% Output     : CNFRefutation 49.35s
% Verified   : 
% Statistics : Number of formulae       :  347 (1046505 expanded)
%              Number of clauses        :  267 (357975 expanded)
%              Number of leaves         :   40 (256206 expanded)
%              Depth                    :   37
%              Number of atoms          : 1046 (13928142 expanded)
%              Number of equality atoms :  330 (2823381 expanded)
%              Maximal formula depth    :   21 (   4 average)
%              Maximal clause size      :   27 (   2 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ~ segmentP(X4,X3)
                    | ~ totalorderedP(X3)
                    | ? [X5] :
                        ( ssList(X5)
                        & neq(X3,X5)
                        & segmentP(X4,X5)
                        & segmentP(X5,X3)
                        & totalorderedP(X5) )
                    | ( ( nil != X2
                        | nil = X1 )
                      & ( ~ neq(X2,nil)
                        | ? [X6] :
                            ( ssList(X6)
                            & neq(X6,nil)
                            & segmentP(X2,X6)
                            & segmentP(X1,X6) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( segmentP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax7)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( frontsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X2,X3) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax5)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( rearsegP(X1,X2)
               => rearsegP(app(X3,X1),X2) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax50)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => rearsegP(X1,nil) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax51)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => cons(hd(X1),tl(X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax78)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => nil != cons(X2,X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax21)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( singletonP(X1)
      &lt;=> ? [X2] :
            ( ssItem(X2)
            & cons(X2,nil) = X1 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax4)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( rearsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax6)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( frontsegP(X1,X2)
               => frontsegP(app(X1,X3),X2) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax43)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,nil) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax45)).

<A NAME="c_0_17"></A>fof(c_0_17,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X3,X2) = app(X1,X2)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax79)).

<A NAME="c_0_18"></A>fof(c_0_18,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_19"></A>fof(c_0_19,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X2,X3) = app(X2,X1)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax80)).

<A NAME="c_0_20"></A>fof(c_0_20,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssList(tl(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

<A NAME="c_0_21"></A>fof(c_0_21,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax86)).

<A NAME="c_0_22"></A>fof(c_0_22,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_23"></A>fof(c_0_23,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => hd(cons(X2,X1)) = X2 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax23)).

<A NAME="c_0_24"></A>fof(c_0_24,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => hd(app(X1,X2)) = hd(X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax85)).

<A NAME="c_0_25"></A>fof(c_0_25,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_26"></A>fof(c_0_26,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssItem(hd(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax22)).

<A NAME="c_0_27"></A>fof(c_0_27,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( frontsegP(cons(X1,X3),cons(X2,X4))
                  &lt;=> ( X1 = X2
                      & frontsegP(X3,X4) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax44)).

<A NAME="c_0_28"></A>fof(c_0_28,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil = X1
        | ? [X2] :
            ( ssList(X2)
            & ? [X3] :
                ( ssItem(X3)
                & cons(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax20)).

<A NAME="c_0_29"></A>fof(c_0_29,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax55)).

<A NAME="c_0_30"></A>fof(c_0_30,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

<A NAME="c_0_31"></A>fof(c_0_31,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_32"></A>fof(c_0_32,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( segmentP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax58)).

<A NAME="c_0_33"></A>fof(c_0_33,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( segmentP(X1,X2)
                   => segmentP(app(app(X3,X1),X4),X2) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax56)).

<A NAME="c_0_34"></A>fof(c_0_34,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax42)).

<A NAME="c_0_35"></A>fof(c_0_35,axiom,(
    ! [X1] :
      ( ssList(X1)
     => rearsegP(X1,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax49)).

<A NAME="c_0_36"></A>fof(c_0_36,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( rearsegP(nil,X1)
      &lt;=> nil = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax52)).

<A NAME="c_0_37"></A>fof(c_0_37,axiom,(
    ! [X1] :
      ( ssList(X1)
     => segmentP(X1,nil) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax57)).

<A NAME="c_0_38"></A>fof(c_0_38,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => totalorderedP(cons(X1,nil)) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax65)).

<A NAME="c_0_39"></A>fof(c_0_39,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) != X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax18)).

<A NAME="c_0_40"></A>fof(c_0_40,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ~ segmentP(X4,X3)
                      | ~ totalorderedP(X3)
                      | ? [X5] :
                          ( ssList(X5)
                          & neq(X3,X5)
                          & segmentP(X4,X5)
                          & segmentP(X5,X3)
                          & totalorderedP(X5) )
                      | ( ( nil != X2
                          | nil = X1 )
                        & ( ~ neq(X2,nil)
                          | ? [X6] :
                              ( ssList(X6)
                              & neq(X6,nil)
                              & segmentP(X2,X6)
                              & segmentP(X1,X6) ) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_41"></A>fof(c_0_41,negated_conjecture,(
    ! [X11,X12] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & segmentP(esk51_0,esk50_0)
      & totalorderedP(esk50_0)
      & ( ~ ssList(X11)
        | ~ neq(esk50_0,X11)
        | ~ segmentP(esk51_0,X11)
        | ~ segmentP(X11,esk50_0)
        | ~ totalorderedP(X11) )
      & ( neq(esk49_0,nil)
        | nil = esk49_0 )
      & ( ~ ssList(X12)
        | ~ neq(X12,nil)
        | ~ segmentP(esk49_0,X12)
        | ~ segmentP(esk48_0,X12)
        | nil = esk49_0 )
      & ( neq(esk49_0,nil)
        | nil != esk48_0 )
      & ( ~ ssList(X12)
        | ~ neq(X12,nil)
        | ~ segmentP(esk49_0,X12)
        | ~ segmentP(esk48_0,X12)
        | nil != esk48_0 ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])])])])).

<A NAME="c_0_42"></A>fof(c_0_42,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk8_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ssList(esk9_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(app(X9,X6),X10) != X5
        | segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( segmentP(esk51_0,esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_46"></A>fof(c_0_46,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk6_2(X4,X5))
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(X5,esk6_2(X4,X5)) = X4
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X5,X7) != X4
        | frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).

<A NAME="c_0_47"></A>fof(c_0_47,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).

<A NAME="c_0_48"></A>fof(c_0_48,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_49"></A>cnf(c_0_49,plain,
    ( app(app(esk8_2(X1,X2),X2),esk9_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ segmentP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( segmentP(esk49_0,esk48_0) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_44]),<A HREF="#c_0_45">c_0_45</A>])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( ssList(esk49_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( ssList(esk9_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ segmentP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_54"></A>cnf(c_0_54,plain,
    ( ssList(esk8_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ segmentP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_55"></A>cnf(c_0_55,plain,
    ( frontsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X2,X3) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,negated_conjecture,
    ( app(app(esk8_2(esk49_0,esk48_0),esk48_0),esk9_2(esk49_0,esk48_0)) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_50]),c_0_51]),c_0_52])])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( ssList(esk9_2(esk49_0,esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_50]),c_0_51]),c_0_52])])).

<A NAME="c_0_60"></A>cnf(c_0_60,negated_conjecture,
    ( ssList(esk8_2(esk49_0,esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_50]),c_0_51]),c_0_52])])).

<A NAME="c_0_61"></A>cnf(c_0_61,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_55]),<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( app(esk8_2(esk49_0,esk48_0),app(esk48_0,esk9_2(esk49_0,esk48_0))) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,c_0_58]),c_0_59]),c_0_51]),c_0_60])])).

<A NAME="c_0_63"></A>fof(c_0_63,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ rearsegP(X4,X5)
      | rearsegP(app(X6,X4),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).

<A NAME="c_0_64"></A>fof(c_0_64,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | rearsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).

<A NAME="c_0_65"></A>cnf(c_0_65,negated_conjecture,
    ( frontsegP(esk49_0,esk8_2(esk49_0,esk48_0))
    | ~ ssList(app(esk48_0,esk9_2(esk49_0,esk48_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_62]),c_0_60])])).

<A NAME="c_0_66"></A>fof(c_0_66,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | cons(hd(X2),tl(X2)) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

<A NAME="c_0_67"></A>fof(c_0_67,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

<A NAME="c_0_68"></A>fof(c_0_68,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_69"></A>fof(c_0_69,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | nil != cons(X4,X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).

<A NAME="c_0_70"></A>fof(c_0_70,plain,(
    ! [X3,X5] :
      ( ( ssItem(esk5_1(X3))
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( cons(esk5_1(X3),nil) = X3
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( ~ ssItem(X5)
        | cons(X5,nil) != X3
        | singletonP(X3)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).

<A NAME="c_0_71"></A>fof(c_0_71,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk7_2(X4,X5))
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(esk7_2(X4,X5),X5) = X4
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X7,X5) != X4
        | rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( rearsegP(app(X1,X2),X3)
    | ~ rearsegP(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,plain,
    ( rearsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>])).

<A NAME="c_0_74"></A>cnf(c_0_74,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_75"></A>cnf(c_0_75,plain,
    ( app(X2,esk6_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,negated_conjecture,
    ( frontsegP(esk49_0,esk8_2(esk49_0,esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_56]),c_0_59]),c_0_51])])).

<A NAME="c_0_77"></A>cnf(c_0_77,plain,
    ( ssList(esk6_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>])).

<A NAME="c_0_78"></A>cnf(c_0_78,plain,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_80"></A>cnf(c_0_80,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( nil != cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,plain,
    ( singletonP(X1)
    | ~ ssList(X1)
    | cons(X2,nil) != X1
    | ~ ssItem(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,plain,
    ( rearsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X3,X2) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>])).

<A NAME="c_0_84"></A>cnf(c_0_84,plain,
    ( rearsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_73]),c_0_74])])).

<A NAME="c_0_85"></A>cnf(c_0_85,negated_conjecture,
    ( app(esk8_2(esk49_0,esk48_0),esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_76]),c_0_60]),c_0_52])])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( ssList(esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_76]),c_0_60]),c_0_52])])).

<A NAME="c_0_87"></A>cnf(c_0_87,plain,
    ( cons(hd(cons(X1,X2)),X2) = cons(X1,X2)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_79]),c_0_80]),<A HREF="#c_0_81">c_0_81</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,plain,
    ( cons(esk5_1(X1),nil) = X1
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,plain,
    ( ssItem(esk5_1(X1))
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>])).

<A NAME="c_0_90"></A>cnf(c_0_90,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>])).

<A NAME="c_0_91"></A>cnf(c_0_91,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_83]),<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_92"></A>fof(c_0_92,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).

<A NAME="c_0_93"></A>cnf(c_0_93,plain,
    ( app(esk7_2(X1,X2),X2) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>])).

<A NAME="c_0_94"></A>cnf(c_0_94,negated_conjecture,
    ( rearsegP(esk49_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,c_0_85]),c_0_86]),c_0_60])])).

<A NAME="c_0_95"></A>cnf(c_0_95,plain,
    ( ssList(esk7_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>])).

<A NAME="c_0_96"></A>cnf(c_0_96,plain,
    ( cons(hd(X1),nil) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>,c_0_88]),c_0_74])]),<A HREF="#c_0_89">c_0_89</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,c_0_80]),c_0_74])])).

<A NAME="c_0_98"></A>cnf(c_0_98,negated_conjecture,
    ( rearsegP(esk49_0,esk9_2(esk49_0,esk48_0))
    | ~ ssList(app(esk8_2(esk49_0,esk48_0),esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_58]),c_0_59])])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,negated_conjecture,
    ( app(esk7_2(esk49_0,nil),nil) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_94]),c_0_74]),c_0_52])])).

<A NAME="c_0_101"></A>cnf(c_0_101,negated_conjecture,
    ( ssList(esk7_2(esk49_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_94]),c_0_74]),c_0_52])])).

<A NAME="c_0_102"></A>fof(c_0_102,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ frontsegP(X4,X5)
      | frontsegP(app(X4,X6),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).

<A NAME="c_0_103"></A>fof(c_0_103,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).

<A NAME="c_0_104"></A>cnf(c_0_104,plain,
    ( cons(hd(cons(X1,nil)),nil) = cons(X1,nil)
    | ~ ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,<A HREF="#c_0_97">c_0_97</A>])).

<A NAME="c_0_105"></A>fof(c_0_105,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X6,X5) != app(X4,X5)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).

<A NAME="c_0_106"></A>fof(c_0_106,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])).

<A NAME="c_0_107"></A>fof(c_0_107,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X5,X6) != app(X5,X4)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).

<A NAME="c_0_108"></A>fof(c_0_108,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssList(tl(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])).

<A NAME="c_0_109"></A>fof(c_0_109,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).

<A NAME="c_0_110"></A>fof(c_0_110,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).

<A NAME="c_0_111"></A>cnf(c_0_111,negated_conjecture,
    ( rearsegP(esk49_0,esk9_2(esk49_0,esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_98">c_0_98</A>,c_0_56]),c_0_51]),c_0_60])])).

<A NAME="c_0_112"></A>cnf(c_0_112,negated_conjecture,
    ( esk7_2(esk49_0,nil) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_100]),c_0_101])])).

<A NAME="c_0_113"></A>cnf(c_0_113,plain,
    ( frontsegP(app(X1,X2),X3)
    | ~ frontsegP(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>])).

<A NAME="c_0_114"></A>cnf(c_0_114,plain,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_103">c_0_103</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,plain,
    ( cons(hd(cons(X1,nil)),nil) = cons(X1,nil)
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>,c_0_80]),c_0_74])])).

<A NAME="c_0_116"></A>fof(c_0_116,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | hd(cons(X4,X3)) = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).

<A NAME="c_0_117"></A>fof(c_0_117,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | hd(app(X3,X4)) = hd(X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).

<A NAME="c_0_118"></A>cnf(c_0_118,plain,
    ( X1 = X2
    | app(X1,X3) != app(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_105">c_0_105</A>])).

<A NAME="c_0_119"></A>cnf(c_0_119,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>])).

<A NAME="c_0_120"></A>cnf(c_0_120,plain,
    ( X1 = X2
    | app(X3,X1) != app(X3,X2)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_107">c_0_107</A>])).

<A NAME="c_0_121"></A>cnf(c_0_121,plain,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>])).

<A NAME="c_0_122"></A>cnf(c_0_122,plain,
    ( tl(app(X1,X2)) = app(tl(X1),X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_109">c_0_109</A>])).

<A NAME="c_0_123"></A>cnf(c_0_123,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_124"></A>cnf(c_0_124,negated_conjecture,
    ( app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),esk9_2(esk49_0,esk48_0)) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_111]),c_0_59]),c_0_52])])).

<A NAME="c_0_125"></A>cnf(c_0_125,negated_conjecture,
    ( ssList(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_111]),c_0_59]),c_0_52])])).

<A NAME="c_0_126"></A>cnf(c_0_126,negated_conjecture,
    ( app(esk49_0,nil) = esk49_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>,<A HREF="#c_0_112">c_0_112</A>])).

<A NAME="c_0_127"></A>cnf(c_0_127,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_113">c_0_113</A>,c_0_114]),c_0_74])])).

<A NAME="c_0_128"></A>fof(c_0_128,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).

<A NAME="c_0_129"></A>cnf(c_0_129,plain,
    ( ssList(cons(X1,nil))
    | ~ ssItem(hd(cons(X1,nil)))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,c_0_115]),c_0_74])])).

<A NAME="c_0_130"></A>cnf(c_0_130,plain,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_116">c_0_116</A>])).

<A NAME="c_0_131"></A>cnf(c_0_131,plain,
    ( hd(app(X1,X2)) = hd(X1)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_117">c_0_117</A>])).

<A NAME="c_0_132"></A>cnf(c_0_132,plain,
    ( X1 = nil
    | app(X1,X2) != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_118">c_0_118</A>,c_0_119]),c_0_74])])).

<A NAME="c_0_133"></A>cnf(c_0_133,plain,
    ( X1 = nil
    | app(X2,X1) != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_120">c_0_120</A>,c_0_99]),c_0_74])])).

<A NAME="c_0_134"></A>cnf(c_0_134,plain,
    ( nil = X1
    | ssList(app(tl(X1),X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_121">c_0_121</A>,c_0_122]),c_0_56]),<A HREF="#c_0_123">c_0_123</A>])).

<A NAME="c_0_135"></A>cnf(c_0_135,negated_conjecture,
    ( app(tl(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))),esk9_2(esk49_0,esk48_0)) = tl(esk49_0)
    | esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_122">c_0_122</A>,c_0_124]),c_0_59]),c_0_125])])).

<A NAME="c_0_136"></A>cnf(c_0_136,negated_conjecture,
    ( app(tl(esk49_0),nil) = tl(esk49_0)
    | esk49_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_122">c_0_122</A>,c_0_126]),c_0_74]),c_0_52])])).

<A NAME="c_0_137"></A>cnf(c_0_137,negated_conjecture,
    ( frontsegP(esk49_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_127">c_0_127</A>,c_0_85]),c_0_86]),c_0_60])])).

<A NAME="c_0_138"></A>cnf(c_0_138,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_128">c_0_128</A>])).

<A NAME="c_0_139"></A>cnf(c_0_139,plain,
    ( ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_129">c_0_129</A>,c_0_130]),c_0_74])])).

<A NAME="c_0_140"></A>cnf(c_0_140,plain,
    ( cons(hd(X1),tl(app(X1,X2))) = app(X1,X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_131]),c_0_56]),<A HREF="#c_0_123">c_0_123</A>])).

<A NAME="c_0_141"></A>cnf(c_0_141,negated_conjecture,
    ( app(esk8_2(esk49_0,esk48_0),esk48_0) = nil
    | esk9_2(esk49_0,esk48_0) != esk49_0
    | ~ ssList(app(esk8_2(esk49_0,esk48_0),esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_132">c_0_132</A>,c_0_58]),c_0_59])])).

<A NAME="c_0_142"></A>cnf(c_0_142,negated_conjecture,
    ( esk9_2(esk49_0,esk48_0) = nil
    | esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)) != esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_133">c_0_133</A>,c_0_124]),c_0_125]),c_0_59])])).

<A NAME="c_0_143"></A>cnf(c_0_143,negated_conjecture,
    ( esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)) = nil
    | ssList(tl(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_134">c_0_134</A>,c_0_135]),c_0_59]),c_0_125])])).

<A NAME="c_0_144"></A>cnf(c_0_144,negated_conjecture,
    ( esk49_0 = nil
    | ssList(tl(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_134">c_0_134</A>,c_0_136]),c_0_74]),c_0_52])])).

<A NAME="c_0_145"></A>cnf(c_0_145,plain,
    ( app(tl(X1),nil) = tl(X1)
    | nil = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_122">c_0_122</A>,c_0_99]),c_0_74])])).

<A NAME="c_0_146"></A>fof(c_0_146,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssItem(hd(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])).

<A NAME="c_0_147"></A>cnf(c_0_147,negated_conjecture,
    ( app(nil,esk6_2(esk49_0,nil)) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_137]),c_0_74]),c_0_52])])).

<A NAME="c_0_148"></A>cnf(c_0_148,negated_conjecture,
    ( ssList(esk6_2(esk49_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_137]),c_0_74]),c_0_52])])).

<A NAME="c_0_149"></A>fof(c_0_149,plain,(
    ! [X5,X6,X7,X8] :
      ( ( X5 = X6
        | ~ frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) )
      & ( frontsegP(X7,X8)
        | ~ frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) )
      & ( X5 != X6
        | ~ frontsegP(X7,X8)
        | frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])).

<A NAME="c_0_150"></A>fof(c_0_150,plain,(
    ! [X4] :
      ( ( ssList(esk44_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( ssItem(esk45_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( cons(esk45_1(X4),esk44_1(X4)) = X4
        | nil = X4
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])])).

<A NAME="c_0_151"></A>cnf(c_0_151,plain,
    ( frontsegP(cons(X1,X2),cons(X1,nil))
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_138]),<A HREF="#c_0_139">c_0_139</A>])).

<A NAME="c_0_152"></A>cnf(c_0_152,negated_conjecture,
    ( cons(hd(esk49_0),tl(esk49_0)) = esk49_0
    | esk49_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_140">c_0_140</A>,c_0_126]),c_0_74]),c_0_52])])).

<A NAME="c_0_153"></A>cnf(c_0_153,negated_conjecture,
    ( app(esk8_2(esk49_0,esk48_0),esk48_0) = nil
    | esk9_2(esk49_0,esk48_0) != esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_141">c_0_141</A>,c_0_56]),c_0_51]),c_0_60])])).

<A NAME="c_0_154"></A>cnf(c_0_154,negated_conjecture,
    ( esk9_2(esk49_0,esk48_0) = nil
    | ssList(tl(esk49_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_142">c_0_142</A>,c_0_143]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_155"></A>cnf(c_0_155,negated_conjecture,
    ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil
    | esk8_2(esk49_0,esk48_0) != esk49_0
    | ~ ssList(app(esk48_0,esk9_2(esk49_0,esk48_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_133">c_0_133</A>,c_0_62]),c_0_60])])).

<A NAME="c_0_156"></A>cnf(c_0_156,plain,
    ( app(app(tl(X1),X2),nil) = app(tl(X1),X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_145">c_0_145</A>,c_0_122]),c_0_56]),<A HREF="#c_0_123">c_0_123</A>])).

<A NAME="c_0_157"></A>cnf(c_0_157,negated_conjecture,
    ( app(tl(esk8_2(esk49_0,esk48_0)),esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))) = tl(esk49_0)
    | esk8_2(esk49_0,esk48_0) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_122">c_0_122</A>,c_0_85]),c_0_86]),c_0_60])])).

<A NAME="c_0_158"></A>cnf(c_0_158,plain,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_146">c_0_146</A>])).

<A NAME="c_0_159"></A>cnf(c_0_159,negated_conjecture,
    ( hd(esk8_2(esk49_0,esk48_0)) = hd(esk49_0)
    | esk8_2(esk49_0,esk48_0) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>,c_0_85]),c_0_86]),c_0_60])])).

<A NAME="c_0_160"></A>fof(c_0_160,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])).

<A NAME="c_0_161"></A>cnf(c_0_161,plain,
    ( segmentP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(app(X3,X2),X4) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_162"></A>cnf(c_0_162,negated_conjecture,
    ( esk6_2(esk49_0,nil) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>,c_0_147]),c_0_148])])).

<A NAME="c_0_163"></A>cnf(c_0_163,plain,
    ( X1 = X2
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ ssList(X4)
    | ~ frontsegP(cons(X1,X3),cons(X2,X4)) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_149">c_0_149</A>])).

<A NAME="c_0_164"></A>cnf(c_0_164,plain,
    ( nil = X1
    | cons(esk45_1(X1),esk44_1(X1)) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_150">c_0_150</A>])).

<A NAME="c_0_165"></A>cnf(c_0_165,plain,
    ( nil = X1
    | ssItem(esk45_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_150">c_0_150</A>])).

<A NAME="c_0_166"></A>cnf(c_0_166,plain,
    ( nil = X1
    | ssList(esk44_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_150">c_0_150</A>])).

<A NAME="c_0_167"></A>cnf(c_0_167,negated_conjecture,
    ( esk49_0 = nil
    | frontsegP(esk49_0,cons(hd(esk49_0),nil))
    | ~ ssItem(hd(esk49_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_151">c_0_151</A>,c_0_152]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_168"></A>cnf(c_0_168,plain,
    ( nil = X2
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_169"></A>cnf(c_0_169,negated_conjecture,
    ( app(esk8_2(esk49_0,esk48_0),esk48_0) = nil
    | ssList(tl(esk49_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_153">c_0_153</A>,c_0_154]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_170"></A>fof(c_0_170,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssItem(X6)
      | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).

<A NAME="c_0_171"></A>cnf(c_0_171,negated_conjecture,
    ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil
    | esk8_2(esk49_0,esk48_0) != esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_155">c_0_155</A>,c_0_56]),c_0_59]),c_0_51])])).

<A NAME="c_0_172"></A>cnf(c_0_172,negated_conjecture,
    ( app(tl(esk49_0),nil) = tl(esk49_0)
    | esk8_2(esk49_0,esk48_0) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_156">c_0_156</A>,c_0_157]),c_0_86]),c_0_60])])).

<A NAME="c_0_173"></A>cnf(c_0_173,negated_conjecture,
    ( esk8_2(esk49_0,esk48_0) = nil
    | ssItem(hd(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_158">c_0_158</A>,c_0_159]),c_0_60])])).

<A NAME="c_0_174"></A>cnf(c_0_174,negated_conjecture,
    ( nil = esk49_0
    | ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_175"></A>cnf(c_0_175,plain,
    ( segmentP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_160">c_0_160</A>])).

<A NAME="c_0_176"></A>fof(c_0_176,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])])).

<A NAME="c_0_177"></A>cnf(c_0_177,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_161">c_0_161</A>])).

<A NAME="c_0_178"></A>cnf(c_0_178,negated_conjecture,
    ( app(nil,esk49_0) = esk49_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_147">c_0_147</A>,<A HREF="#c_0_162">c_0_162</A>])).

<A NAME="c_0_179"></A>cnf(c_0_179,plain,
    ( esk45_1(X1) = X2
    | nil = X1
    | ~ frontsegP(X1,cons(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_163">c_0_163</A>,c_0_164]),c_0_165]),<A HREF="#c_0_166">c_0_166</A>])).

<A NAME="c_0_180"></A>cnf(c_0_180,negated_conjecture,
    ( esk49_0 = nil
    | frontsegP(esk49_0,cons(hd(esk49_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_167">c_0_167</A>,c_0_158]),c_0_52])])).

<A NAME="c_0_181"></A>cnf(c_0_181,negated_conjecture,
    ( esk48_0 = nil
    | ssList(tl(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_168">c_0_168</A>,c_0_169]),c_0_51]),c_0_60])])).

<A NAME="c_0_182"></A>cnf(c_0_182,plain,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_170">c_0_170</A>])).

<A NAME="c_0_183"></A>cnf(c_0_183,negated_conjecture,
    ( app(tl(esk49_0),nil) = tl(esk49_0)
    | app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_171">c_0_171</A>,c_0_172]),<A HREF="#c_0_136">c_0_136</A>])).

<A NAME="c_0_184"></A>cnf(c_0_184,negated_conjecture,
    ( app(app(nil,esk48_0),esk9_2(esk49_0,esk48_0)) = esk49_0
    | ssItem(hd(esk49_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,<A HREF="#c_0_173">c_0_173</A>])).

<A NAME="c_0_185"></A>cnf(c_0_185,negated_conjecture,
    ( esk49_0 = nil
    | ~ neq(esk48_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_174">c_0_174</A>,c_0_175]),c_0_50]),c_0_51])])).

<A NAME="c_0_186"></A>cnf(c_0_186,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_176">c_0_176</A>])).

<A NAME="c_0_187"></A>fof(c_0_187,plain,(
    ! [X2] :
      ( ( ~ segmentP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | segmentP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])).

<A NAME="c_0_188"></A>cnf(c_0_188,negated_conjecture,
    ( segmentP(app(esk49_0,X1),esk49_0)
    | ~ ssList(app(esk49_0,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_177">c_0_177</A>,c_0_178]),c_0_74]),c_0_52])])).

<A NAME="c_0_189"></A>cnf(c_0_189,negated_conjecture,
    ( esk45_1(esk49_0) = hd(esk49_0)
    | esk49_0 = nil
    | ~ ssItem(hd(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_179">c_0_179</A>,c_0_180]),c_0_74]),c_0_52])])).

<A NAME="c_0_190"></A>cnf(c_0_190,negated_conjecture,
    ( esk8_2(esk49_0,esk48_0) = nil
    | ssList(tl(esk49_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_132">c_0_132</A>,c_0_169]),c_0_51]),c_0_60])]),<A HREF="#c_0_181">c_0_181</A>])).

<A NAME="c_0_191"></A>cnf(c_0_191,plain,
    ( cons(X1,app(X2,X3)) != nil
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>,c_0_182]),c_0_80]),<A HREF="#c_0_81">c_0_81</A>])).

<A NAME="c_0_192"></A>cnf(c_0_192,negated_conjecture,
    ( app(tl(esk49_0),nil) = tl(esk49_0)
    | esk48_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>,c_0_183]),c_0_59]),c_0_51])])).

<A NAME="c_0_193"></A>cnf(c_0_193,plain,
    ( cons(hd(app(X1,X2)),app(tl(X1),X2)) = app(X1,X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_122]),c_0_56]),<A HREF="#c_0_123">c_0_123</A>])).

<A NAME="c_0_194"></A>cnf(c_0_194,negated_conjecture,
    ( app(nil,esk48_0) = nil
    | ssItem(hd(esk49_0))
    | esk49_0 != nil
    | ~ ssList(app(nil,esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>,c_0_184]),c_0_59])])).

<A NAME="c_0_195"></A>cnf(c_0_195,negated_conjecture,
    ( esk48_0 = nil
    | esk49_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_185">c_0_185</A>,c_0_186]),c_0_74]),c_0_51])])).

<A NAME="c_0_196"></A>fof(c_0_196,plain,(
    ! [X5,X6,X7,X8] :
      ( ~ ssList(X5)
      | ~ ssList(X6)
      | ~ ssList(X7)
      | ~ ssList(X8)
      | ~ segmentP(X5,X6)
      | segmentP(app(app(X7,X5),X8),X6) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])).

<A NAME="c_0_197"></A>cnf(c_0_197,plain,
    ( segmentP(nil,X1)
    | ~ ssList(X1)
    | nil != X1 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_187">c_0_187</A>])).

<A NAME="c_0_198"></A>cnf(c_0_198,negated_conjecture,
    ( segmentP(esk49_0,esk49_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_188">c_0_188</A>,c_0_99]),c_0_52]),c_0_74])])).

<A NAME="c_0_199"></A>cnf(c_0_199,plain,
    ( nil = X1
    | app(X2,app(X3,X1)) != nil
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_168">c_0_168</A>,c_0_57]),<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_200"></A>cnf(c_0_200,negated_conjecture,
    ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = esk49_0
    | ssItem(hd(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_184">c_0_184</A>,c_0_119]),c_0_51])])).

<A NAME="c_0_201"></A>cnf(c_0_201,negated_conjecture,
    ( esk45_1(esk49_0) = hd(esk49_0)
    | esk49_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_189">c_0_189</A>,c_0_158]),c_0_52])])).

<A NAME="c_0_202"></A>cnf(c_0_202,negated_conjecture,
    ( X1 = esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))
    | app(X1,esk9_2(esk49_0,esk48_0)) != esk49_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_118">c_0_118</A>,c_0_124]),c_0_59]),c_0_125])])).

<A NAME="c_0_203"></A>cnf(c_0_203,negated_conjecture,
    ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil
    | ssList(tl(esk49_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_171">c_0_171</A>,c_0_190]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_204"></A>cnf(c_0_204,negated_conjecture,
    ( esk48_0 = nil
    | cons(X1,tl(esk49_0)) != nil
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_191">c_0_191</A>,c_0_192]),c_0_74])]),<A HREF="#c_0_181">c_0_181</A>])).

<A NAME="c_0_205"></A>cnf(c_0_205,negated_conjecture,
    ( cons(hd(esk49_0),tl(esk49_0)) = esk49_0
    | esk8_2(esk49_0,esk48_0) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_193">c_0_193</A>,c_0_157]),c_0_85]),c_0_85]),c_0_86]),c_0_60])])).

<A NAME="c_0_206"></A>cnf(c_0_206,negated_conjecture,
    ( esk48_0 = nil
    | ssItem(hd(esk49_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_194">c_0_194</A>,c_0_119]),c_0_51])]),<A HREF="#c_0_195">c_0_195</A>])).

<A NAME="c_0_207"></A>cnf(c_0_207,plain,
    ( segmentP(app(app(X1,X2),X3),X4)
    | ~ segmentP(X2,X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X4)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_196">c_0_196</A>])).

<A NAME="c_0_208"></A>cnf(c_0_208,plain,
    ( segmentP(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_197]),c_0_74])])).

<A NAME="c_0_209"></A>fof(c_0_209,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])).

<A NAME="c_0_210"></A>cnf(c_0_210,negated_conjecture,
    ( frontsegP(esk49_0,esk49_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_126]),c_0_74]),c_0_52])])).

<A NAME="c_0_211"></A>cnf(c_0_211,plain,
    ( frontsegP(app(app(X1,X2),X3),X1)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_113">c_0_113</A>,c_0_61]),<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_212"></A>cnf(c_0_212,negated_conjecture,
    ( app(app(esk8_2(esk49_0,esk49_0),esk49_0),esk9_2(esk49_0,esk49_0)) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_198]),c_0_52])])).

<A NAME="c_0_213"></A>cnf(c_0_213,negated_conjecture,
    ( ssList(esk8_2(esk49_0,esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_198]),c_0_52])])).

<A NAME="c_0_214"></A>cnf(c_0_214,negated_conjecture,
    ( ssList(esk9_2(esk49_0,esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_198]),c_0_52])])).

<A NAME="c_0_215"></A>cnf(c_0_215,negated_conjecture,
    ( esk6_2(esk49_0,esk8_2(esk49_0,esk48_0)) = nil
    | app(X1,esk49_0) != nil
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_199">c_0_199</A>,c_0_85]),c_0_86]),c_0_60])])).

<A NAME="c_0_216"></A>cnf(c_0_216,negated_conjecture,
    ( app(esk8_2(esk49_0,esk48_0),esk49_0) = esk49_0
    | ssItem(hd(esk49_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,<A HREF="#c_0_200">c_0_200</A>])).

<A NAME="c_0_217"></A>cnf(c_0_217,negated_conjecture,
    ( esk49_0 = nil
    | ssItem(hd(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_165">c_0_165</A>,c_0_201]),c_0_52])])).

<A NAME="c_0_218"></A>cnf(c_0_218,negated_conjecture,
    ( esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)) = esk48_0
    | ssList(tl(esk49_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_202">c_0_202</A>,c_0_203]),c_0_51])]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_219"></A>cnf(c_0_219,negated_conjecture,
    ( esk8_2(esk49_0,esk48_0) = nil
    | esk48_0 = nil ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_204">c_0_204</A>,c_0_205]),c_0_206]),<A HREF="#c_0_195">c_0_195</A>])).

<A NAME="c_0_220"></A>cnf(c_0_220,plain,
    ( segmentP(app(app(X1,nil),X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_207">c_0_207</A>,c_0_208]),c_0_74])])).

<A NAME="c_0_221"></A>cnf(c_0_221,plain,
    ( frontsegP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_209">c_0_209</A>])).

<A NAME="c_0_222"></A>cnf(c_0_222,negated_conjecture,
    ( app(esk49_0,esk6_2(esk49_0,esk49_0)) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_210]),c_0_52])])).

<A NAME="c_0_223"></A>cnf(c_0_223,negated_conjecture,
    ( ssList(esk6_2(esk49_0,esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_210]),c_0_52])])).

<A NAME="c_0_224"></A>cnf(c_0_224,negated_conjecture,
    ( frontsegP(esk49_0,esk8_2(esk49_0,esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_211">c_0_211</A>,c_0_212]),c_0_213]),c_0_214]),c_0_52])])).

<A NAME="c_0_225"></A>cnf(c_0_225,negated_conjecture,
    ( esk6_2(esk49_0,esk8_2(esk49_0,esk48_0)) = nil
    | ssItem(hd(esk49_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_215">c_0_215</A>,c_0_216]),c_0_60])]),<A HREF="#c_0_217">c_0_217</A>])).

<A NAME="c_0_226"></A>cnf(c_0_226,negated_conjecture,
    ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = esk49_0
    | ssList(tl(esk49_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_124">c_0_124</A>,<A HREF="#c_0_218">c_0_218</A>])).

<A NAME="c_0_227"></A>cnf(c_0_227,negated_conjecture,
    ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil
    | esk48_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_171">c_0_171</A>,c_0_219]),<A HREF="#c_0_195">c_0_195</A>])).

<A NAME="c_0_228"></A>cnf(c_0_228,negated_conjecture,
    ( segmentP(app(esk49_0,X1),nil)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_220">c_0_220</A>,c_0_126]),c_0_52])])).

<A NAME="c_0_229"></A>cnf(c_0_229,plain,
    ( app(X1,esk6_2(X1,X1)) = X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,<A HREF="#c_0_221">c_0_221</A>])).

<A NAME="c_0_230"></A>cnf(c_0_230,negated_conjecture,
    ( esk6_2(esk49_0,esk49_0) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_133">c_0_133</A>,c_0_222]),c_0_52]),c_0_223])])).

<A NAME="c_0_231"></A>fof(c_0_231,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | rearsegP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])).

<A NAME="c_0_232"></A>fof(c_0_232,plain,(
    ! [X2] :
      ( ( ~ rearsegP(nil,X2)
        | nil = X2
        | ~ ssList(X2) )
      & ( nil != X2
        | rearsegP(nil,X2)
        | ~ ssList(X2) ) ) ),
    inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).

<A NAME="c_0_233"></A>cnf(c_0_233,negated_conjecture,
    ( app(esk8_2(esk49_0,esk49_0),app(esk49_0,esk9_2(esk49_0,esk49_0))) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,c_0_212]),c_0_214]),c_0_52]),c_0_213])])).

<A NAME="c_0_234"></A>cnf(c_0_234,negated_conjecture,
    ( app(esk8_2(esk49_0,esk49_0),esk6_2(esk49_0,esk8_2(esk49_0,esk49_0))) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_224]),c_0_213]),c_0_52])])).

<A NAME="c_0_235"></A>cnf(c_0_235,negated_conjecture,
    ( ssList(esk6_2(esk49_0,esk8_2(esk49_0,esk49_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_224]),c_0_213]),c_0_52])])).

<A NAME="c_0_236"></A>cnf(c_0_236,negated_conjecture,
    ( app(esk8_2(esk49_0,esk48_0),nil) = esk49_0
    | ssItem(hd(esk49_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,<A HREF="#c_0_225">c_0_225</A>])).

<A NAME="c_0_237"></A>cnf(c_0_237,negated_conjecture,
    ( app(esk8_2(esk49_0,esk48_0),esk49_0) = esk49_0
    | ssList(tl(esk49_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,<A HREF="#c_0_226">c_0_226</A>])).

<A NAME="c_0_238"></A>cnf(c_0_238,negated_conjecture,
    ( esk48_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>,c_0_227]),c_0_59]),c_0_51])])).

<A NAME="c_0_239"></A>cnf(c_0_239,negated_conjecture,
    ( segmentP(esk49_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_228">c_0_228</A>,c_0_229]),c_0_230]),c_0_74]),c_0_52])])).

<A NAME="c_0_240"></A>cnf(c_0_240,plain,
    ( rearsegP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_231">c_0_231</A>])).

<A NAME="c_0_241"></A>cnf(c_0_241,plain,
    ( rearsegP(nil,X1)
    | ~ ssList(X1)
    | nil != X1 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_232">c_0_232</A>])).

<A NAME="c_0_242"></A>cnf(c_0_242,negated_conjecture,
    ( segmentP(app(esk49_0,X1),esk9_2(esk49_0,esk48_0))
    | ~ ssList(app(esk49_0,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_177">c_0_177</A>,c_0_124]),c_0_125]),c_0_59])])).

<A NAME="c_0_243"></A>cnf(c_0_243,negated_conjecture,
    ( app(esk49_0,esk9_2(esk49_0,esk49_0)) = nil
    | esk8_2(esk49_0,esk49_0) != esk49_0
    | ~ ssList(app(esk49_0,esk9_2(esk49_0,esk49_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_133">c_0_133</A>,c_0_233]),c_0_213])])).

<A NAME="c_0_244"></A>cnf(c_0_244,negated_conjecture,
    ( hd(esk8_2(esk49_0,esk49_0)) = hd(esk49_0)
    | esk8_2(esk49_0,esk49_0) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>,c_0_234]),c_0_235]),c_0_213])])).

<A NAME="c_0_245"></A>cnf(c_0_245,negated_conjecture,
    ( frontsegP(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,c_0_124]),c_0_59]),c_0_125])])).

<A NAME="c_0_246"></A>cnf(c_0_246,negated_conjecture,
    ( esk8_2(esk49_0,esk48_0) = esk49_0
    | ssItem(hd(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_236]),c_0_60])])).

<A NAME="c_0_247"></A>cnf(c_0_247,negated_conjecture,
    ( esk6_2(esk49_0,esk8_2(esk49_0,esk49_0)) = nil
    | app(X1,esk49_0) != nil
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_199">c_0_199</A>,c_0_234]),c_0_235]),c_0_213])])).

<A NAME="c_0_248"></A>cnf(c_0_248,negated_conjecture,
    ( app(esk8_2(esk49_0,nil),esk49_0) = esk49_0
    | ssList(tl(esk49_0)) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_237">c_0_237</A>,<A HREF="#c_0_238">c_0_238</A>])).

<A NAME="c_0_249"></A>cnf(c_0_249,negated_conjecture,
    ( ssList(esk8_2(esk49_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_239]),c_0_74]),c_0_52])])).

<A NAME="c_0_250"></A>cnf(c_0_250,plain,
    ( app(esk7_2(X1,X1),X1) = X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,<A HREF="#c_0_240">c_0_240</A>])).

<A NAME="c_0_251"></A>cnf(c_0_251,plain,
    ( ssList(esk7_2(X1,X1))
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,<A HREF="#c_0_240">c_0_240</A>])).

<A NAME="c_0_252"></A>cnf(c_0_252,plain,
    ( rearsegP(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_241]),c_0_74])])).

<A NAME="c_0_253"></A>cnf(c_0_253,plain,
    ( rearsegP(cons(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_138]),<A HREF="#c_0_139">c_0_139</A>])).

<A NAME="c_0_254"></A>cnf(c_0_254,negated_conjecture,
    ( segmentP(app(esk49_0,X1),esk9_2(esk49_0,esk48_0))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_242">c_0_242</A>,c_0_56]),c_0_52])])).

<A NAME="c_0_255"></A>cnf(c_0_255,negated_conjecture,
    ( app(esk49_0,esk9_2(esk49_0,esk49_0)) = nil
    | esk8_2(esk49_0,esk49_0) != esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_243">c_0_243</A>,c_0_56]),c_0_214]),c_0_52])])).

<A NAME="c_0_256"></A>cnf(c_0_256,negated_conjecture,
    ( esk8_2(esk49_0,esk49_0) = nil
    | ssItem(hd(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_158">c_0_158</A>,c_0_244]),c_0_213])])).

<A NAME="c_0_257"></A>cnf(c_0_257,negated_conjecture,
    ( X1 = esk9_2(esk49_0,esk48_0)
    | app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),X1) != esk49_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_120">c_0_120</A>,c_0_124]),c_0_125]),c_0_59])])).

<A NAME="c_0_258"></A>cnf(c_0_258,negated_conjecture,
    ( app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)))) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_245]),c_0_125]),c_0_52])])).

<A NAME="c_0_259"></A>cnf(c_0_259,negated_conjecture,
    ( ssList(esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_245]),c_0_125]),c_0_52])])).

<A NAME="c_0_260"></A>cnf(c_0_260,negated_conjecture,
    ( app(esk48_0,esk9_2(esk49_0,esk48_0)) = nil
    | ssItem(hd(esk49_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_171">c_0_171</A>,<A HREF="#c_0_246">c_0_246</A>])).

<A NAME="c_0_261"></A>cnf(c_0_261,negated_conjecture,
    ( esk8_2(esk49_0,esk49_0) = nil
    | esk6_2(esk49_0,esk8_2(esk49_0,esk49_0)) != esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_132">c_0_132</A>,c_0_234]),c_0_235]),c_0_213])])).

<A NAME="c_0_262"></A>cnf(c_0_262,negated_conjecture,
    ( esk6_2(esk49_0,esk8_2(esk49_0,esk49_0)) = nil
    | ssList(tl(esk49_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_247">c_0_247</A>,c_0_248]),c_0_249])]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_263"></A>cnf(c_0_263,plain,
    ( X1 = esk7_2(X2,X2)
    | app(X1,X2) != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_118">c_0_118</A>,c_0_250]),<A HREF="#c_0_251">c_0_251</A>])).

<A NAME="c_0_264"></A>cnf(c_0_264,plain,
    ( app(esk7_2(nil,nil),nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_252]),c_0_74])])).

<A NAME="c_0_265"></A>cnf(c_0_265,plain,
    ( ssList(esk7_2(nil,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_252]),c_0_74])])).

<A NAME="c_0_266"></A>cnf(c_0_266,negated_conjecture,
    ( esk49_0 = nil
    | rearsegP(esk49_0,tl(esk49_0))
    | ~ ssItem(hd(esk49_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_253">c_0_253</A>,c_0_152]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_267"></A>cnf(c_0_267,negated_conjecture,
    ( nil != esk48_0
    | ~ segmentP(esk48_0,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_268"></A>cnf(c_0_268,negated_conjecture,
    ( segmentP(app(esk49_0,X1),esk9_2(esk49_0,nil))
    | ~ ssList(X1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_254">c_0_254</A>,<A HREF="#c_0_238">c_0_238</A>])).

<A NAME="c_0_269"></A>cnf(c_0_269,negated_conjecture,
    ( app(esk49_0,esk9_2(esk49_0,esk49_0)) = nil
    | ssItem(hd(esk49_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_255">c_0_255</A>,c_0_256]),<A HREF="#c_0_217">c_0_217</A>])).

<A NAME="c_0_270"></A>cnf(c_0_270,negated_conjecture,
    ( segmentP(esk49_0,esk9_2(esk49_0,esk48_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_242">c_0_242</A>,c_0_99]),c_0_52]),c_0_74])])).

<A NAME="c_0_271"></A>cnf(c_0_271,negated_conjecture,
    ( esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))) = esk9_2(esk49_0,esk48_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_257">c_0_257</A>,c_0_258]),c_0_259])])).

<A NAME="c_0_272"></A>cnf(c_0_272,negated_conjecture,
    ( esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)) = esk48_0
    | ssItem(hd(esk49_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_202">c_0_202</A>,c_0_260]),c_0_51])]),<A HREF="#c_0_217">c_0_217</A>])).

<A NAME="c_0_273"></A>cnf(c_0_273,negated_conjecture,
    ( esk8_2(esk49_0,esk49_0) = nil
    | ssList(tl(esk49_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_261">c_0_261</A>,c_0_262]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_274"></A>cnf(c_0_274,plain,
    ( app(X1,app(nil,X2)) = app(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,c_0_99]),c_0_74])])).

<A NAME="c_0_275"></A>cnf(c_0_275,negated_conjecture,
    ( app(esk8_2(esk49_0,nil),app(nil,esk9_2(esk49_0,nil))) = esk49_0 ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_238]),c_0_238]),<A HREF="#c_0_238">c_0_238</A>])).

<A NAME="c_0_276"></A>cnf(c_0_276,negated_conjecture,
    ( ssList(esk9_2(esk49_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_239]),c_0_74]),c_0_52])])).

<A NAME="c_0_277"></A>cnf(c_0_277,negated_conjecture,
    ( esk8_2(esk49_0,esk48_0) = esk7_2(esk48_0,esk48_0)
    | ssList(tl(esk49_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_263">c_0_263</A>,c_0_169]),c_0_51]),c_0_60])]),<A HREF="#c_0_181">c_0_181</A>])).

<A NAME="c_0_278"></A>cnf(c_0_278,plain,
    ( esk7_2(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>,c_0_264]),c_0_74]),c_0_265])])).

<A NAME="c_0_279"></A>cnf(c_0_279,negated_conjecture,
    ( esk49_0 = nil
    | rearsegP(esk49_0,tl(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_266">c_0_266</A>,c_0_158]),c_0_52])])).

<A NAME="c_0_280"></A>cnf(c_0_280,plain,
    ( app(esk7_2(X1,nil),nil) = X1
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_73]),c_0_74])])).

<A NAME="c_0_281"></A>cnf(c_0_281,plain,
    ( ssList(esk7_2(X1,nil))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_73]),c_0_74])])).

<A NAME="c_0_282"></A>cnf(c_0_282,negated_conjecture,
    ( ~ segmentP(nil,X1)
    | ~ segmentP(esk49_0,X1)
    | ~ ssList(X1)
    | ~ neq(X1,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_267">c_0_267</A>,c_0_238]),c_0_238])])).

<A NAME="c_0_283"></A>cnf(c_0_283,negated_conjecture,
    ( segmentP(nil,esk9_2(esk49_0,nil))
    | ssItem(hd(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_268">c_0_268</A>,c_0_269]),c_0_214])])).

<A NAME="c_0_284"></A>cnf(c_0_284,negated_conjecture,
    ( segmentP(esk49_0,esk9_2(esk49_0,nil)) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_270">c_0_270</A>,<A HREF="#c_0_238">c_0_238</A>])).

<A NAME="c_0_285"></A>cnf(c_0_285,negated_conjecture,
    ( esk9_2(esk49_0,esk48_0) = esk6_2(esk49_0,esk48_0)
    | ssItem(hd(esk49_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_271">c_0_271</A>,<A HREF="#c_0_272">c_0_272</A>])).

<A NAME="c_0_286"></A>cnf(c_0_286,negated_conjecture,
    ( neq(esk49_0,nil)
    | nil != esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_287"></A>cnf(c_0_287,negated_conjecture,
    ( app(esk49_0,esk9_2(esk49_0,esk49_0)) = nil
    | ssList(tl(esk49_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_255">c_0_255</A>,c_0_273]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_288"></A>cnf(c_0_288,negated_conjecture,
    ( app(nil,esk9_2(esk49_0,esk48_0)) = esk49_0
    | ssList(tl(esk49_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_124">c_0_124</A>,<A HREF="#c_0_143">c_0_143</A>])).

<A NAME="c_0_289"></A>cnf(c_0_289,plain,
    ( nil = X1
    | rearsegP(app(X1,X2),tl(app(X1,X2)))
    | ~ ssList(tl(app(X1,X2)))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_253">c_0_253</A>,c_0_140]),<A HREF="#c_0_158">c_0_158</A>])).

<A NAME="c_0_290"></A>cnf(c_0_290,negated_conjecture,
    ( app(esk8_2(esk49_0,nil),esk9_2(esk49_0,nil)) = esk49_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_274">c_0_274</A>,c_0_275]),c_0_276]),c_0_249])])).

<A NAME="c_0_291"></A>cnf(c_0_291,negated_conjecture,
    ( esk8_2(esk49_0,nil) = nil
    | ssList(tl(esk49_0)) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_277">c_0_277</A>,c_0_238]),c_0_238]),c_0_238]),<A HREF="#c_0_278">c_0_278</A>])).

<A NAME="c_0_292"></A>cnf(c_0_292,plain,
    ( X1 = cons(X2,nil)
    | app(X1,X3) != cons(X2,X3)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_118">c_0_118</A>,c_0_138]),<A HREF="#c_0_139">c_0_139</A>])).

<A NAME="c_0_293"></A>cnf(c_0_293,negated_conjecture,
    ( app(esk7_2(esk49_0,tl(esk49_0)),tl(esk49_0)) = esk49_0
    | esk49_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_279]),c_0_52])]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_294"></A>cnf(c_0_294,negated_conjecture,
    ( esk49_0 = nil
    | ssList(esk7_2(esk49_0,tl(esk49_0))) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_279]),c_0_52])]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_295"></A>cnf(c_0_295,plain,
    ( app(tl(esk7_2(X1,nil)),nil) = tl(X1)
    | esk7_2(X1,nil) = nil
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_122">c_0_122</A>,c_0_280]),c_0_74])]),<A HREF="#c_0_281">c_0_281</A>])).

<A NAME="c_0_296"></A>cnf(c_0_296,negated_conjecture,
    ( ssItem(hd(esk49_0))
    | ~ neq(esk9_2(esk49_0,nil),nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_282">c_0_282</A>,c_0_283]),c_0_284]),c_0_276])])).

<A NAME="c_0_297"></A>cnf(c_0_297,negated_conjecture,
    ( esk9_2(esk49_0,nil) = esk49_0
    | ssItem(hd(esk49_0)) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_285">c_0_285</A>,c_0_238]),c_0_238]),<A HREF="#c_0_162">c_0_162</A>])).

<A NAME="c_0_298"></A>cnf(c_0_298,negated_conjecture,
    ( neq(esk49_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_286">c_0_286</A>,c_0_238])])).

<A NAME="c_0_299"></A>cnf(c_0_299,negated_conjecture,
    ( esk8_2(esk49_0,esk48_0) = nil
    | esk49_0 != nil
    | ~ ssList(app(esk48_0,esk9_2(esk49_0,esk48_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>,c_0_62]),c_0_60])])).

<A NAME="c_0_300"></A>cnf(c_0_300,negated_conjecture,
    ( segmentP(nil,esk9_2(esk49_0,nil))
    | ssList(tl(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_268">c_0_268</A>,c_0_287]),c_0_214])])).

<A NAME="c_0_301"></A>cnf(c_0_301,negated_conjecture,
    ( esk9_2(esk49_0,esk48_0) = esk49_0
    | ssList(tl(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_119">c_0_119</A>,c_0_288]),c_0_59])])).

<A NAME="c_0_302"></A>cnf(c_0_302,negated_conjecture,
    ( esk8_2(esk49_0,nil) = nil
    | rearsegP(esk49_0,tl(esk49_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_289">c_0_289</A>,c_0_290]),c_0_276]),c_0_249])]),<A HREF="#c_0_291">c_0_291</A>])).

<A NAME="c_0_303"></A>cnf(c_0_303,negated_conjecture,
    ( esk7_2(esk49_0,tl(esk49_0)) = cons(X1,nil)
    | esk49_0 = nil
    | cons(X1,tl(esk49_0)) != esk49_0
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_292">c_0_292</A>,c_0_293]),c_0_294]),<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_304"></A>cnf(c_0_304,plain,
    ( cons(hd(app(esk7_2(X1,nil),nil)),tl(X1)) = app(esk7_2(X1,nil),nil)
    | esk7_2(X1,nil) = nil
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_193">c_0_193</A>,c_0_295]),c_0_74])]),<A HREF="#c_0_281">c_0_281</A>])).

<A NAME="c_0_305"></A>cnf(c_0_305,negated_conjecture,
    ( ssItem(hd(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_296">c_0_296</A>,c_0_297]),c_0_298])])).

<A NAME="c_0_306"></A>cnf(c_0_306,negated_conjecture,
    ( esk8_2(esk49_0,esk48_0) = nil
    | esk49_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_299">c_0_299</A>,c_0_56]),c_0_59]),c_0_51])])).

<A NAME="c_0_307"></A>cnf(c_0_307,negated_conjecture,
    ( ~ totalorderedP(X1)
    | ~ segmentP(X1,esk50_0)
    | ~ segmentP(esk51_0,X1)
    | ~ neq(esk50_0,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_308"></A>fof(c_0_308,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | segmentP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])).

<A NAME="c_0_309"></A>cnf(c_0_309,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_207">c_0_207</A>,<A HREF="#c_0_175">c_0_175</A>])).

<A NAME="c_0_310"></A>cnf(c_0_310,negated_conjecture,
    ( ssList(tl(esk49_0))
    | ~ neq(esk9_2(esk49_0,nil),nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_282">c_0_282</A>,c_0_300]),c_0_284]),c_0_276])])).

<A NAME="c_0_311"></A>cnf(c_0_311,negated_conjecture,
    ( esk9_2(esk49_0,nil) = esk49_0
    | ssList(tl(esk49_0)) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_301">c_0_301</A>,<A HREF="#c_0_238">c_0_238</A>])).

<A NAME="c_0_312"></A>cnf(c_0_312,negated_conjecture,
    ( esk8_2(esk49_0,nil) = nil
    | ssList(esk7_2(esk49_0,tl(esk49_0))) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_302]),c_0_52])]),<A HREF="#c_0_291">c_0_291</A>])).

<A NAME="c_0_313"></A>cnf(c_0_313,negated_conjecture,
    ( esk7_2(esk49_0,tl(esk49_0)) = cons(hd(esk49_0),nil)
    | esk49_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_303">c_0_303</A>,c_0_304]),c_0_112]),c_0_126]),c_0_112]),c_0_112]),c_0_126]),c_0_112]),c_0_126]),c_0_305]),c_0_52])])).

<A NAME="c_0_314"></A>cnf(c_0_314,negated_conjecture,
    ( esk8_2(esk49_0,nil) = nil
    | esk49_0 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_306">c_0_306</A>,<A HREF="#c_0_238">c_0_238</A>])).

<A NAME="c_0_315"></A>cnf(c_0_315,negated_conjecture,
    ( ~ totalorderedP(X1)
    | ~ segmentP(X1,esk48_0)
    | ~ segmentP(esk49_0,X1)
    | ~ ssList(X1)
    | ~ neq(esk48_0,X1) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_307">c_0_307</A>,c_0_45]),c_0_45]),<A HREF="#c_0_44">c_0_44</A>])).

<A NAME="c_0_316"></A>cnf(c_0_316,plain,
    ( segmentP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_308">c_0_308</A>])).

<A NAME="c_0_317"></A>cnf(c_0_317,plain,
    ( segmentP(app(X1,X2),X1)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_309">c_0_309</A>,c_0_250]),<A HREF="#c_0_251">c_0_251</A>])).

<A NAME="c_0_318"></A>cnf(c_0_318,negated_conjecture,
    ( ssList(tl(esk49_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_310">c_0_310</A>,c_0_311]),c_0_298])])).

<A NAME="c_0_319"></A>cnf(c_0_319,negated_conjecture,
    ( app(nil,esk9_2(esk49_0,nil)) = nil
    | esk8_2(esk49_0,nil) != esk49_0 ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_171">c_0_171</A>,c_0_238]),c_0_238]),<A HREF="#c_0_238">c_0_238</A>])).

<A NAME="c_0_320"></A>cnf(c_0_320,negated_conjecture,
    ( esk8_2(esk49_0,nil) = nil
    | ssList(cons(hd(esk49_0),nil)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_312">c_0_312</A>,c_0_313]),<A HREF="#c_0_314">c_0_314</A>])).

<A NAME="c_0_321"></A>cnf(c_0_321,negated_conjecture,
    ( esk49_0 = nil
    | ssList(cons(hd(esk49_0),nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_294">c_0_294</A>,<A HREF="#c_0_313">c_0_313</A>])).

<A NAME="c_0_322"></A>cnf(c_0_322,negated_conjecture,
    ( X1 = esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))
    | app(esk8_2(esk49_0,esk48_0),X1) != esk49_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_120">c_0_120</A>,c_0_85]),c_0_60]),c_0_86])])).

<A NAME="c_0_323"></A>cnf(c_0_323,negated_conjecture,
    ( ~ totalorderedP(X1)
    | ~ segmentP(esk49_0,X1)
    | ~ ssList(X1)
    | ~ neq(nil,X1) ),
    inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_315">c_0_315</A>,c_0_238]),c_0_238]),<A HREF="#c_0_316">c_0_316</A>])).

<A NAME="c_0_324"></A>cnf(c_0_324,negated_conjecture,
    ( esk49_0 = nil
    | segmentP(esk49_0,esk7_2(esk49_0,tl(esk49_0))) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_317">c_0_317</A>,c_0_293]),c_0_318])]),<A HREF="#c_0_294">c_0_294</A>])).

<A NAME="c_0_325"></A>cnf(c_0_325,plain,
    ( segmentP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_309">c_0_309</A>,c_0_99]),c_0_74])]),<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_326"></A>cnf(c_0_326,negated_conjecture,
    ( app(nil,esk9_2(esk49_0,nil)) = nil
    | ssList(cons(hd(esk49_0),nil)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_319">c_0_319</A>,c_0_320]),<A HREF="#c_0_321">c_0_321</A>])).

<A NAME="c_0_327"></A>cnf(c_0_327,negated_conjecture,
    ( X1 = esk6_2(esk49_0,esk8_2(esk49_0,nil))
    | app(esk8_2(esk49_0,nil),X1) != esk49_0
    | ~ ssList(X1) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_322">c_0_322</A>,c_0_238]),<A HREF="#c_0_238">c_0_238</A>])).

<A NAME="c_0_328"></A>cnf(c_0_328,negated_conjecture,
    ( esk49_0 = nil
    | ~ totalorderedP(esk7_2(esk49_0,tl(esk49_0)))
    | ~ neq(nil,esk7_2(esk49_0,tl(esk49_0))) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_323">c_0_323</A>,c_0_324]),<A HREF="#c_0_294">c_0_294</A>])).

<A NAME="c_0_329"></A>fof(c_0_329,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | totalorderedP(cons(X2,nil)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])).

<A NAME="c_0_330"></A>cnf(c_0_330,negated_conjecture,
    ( segmentP(nil,esk9_2(esk49_0,nil))
    | ssList(cons(hd(esk49_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_325">c_0_325</A>,c_0_326]),c_0_276]),c_0_74])])).

<A NAME="c_0_331"></A>cnf(c_0_331,negated_conjecture,
    ( esk6_2(esk49_0,esk8_2(esk49_0,nil)) = esk9_2(esk49_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_327">c_0_327</A>,c_0_290]),c_0_276])])).

<A NAME="c_0_332"></A>cnf(c_0_332,negated_conjecture,
    ( esk49_0 = nil
    | ~ totalorderedP(cons(hd(esk49_0),nil))
    | ~ neq(nil,cons(hd(esk49_0),nil)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_328">c_0_328</A>,<A HREF="#c_0_313">c_0_313</A>])).

<A NAME="c_0_333"></A>cnf(c_0_333,plain,
    ( totalorderedP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_329">c_0_329</A>])).

<A NAME="c_0_334"></A>cnf(c_0_334,negated_conjecture,
    ( ssList(cons(hd(esk49_0),nil))
    | ~ neq(esk9_2(esk49_0,nil),nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_282">c_0_282</A>,c_0_330]),c_0_284]),c_0_276])])).

<A NAME="c_0_335"></A>cnf(c_0_335,negated_conjecture,
    ( esk9_2(esk49_0,nil) = esk49_0
    | ssList(cons(hd(esk49_0),nil)) ),
    inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_331">c_0_331</A>,c_0_320]),<A HREF="#c_0_162">c_0_162</A>])).

<A NAME="c_0_336"></A>cnf(c_0_336,negated_conjecture,
    ( totalorderedP(esk50_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_337"></A>fof(c_0_337,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) != X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_39])])])).

<A NAME="c_0_338"></A>cnf(c_0_338,negated_conjecture,
    ( esk49_0 = nil
    | ~ neq(nil,cons(hd(esk49_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_332">c_0_332</A>,c_0_333]),c_0_305])])).

<A NAME="c_0_339"></A>cnf(c_0_339,negated_conjecture,
    ( ssList(cons(hd(esk49_0),nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_334">c_0_334</A>,c_0_335]),c_0_298])])).

<A NAME="c_0_340"></A>cnf(c_0_340,negated_conjecture,
    ( totalorderedP(esk48_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_336">c_0_336</A>,<A HREF="#c_0_45">c_0_45</A>])).

<A NAME="c_0_341"></A>cnf(c_0_341,plain,
    ( cons(X1,X2) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_337">c_0_337</A>])).

<A NAME="c_0_342"></A>cnf(c_0_342,negated_conjecture,
    ( cons(hd(esk49_0),nil) = nil
    | esk49_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_338">c_0_338</A>,c_0_186]),c_0_339]),c_0_74])])).

<A NAME="c_0_343"></A>cnf(c_0_343,negated_conjecture,
    ( ~ neq(esk48_0,esk48_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_315">c_0_315</A>,c_0_175]),c_0_340]),c_0_50]),c_0_51])])).

<A NAME="c_0_344"></A>cnf(c_0_344,negated_conjecture,
    ( esk49_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_341">c_0_341</A>,c_0_342]),c_0_74]),c_0_305])])).

<A NAME="c_0_345"></A>cnf(c_0_345,negated_conjecture,
    ( ~ neq(nil,nil) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_343">c_0_343</A>,c_0_238]),<A HREF="#c_0_238">c_0_238</A>])).

<A NAME="c_0_346"></A>cnf(c_0_346,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_298">c_0_298</A>,c_0_344]),<A HREF="#c_0_345">c_0_345</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC066+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n128.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:35:07 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 49.35/49.55  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 49.35/49.55  # and selection function SelectCQPrecWNTNp.
% 49.35/49.55  #
% 49.35/49.55  # Preprocessing time       : 0.014 s
% 49.35/49.55  # Presaturation interreduction done
% 49.35/49.55  
% 49.35/49.55  # Proof found!
% 49.35/49.55  # SZS status Theorem
% 49.35/49.55  # SZS output start CNFRefutation
% 49.35/49.55  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~(segmentP(X4,X3)))|~(totalorderedP(X3)))|?[X5]:((((ssList(X5)&neq(X3,X5))&segmentP(X4,X5))&segmentP(X5,X3))&totalorderedP(X5)))|((nil!=X2|nil=X1)&(~(neq(X2,nil))|?[X6]:(((ssList(X6)&neq(X6,nil))&segmentP(X2,X6))&segmentP(X1,X6)))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 49.35/49.55  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax7)).
% 49.35/49.55  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax5)).
% 49.35/49.55  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 49.35/49.55  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 49.35/49.55  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(rearsegP(X1,X2)=>rearsegP(app(X3,X1),X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax50)).
% 49.35/49.55  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>rearsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax51)).
% 49.35/49.55  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax78)).
% 49.35/49.55  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
% 49.35/49.55  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
% 49.35/49.55  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>nil!=cons(X2,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax21)).
% 49.35/49.55  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)&lt;=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax4)).
% 49.35/49.55  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax6)).
% 49.35/49.55  fof(c_0_13, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 49.35/49.55  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
% 49.35/49.55  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(frontsegP(X1,X2)=>frontsegP(app(X1,X3),X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax43)).
% 49.35/49.55  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax45)).
% 49.35/49.55  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax79)).
% 49.35/49.55  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
% 49.35/49.55  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X2,X3)=app(X2,X1)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax80)).
% 49.35/49.55  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
% 49.35/49.55  fof(c_0_21, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax86)).
% 49.35/49.55  fof(c_0_22, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 49.35/49.55  fof(c_0_23, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax23)).
% 49.35/49.55  fof(c_0_24, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>hd(app(X1,X2))=hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax85)).
% 49.35/49.55  fof(c_0_25, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
% 49.35/49.55  fof(c_0_26, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax22)).
% 49.35/49.55  fof(c_0_27, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(frontsegP(cons(X1,X3),cons(X2,X4))&lt;=>(X1=X2&frontsegP(X3,X4)))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax44)).
% 49.35/49.55  fof(c_0_28, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax20)).
% 49.35/49.55  fof(c_0_29, axiom, (![X1]:(ssList(X1)=>segmentP(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax55)).
% 49.35/49.55  fof(c_0_30, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
% 49.35/49.55  fof(c_0_31, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
% 49.35/49.55  fof(c_0_32, axiom, (![X1]:(ssList(X1)=>(segmentP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax58)).
% 49.35/49.55  fof(c_0_33, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(segmentP(X1,X2)=>segmentP(app(app(X3,X1),X4),X2))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax56)).
% 49.35/49.55  fof(c_0_34, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax42)).
% 49.35/49.55  fof(c_0_35, axiom, (![X1]:(ssList(X1)=>rearsegP(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax49)).
% 49.35/49.55  fof(c_0_36, axiom, (![X1]:(ssList(X1)=>(rearsegP(nil,X1)&lt;=>nil=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax52)).
% 49.35/49.55  fof(c_0_37, axiom, (![X1]:(ssList(X1)=>segmentP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax57)).
% 49.35/49.55  fof(c_0_38, axiom, (![X1]:(ssItem(X1)=>totalorderedP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax65)).
% 49.35/49.55  fof(c_0_39, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)!=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax18)).
% 49.35/49.55  fof(c_0_40, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((((X2!=X4|X1!=X3)|~segmentP(X4,X3))|~totalorderedP(X3))|?[X5]:((((ssList(X5)&neq(X3,X5))&segmentP(X4,X5))&segmentP(X5,X3))&totalorderedP(X5)))|((nil!=X2|nil=X1)&(~neq(X2,nil)|?[X6]:(((ssList(X6)&neq(X6,nil))&segmentP(X2,X6))&segmentP(X1,X6))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 49.35/49.55  fof(c_0_41, negated_conjecture, (![X11]:![X12]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((((esk49_0=esk51_0&esk48_0=esk50_0)&segmentP(esk51_0,esk50_0))&totalorderedP(esk50_0))&((((~ssList(X11)|~neq(esk50_0,X11))|~segmentP(esk51_0,X11))|~segmentP(X11,esk50_0))|~totalorderedP(X11)))&(((neq(esk49_0,nil)|nil=esk49_0)&((((~ssList(X12)|~neq(X12,nil))|~segmentP(esk49_0,X12))|~segmentP(esk48_0,X12))|nil=esk49_0))&((neq(esk49_0,nil)|nil!=esk48_0)&((((~ssList(X12)|~neq(X12,nil))|~segmentP(esk49_0,X12))|~segmentP(esk48_0,X12))|nil!=esk48_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])])])])).
% 49.35/49.55  fof(c_0_42, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
% 49.35/49.55  cnf(c_0_43, negated_conjecture, (segmentP(esk51_0,esk50_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 49.35/49.55  cnf(c_0_44, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_41])).
% 49.35/49.55  cnf(c_0_45, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_41])).
% 49.35/49.55  fof(c_0_46, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])])).
% 49.35/49.55  fof(c_0_47, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
% 49.35/49.55  fof(c_0_48, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 49.35/49.55  cnf(c_0_49, plain, (app(app(esk8_2(X1,X2),X2),esk9_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
% 49.35/49.55  cnf(c_0_50, negated_conjecture, (segmentP(esk49_0,esk48_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44]), c_0_45])).
% 49.35/49.55  cnf(c_0_51, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 49.35/49.55  cnf(c_0_52, negated_conjecture, (ssList(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 49.35/49.55  cnf(c_0_53, plain, (ssList(esk9_2(X1,X2))|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
% 49.35/49.55  cnf(c_0_54, plain, (ssList(esk8_2(X1,X2))|~ssList(X1)|~ssList(X2)|~segmentP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
% 49.35/49.55  cnf(c_0_55, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_46])).
% 49.35/49.55  cnf(c_0_56, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
% 49.35/49.55  cnf(c_0_57, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
% 49.35/49.55  cnf(c_0_58, negated_conjecture, (app(app(esk8_2(esk49_0,esk48_0),esk48_0),esk9_2(esk49_0,esk48_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52])])).
% 49.35/49.55  cnf(c_0_59, negated_conjecture, (ssList(esk9_2(esk49_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_50]), c_0_51]), c_0_52])])).
% 49.35/49.55  cnf(c_0_60, negated_conjecture, (ssList(esk8_2(esk49_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_50]), c_0_51]), c_0_52])])).
% 49.35/49.55  cnf(c_0_61, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_55]), c_0_56])).
% 49.35/49.55  cnf(c_0_62, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),app(esk48_0,esk9_2(esk49_0,esk48_0)))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_51]), c_0_60])])).
% 49.35/49.55  fof(c_0_63, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~rearsegP(X4,X5)|rearsegP(app(X6,X4),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
% 49.35/49.55  fof(c_0_64, plain, (![X2]:(~ssList(X2)|rearsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
% 49.35/49.55  cnf(c_0_65, negated_conjecture, (frontsegP(esk49_0,esk8_2(esk49_0,esk48_0))|~ssList(app(esk48_0,esk9_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_60])])).
% 49.35/49.55  fof(c_0_66, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
% 49.35/49.55  fof(c_0_67, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
% 49.35/49.55  fof(c_0_68, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 49.35/49.55  fof(c_0_69, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|nil!=cons(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
% 49.35/49.55  fof(c_0_70, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
% 49.35/49.55  fof(c_0_71, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])])])).
% 49.35/49.55  cnf(c_0_72, plain, (rearsegP(app(X1,X2),X3)|~rearsegP(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
% 49.35/49.55  cnf(c_0_73, plain, (rearsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
% 49.35/49.55  cnf(c_0_74, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_13])).
% 49.35/49.55  cnf(c_0_75, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
% 49.35/49.55  cnf(c_0_76, negated_conjecture, (frontsegP(esk49_0,esk8_2(esk49_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_56]), c_0_59]), c_0_51])])).
% 49.35/49.55  cnf(c_0_77, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
% 49.35/49.55  cnf(c_0_78, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
% 49.35/49.55  cnf(c_0_79, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
% 49.35/49.55  cnf(c_0_80, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
% 49.35/49.55  cnf(c_0_81, plain, (nil!=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
% 49.35/49.55  cnf(c_0_82, plain, (singletonP(X1)|~ssList(X1)|cons(X2,nil)!=X1|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
% 49.35/49.55  cnf(c_0_83, plain, (rearsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X3,X2)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_71])).
% 49.35/49.55  cnf(c_0_84, plain, (rearsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74])])).
% 49.35/49.55  cnf(c_0_85, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk6_2(esk49_0,esk8_2(esk49_0,esk48_0)))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_60]), c_0_52])])).
% 49.35/49.55  cnf(c_0_86, negated_conjecture, (ssList(esk6_2(esk49_0,esk8_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_76]), c_0_60]), c_0_52])])).
% 49.35/49.55  cnf(c_0_87, plain, (cons(hd(cons(X1,X2)),X2)=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80]), c_0_81])).
% 49.35/49.55  cnf(c_0_88, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
% 49.35/49.55  cnf(c_0_89, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
% 49.35/49.55  cnf(c_0_90, plain, (singletonP(cons(X1,nil))|~ssList(cons(X1,nil))|~ssItem(X1)), inference(er,[status(thm)],[c_0_82])).
% 49.35/49.55  cnf(c_0_91, plain, (rearsegP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_83]), c_0_56])).
% 49.35/49.55  fof(c_0_92, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).
% 49.35/49.55  cnf(c_0_93, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
% 49.35/49.55  cnf(c_0_94, negated_conjecture, (rearsegP(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86]), c_0_60])])).
% 49.35/49.55  cnf(c_0_95, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
% 49.35/49.55  cnf(c_0_96, plain, (cons(hd(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_74])]), c_0_89])).
% 49.35/49.55  cnf(c_0_97, plain, (singletonP(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_80]), c_0_74])])).
% 49.35/49.55  cnf(c_0_98, negated_conjecture, (rearsegP(esk49_0,esk9_2(esk49_0,esk48_0))|~ssList(app(esk8_2(esk49_0,esk48_0),esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_58]), c_0_59])])).
% 49.35/49.55  cnf(c_0_99, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
% 49.35/49.55  cnf(c_0_100, negated_conjecture, (app(esk7_2(esk49_0,nil),nil)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_74]), c_0_52])])).
% 49.35/49.55  cnf(c_0_101, negated_conjecture, (ssList(esk7_2(esk49_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_94]), c_0_74]), c_0_52])])).
% 49.35/49.55  fof(c_0_102, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~frontsegP(X4,X5)|frontsegP(app(X4,X6),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
% 49.35/49.55  fof(c_0_103, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])).
% 49.35/49.55  cnf(c_0_104, plain, (cons(hd(cons(X1,nil)),nil)=cons(X1,nil)|~ssList(cons(X1,nil))|~ssItem(X1)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
% 49.35/49.55  fof(c_0_105, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
% 49.35/49.55  fof(c_0_106, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])).
% 49.35/49.55  fof(c_0_107, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X5,X6)!=app(X5,X4)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
% 49.35/49.55  fof(c_0_108, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])).
% 49.35/49.55  fof(c_0_109, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
% 49.35/49.55  fof(c_0_110, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])).
% 49.35/49.55  cnf(c_0_111, negated_conjecture, (rearsegP(esk49_0,esk9_2(esk49_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_56]), c_0_51]), c_0_60])])).
% 49.35/49.55  cnf(c_0_112, negated_conjecture, (esk7_2(esk49_0,nil)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101])])).
% 49.35/49.55  cnf(c_0_113, plain, (frontsegP(app(X1,X2),X3)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
% 49.35/49.55  cnf(c_0_114, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
% 49.35/49.55  cnf(c_0_115, plain, (cons(hd(cons(X1,nil)),nil)=cons(X1,nil)|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_80]), c_0_74])])).
% 49.35/49.55  fof(c_0_116, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).
% 49.35/49.55  fof(c_0_117, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|hd(app(X3,X4))=hd(X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
% 49.35/49.55  cnf(c_0_118, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_105])).
% 49.35/49.55  cnf(c_0_119, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
% 49.35/49.55  cnf(c_0_120, plain, (X1=X2|app(X3,X1)!=app(X3,X2)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
% 49.35/49.55  cnf(c_0_121, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
% 49.35/49.55  cnf(c_0_122, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
% 49.35/49.55  cnf(c_0_123, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_110])).
% 49.35/49.55  cnf(c_0_124, negated_conjecture, (app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),esk9_2(esk49_0,esk48_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_111]), c_0_59]), c_0_52])])).
% 49.35/49.55  cnf(c_0_125, negated_conjecture, (ssList(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_111]), c_0_59]), c_0_52])])).
% 49.35/49.55  cnf(c_0_126, negated_conjecture, (app(esk49_0,nil)=esk49_0), inference(rw,[status(thm)],[c_0_100, c_0_112])).
% 49.35/49.55  cnf(c_0_127, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_74])])).
% 49.35/49.55  fof(c_0_128, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
% 49.35/49.55  cnf(c_0_129, plain, (ssList(cons(X1,nil))|~ssItem(hd(cons(X1,nil)))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_115]), c_0_74])])).
% 49.35/49.55  cnf(c_0_130, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_116])).
% 49.35/49.55  cnf(c_0_131, plain, (hd(app(X1,X2))=hd(X1)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_117])).
% 49.35/49.55  cnf(c_0_132, plain, (X1=nil|app(X1,X2)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_74])])).
% 49.35/49.55  cnf(c_0_133, plain, (X1=nil|app(X2,X1)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_99]), c_0_74])])).
% 49.35/49.55  cnf(c_0_134, plain, (nil=X1|ssList(app(tl(X1),X2))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_56]), c_0_123])).
% 49.35/49.55  cnf(c_0_135, negated_conjecture, (app(tl(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))),esk9_2(esk49_0,esk48_0))=tl(esk49_0)|esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_124]), c_0_59]), c_0_125])])).
% 49.35/49.55  cnf(c_0_136, negated_conjecture, (app(tl(esk49_0),nil)=tl(esk49_0)|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_126]), c_0_74]), c_0_52])])).
% 49.35/49.55  cnf(c_0_137, negated_conjecture, (frontsegP(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_85]), c_0_86]), c_0_60])])).
% 49.35/49.55  cnf(c_0_138, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_128])).
% 49.35/49.55  cnf(c_0_139, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_74])])).
% 49.35/49.55  cnf(c_0_140, plain, (cons(hd(X1),tl(app(X1,X2)))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_131]), c_0_56]), c_0_123])).
% 49.35/49.55  cnf(c_0_141, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk48_0)=nil|esk9_2(esk49_0,esk48_0)!=esk49_0|~ssList(app(esk8_2(esk49_0,esk48_0),esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_58]), c_0_59])])).
% 49.35/49.55  cnf(c_0_142, negated_conjecture, (esk9_2(esk49_0,esk48_0)=nil|esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_124]), c_0_125]), c_0_59])])).
% 49.35/49.55  cnf(c_0_143, negated_conjecture, (esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))=nil|ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_59]), c_0_125])])).
% 49.35/49.55  cnf(c_0_144, negated_conjecture, (esk49_0=nil|ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_136]), c_0_74]), c_0_52])])).
% 49.35/49.55  cnf(c_0_145, plain, (app(tl(X1),nil)=tl(X1)|nil=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_99]), c_0_74])])).
% 49.35/49.55  fof(c_0_146, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])).
% 49.35/49.55  cnf(c_0_147, negated_conjecture, (app(nil,esk6_2(esk49_0,nil))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_137]), c_0_74]), c_0_52])])).
% 49.35/49.55  cnf(c_0_148, negated_conjecture, (ssList(esk6_2(esk49_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_137]), c_0_74]), c_0_52])])).
% 49.35/49.55  fof(c_0_149, plain, (![X5]:![X6]:![X7]:![X8]:(((((((X5=X6|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5))&(((((frontsegP(X7,X8)|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))&((((((X5!=X6|~frontsegP(X7,X8))|frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])).
% 49.35/49.55  fof(c_0_150, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])])).
% 49.35/49.55  cnf(c_0_151, plain, (frontsegP(cons(X1,X2),cons(X1,nil))|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_138]), c_0_139])).
% 49.35/49.55  cnf(c_0_152, negated_conjecture, (cons(hd(esk49_0),tl(esk49_0))=esk49_0|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_126]), c_0_74]), c_0_52])])).
% 49.35/49.55  cnf(c_0_153, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk48_0)=nil|esk9_2(esk49_0,esk48_0)!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_56]), c_0_51]), c_0_60])])).
% 49.35/49.55  cnf(c_0_154, negated_conjecture, (esk9_2(esk49_0,esk48_0)=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143]), c_0_144])).
% 49.35/49.55  cnf(c_0_155, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=nil|esk8_2(esk49_0,esk48_0)!=esk49_0|~ssList(app(esk48_0,esk9_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_62]), c_0_60])])).
% 49.35/49.55  cnf(c_0_156, plain, (app(app(tl(X1),X2),nil)=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_122]), c_0_56]), c_0_123])).
% 49.35/49.55  cnf(c_0_157, negated_conjecture, (app(tl(esk8_2(esk49_0,esk48_0)),esk6_2(esk49_0,esk8_2(esk49_0,esk48_0)))=tl(esk49_0)|esk8_2(esk49_0,esk48_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_85]), c_0_86]), c_0_60])])).
% 49.35/49.55  cnf(c_0_158, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_146])).
% 49.35/49.55  cnf(c_0_159, negated_conjecture, (hd(esk8_2(esk49_0,esk48_0))=hd(esk49_0)|esk8_2(esk49_0,esk48_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_85]), c_0_86]), c_0_60])])).
% 49.35/49.55  fof(c_0_160, plain, (![X2]:(~ssList(X2)|segmentP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])).
% 49.35/49.55  cnf(c_0_161, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
% 49.35/49.55  cnf(c_0_162, negated_conjecture, (esk6_2(esk49_0,nil)=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_147]), c_0_148])])).
% 49.35/49.55  cnf(c_0_163, plain, (X1=X2|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~ssList(X4)|~frontsegP(cons(X1,X3),cons(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_149])).
% 49.35/49.55  cnf(c_0_164, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_150])).
% 49.35/49.55  cnf(c_0_165, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_150])).
% 49.35/49.55  cnf(c_0_166, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_150])).
% 49.35/49.55  cnf(c_0_167, negated_conjecture, (esk49_0=nil|frontsegP(esk49_0,cons(hd(esk49_0),nil))|~ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_144])).
% 49.35/49.55  cnf(c_0_168, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_110])).
% 49.35/49.55  cnf(c_0_169, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk48_0)=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154]), c_0_144])).
% 49.35/49.55  fof(c_0_170, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
% 49.35/49.55  cnf(c_0_171, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=nil|esk8_2(esk49_0,esk48_0)!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_56]), c_0_59]), c_0_51])])).
% 49.35/49.55  cnf(c_0_172, negated_conjecture, (app(tl(esk49_0),nil)=tl(esk49_0)|esk8_2(esk49_0,esk48_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_86]), c_0_60])])).
% 49.35/49.55  cnf(c_0_173, negated_conjecture, (esk8_2(esk49_0,esk48_0)=nil|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_60])])).
% 49.35/49.55  cnf(c_0_174, negated_conjecture, (nil=esk49_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 49.35/49.55  cnf(c_0_175, plain, (segmentP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_160])).
% 49.35/49.55  fof(c_0_176, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])])).
% 49.35/49.55  cnf(c_0_177, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_161])).
% 49.35/49.55  cnf(c_0_178, negated_conjecture, (app(nil,esk49_0)=esk49_0), inference(rw,[status(thm)],[c_0_147, c_0_162])).
% 49.35/49.55  cnf(c_0_179, plain, (esk45_1(X1)=X2|nil=X1|~frontsegP(X1,cons(X2,X3))|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164]), c_0_165]), c_0_166])).
% 49.35/49.55  cnf(c_0_180, negated_conjecture, (esk49_0=nil|frontsegP(esk49_0,cons(hd(esk49_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_158]), c_0_52])])).
% 49.35/49.55  cnf(c_0_181, negated_conjecture, (esk48_0=nil|ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169]), c_0_51]), c_0_60])])).
% 49.35/49.55  cnf(c_0_182, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_170])).
% 49.35/49.55  cnf(c_0_183, negated_conjecture, (app(tl(esk49_0),nil)=tl(esk49_0)|app(esk48_0,esk9_2(esk49_0,esk48_0))=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_136])).
% 49.35/49.55  cnf(c_0_184, negated_conjecture, (app(app(nil,esk48_0),esk9_2(esk49_0,esk48_0))=esk49_0|ssItem(hd(esk49_0))), inference(spm,[status(thm)],[c_0_58, c_0_173])).
% 49.35/49.55  cnf(c_0_185, negated_conjecture, (esk49_0=nil|~neq(esk48_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_175]), c_0_50]), c_0_51])])).
% 49.35/49.55  cnf(c_0_186, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_176])).
% 49.35/49.55  fof(c_0_187, plain, (![X2]:(((~segmentP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|segmentP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])).
% 49.35/49.55  cnf(c_0_188, negated_conjecture, (segmentP(app(esk49_0,X1),esk49_0)|~ssList(app(esk49_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_178]), c_0_74]), c_0_52])])).
% 49.35/49.55  cnf(c_0_189, negated_conjecture, (esk45_1(esk49_0)=hd(esk49_0)|esk49_0=nil|~ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_74]), c_0_52])])).
% 49.35/49.55  cnf(c_0_190, negated_conjecture, (esk8_2(esk49_0,esk48_0)=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_169]), c_0_51]), c_0_60])]), c_0_181])).
% 49.35/49.55  cnf(c_0_191, plain, (cons(X1,app(X2,X3))!=nil|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_182]), c_0_80]), c_0_81])).
% 49.35/49.55  cnf(c_0_192, negated_conjecture, (app(tl(esk49_0),nil)=tl(esk49_0)|esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_183]), c_0_59]), c_0_51])])).
% 49.35/49.55  cnf(c_0_193, plain, (cons(hd(app(X1,X2)),app(tl(X1),X2))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_122]), c_0_56]), c_0_123])).
% 49.35/49.55  cnf(c_0_194, negated_conjecture, (app(nil,esk48_0)=nil|ssItem(hd(esk49_0))|esk49_0!=nil|~ssList(app(nil,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_184]), c_0_59])])).
% 49.35/49.55  cnf(c_0_195, negated_conjecture, (esk48_0=nil|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_186]), c_0_74]), c_0_51])])).
% 49.35/49.55  fof(c_0_196, plain, (![X5]:![X6]:![X7]:![X8]:(~ssList(X5)|(~ssList(X6)|(~ssList(X7)|(~ssList(X8)|(~segmentP(X5,X6)|segmentP(app(app(X7,X5),X8),X6))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])).
% 49.35/49.55  cnf(c_0_197, plain, (segmentP(nil,X1)|~ssList(X1)|nil!=X1), inference(split_conjunct,[status(thm)],[c_0_187])).
% 49.35/49.55  cnf(c_0_198, negated_conjecture, (segmentP(esk49_0,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_99]), c_0_52]), c_0_74])])).
% 49.35/49.55  cnf(c_0_199, plain, (nil=X1|app(X2,app(X3,X1))!=nil|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_57]), c_0_56])).
% 49.35/49.55  cnf(c_0_200, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=esk49_0|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_119]), c_0_51])])).
% 49.35/49.55  cnf(c_0_201, negated_conjecture, (esk45_1(esk49_0)=hd(esk49_0)|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_158]), c_0_52])])).
% 49.35/49.55  cnf(c_0_202, negated_conjecture, (X1=esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))|app(X1,esk9_2(esk49_0,esk48_0))!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_124]), c_0_59]), c_0_125])])).
% 49.35/49.55  cnf(c_0_203, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_190]), c_0_144])).
% 49.35/49.55  cnf(c_0_204, negated_conjecture, (esk48_0=nil|cons(X1,tl(esk49_0))!=nil|~ssItem(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192]), c_0_74])]), c_0_181])).
% 49.35/49.55  cnf(c_0_205, negated_conjecture, (cons(hd(esk49_0),tl(esk49_0))=esk49_0|esk8_2(esk49_0,esk48_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_157]), c_0_85]), c_0_85]), c_0_86]), c_0_60])])).
% 49.35/49.55  cnf(c_0_206, negated_conjecture, (esk48_0=nil|ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_119]), c_0_51])]), c_0_195])).
% 49.35/49.55  cnf(c_0_207, plain, (segmentP(app(app(X1,X2),X3),X4)|~segmentP(X2,X4)|~ssList(X3)|~ssList(X1)|~ssList(X4)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_196])).
% 49.35/49.55  cnf(c_0_208, plain, (segmentP(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_197]), c_0_74])])).
% 49.35/49.55  fof(c_0_209, plain, (![X2]:(~ssList(X2)|frontsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])).
% 49.35/49.55  cnf(c_0_210, negated_conjecture, (frontsegP(esk49_0,esk49_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_126]), c_0_74]), c_0_52])])).
% 49.35/49.55  cnf(c_0_211, plain, (frontsegP(app(app(X1,X2),X3),X1)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_61]), c_0_56])).
% 49.35/49.55  cnf(c_0_212, negated_conjecture, (app(app(esk8_2(esk49_0,esk49_0),esk49_0),esk9_2(esk49_0,esk49_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_198]), c_0_52])])).
% 49.35/49.55  cnf(c_0_213, negated_conjecture, (ssList(esk8_2(esk49_0,esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_198]), c_0_52])])).
% 49.35/49.55  cnf(c_0_214, negated_conjecture, (ssList(esk9_2(esk49_0,esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_198]), c_0_52])])).
% 49.35/49.55  cnf(c_0_215, negated_conjecture, (esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))=nil|app(X1,esk49_0)!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_85]), c_0_86]), c_0_60])])).
% 49.35/49.55  cnf(c_0_216, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk49_0)=esk49_0|ssItem(hd(esk49_0))), inference(spm,[status(thm)],[c_0_62, c_0_200])).
% 49.35/49.55  cnf(c_0_217, negated_conjecture, (esk49_0=nil|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_201]), c_0_52])])).
% 49.35/49.55  cnf(c_0_218, negated_conjecture, (esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))=esk48_0|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203]), c_0_51])]), c_0_144])).
% 49.35/49.55  cnf(c_0_219, negated_conjecture, (esk8_2(esk49_0,esk48_0)=nil|esk48_0=nil), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_205]), c_0_206]), c_0_195])).
% 49.35/49.55  cnf(c_0_220, plain, (segmentP(app(app(X1,nil),X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208]), c_0_74])])).
% 49.35/49.55  cnf(c_0_221, plain, (frontsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_209])).
% 49.35/49.55  cnf(c_0_222, negated_conjecture, (app(esk49_0,esk6_2(esk49_0,esk49_0))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_210]), c_0_52])])).
% 49.35/49.55  cnf(c_0_223, negated_conjecture, (ssList(esk6_2(esk49_0,esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_210]), c_0_52])])).
% 49.35/49.55  cnf(c_0_224, negated_conjecture, (frontsegP(esk49_0,esk8_2(esk49_0,esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_212]), c_0_213]), c_0_214]), c_0_52])])).
% 49.35/49.55  cnf(c_0_225, negated_conjecture, (esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))=nil|ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216]), c_0_60])]), c_0_217])).
% 49.35/49.55  cnf(c_0_226, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=esk49_0|ssList(tl(esk49_0))), inference(spm,[status(thm)],[c_0_124, c_0_218])).
% 49.35/49.55  cnf(c_0_227, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=nil|esk48_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_219]), c_0_195])).
% 49.35/49.55  cnf(c_0_228, negated_conjecture, (segmentP(app(esk49_0,X1),nil)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_126]), c_0_52])])).
% 49.35/49.55  cnf(c_0_229, plain, (app(X1,esk6_2(X1,X1))=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_75, c_0_221])).
% 49.35/49.55  cnf(c_0_230, negated_conjecture, (esk6_2(esk49_0,esk49_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_222]), c_0_52]), c_0_223])])).
% 49.35/49.55  fof(c_0_231, plain, (![X2]:(~ssList(X2)|rearsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])).
% 49.35/49.55  fof(c_0_232, plain, (![X2]:(((~rearsegP(nil,X2)|nil=X2)|~ssList(X2))&((nil!=X2|rearsegP(nil,X2))|~ssList(X2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).
% 49.35/49.55  cnf(c_0_233, negated_conjecture, (app(esk8_2(esk49_0,esk49_0),app(esk49_0,esk9_2(esk49_0,esk49_0)))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_212]), c_0_214]), c_0_52]), c_0_213])])).
% 49.35/49.55  cnf(c_0_234, negated_conjecture, (app(esk8_2(esk49_0,esk49_0),esk6_2(esk49_0,esk8_2(esk49_0,esk49_0)))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_224]), c_0_213]), c_0_52])])).
% 49.35/49.55  cnf(c_0_235, negated_conjecture, (ssList(esk6_2(esk49_0,esk8_2(esk49_0,esk49_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_224]), c_0_213]), c_0_52])])).
% 49.35/49.55  cnf(c_0_236, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),nil)=esk49_0|ssItem(hd(esk49_0))), inference(spm,[status(thm)],[c_0_85, c_0_225])).
% 49.35/49.55  cnf(c_0_237, negated_conjecture, (app(esk8_2(esk49_0,esk48_0),esk49_0)=esk49_0|ssList(tl(esk49_0))), inference(spm,[status(thm)],[c_0_62, c_0_226])).
% 49.35/49.55  cnf(c_0_238, negated_conjecture, (esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_227]), c_0_59]), c_0_51])])).
% 49.35/49.55  cnf(c_0_239, negated_conjecture, (segmentP(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_229]), c_0_230]), c_0_74]), c_0_52])])).
% 49.35/49.55  cnf(c_0_240, plain, (rearsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_231])).
% 49.35/49.55  cnf(c_0_241, plain, (rearsegP(nil,X1)|~ssList(X1)|nil!=X1), inference(split_conjunct,[status(thm)],[c_0_232])).
% 49.35/49.55  cnf(c_0_242, negated_conjecture, (segmentP(app(esk49_0,X1),esk9_2(esk49_0,esk48_0))|~ssList(app(esk49_0,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_124]), c_0_125]), c_0_59])])).
% 49.35/49.55  cnf(c_0_243, negated_conjecture, (app(esk49_0,esk9_2(esk49_0,esk49_0))=nil|esk8_2(esk49_0,esk49_0)!=esk49_0|~ssList(app(esk49_0,esk9_2(esk49_0,esk49_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_233]), c_0_213])])).
% 49.35/49.55  cnf(c_0_244, negated_conjecture, (hd(esk8_2(esk49_0,esk49_0))=hd(esk49_0)|esk8_2(esk49_0,esk49_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_234]), c_0_235]), c_0_213])])).
% 49.35/49.55  cnf(c_0_245, negated_conjecture, (frontsegP(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_124]), c_0_59]), c_0_125])])).
% 49.35/49.55  cnf(c_0_246, negated_conjecture, (esk8_2(esk49_0,esk48_0)=esk49_0|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_236]), c_0_60])])).
% 49.35/49.55  cnf(c_0_247, negated_conjecture, (esk6_2(esk49_0,esk8_2(esk49_0,esk49_0))=nil|app(X1,esk49_0)!=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_234]), c_0_235]), c_0_213])])).
% 49.35/49.55  cnf(c_0_248, negated_conjecture, (app(esk8_2(esk49_0,nil),esk49_0)=esk49_0|ssList(tl(esk49_0))), inference(rw,[status(thm)],[c_0_237, c_0_238])).
% 49.35/49.55  cnf(c_0_249, negated_conjecture, (ssList(esk8_2(esk49_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_239]), c_0_74]), c_0_52])])).
% 49.35/49.55  cnf(c_0_250, plain, (app(esk7_2(X1,X1),X1)=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_93, c_0_240])).
% 49.35/49.55  cnf(c_0_251, plain, (ssList(esk7_2(X1,X1))|~ssList(X1)), inference(spm,[status(thm)],[c_0_95, c_0_240])).
% 49.35/49.55  cnf(c_0_252, plain, (rearsegP(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_241]), c_0_74])])).
% 49.35/49.55  cnf(c_0_253, plain, (rearsegP(cons(X1,X2),X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_138]), c_0_139])).
% 49.35/49.55  cnf(c_0_254, negated_conjecture, (segmentP(app(esk49_0,X1),esk9_2(esk49_0,esk48_0))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_56]), c_0_52])])).
% 49.35/49.55  cnf(c_0_255, negated_conjecture, (app(esk49_0,esk9_2(esk49_0,esk49_0))=nil|esk8_2(esk49_0,esk49_0)!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_56]), c_0_214]), c_0_52])])).
% 49.35/49.55  cnf(c_0_256, negated_conjecture, (esk8_2(esk49_0,esk49_0)=nil|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_244]), c_0_213])])).
% 49.35/49.55  cnf(c_0_257, negated_conjecture, (X1=esk9_2(esk49_0,esk48_0)|app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),X1)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_124]), c_0_125]), c_0_59])])).
% 49.35/49.55  cnf(c_0_258, negated_conjecture, (app(esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)),esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_245]), c_0_125]), c_0_52])])).
% 49.35/49.55  cnf(c_0_259, negated_conjecture, (ssList(esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_245]), c_0_125]), c_0_52])])).
% 49.35/49.55  cnf(c_0_260, negated_conjecture, (app(esk48_0,esk9_2(esk49_0,esk48_0))=nil|ssItem(hd(esk49_0))), inference(spm,[status(thm)],[c_0_171, c_0_246])).
% 49.35/49.55  cnf(c_0_261, negated_conjecture, (esk8_2(esk49_0,esk49_0)=nil|esk6_2(esk49_0,esk8_2(esk49_0,esk49_0))!=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_234]), c_0_235]), c_0_213])])).
% 49.35/49.55  cnf(c_0_262, negated_conjecture, (esk6_2(esk49_0,esk8_2(esk49_0,esk49_0))=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_248]), c_0_249])]), c_0_144])).
% 49.35/49.55  cnf(c_0_263, plain, (X1=esk7_2(X2,X2)|app(X1,X2)!=X2|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_250]), c_0_251])).
% 49.35/49.55  cnf(c_0_264, plain, (app(esk7_2(nil,nil),nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_252]), c_0_74])])).
% 49.35/49.55  cnf(c_0_265, plain, (ssList(esk7_2(nil,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_252]), c_0_74])])).
% 49.35/49.55  cnf(c_0_266, negated_conjecture, (esk49_0=nil|rearsegP(esk49_0,tl(esk49_0))|~ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_152]), c_0_144])).
% 49.35/49.55  cnf(c_0_267, negated_conjecture, (nil!=esk48_0|~segmentP(esk48_0,X1)|~segmentP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 49.35/49.55  cnf(c_0_268, negated_conjecture, (segmentP(app(esk49_0,X1),esk9_2(esk49_0,nil))|~ssList(X1)), inference(rw,[status(thm)],[c_0_254, c_0_238])).
% 49.35/49.55  cnf(c_0_269, negated_conjecture, (app(esk49_0,esk9_2(esk49_0,esk49_0))=nil|ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_256]), c_0_217])).
% 49.35/49.55  cnf(c_0_270, negated_conjecture, (segmentP(esk49_0,esk9_2(esk49_0,esk48_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_99]), c_0_52]), c_0_74])])).
% 49.35/49.55  cnf(c_0_271, negated_conjecture, (esk6_2(esk49_0,esk7_2(esk49_0,esk9_2(esk49_0,esk48_0)))=esk9_2(esk49_0,esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_258]), c_0_259])])).
% 49.35/49.55  cnf(c_0_272, negated_conjecture, (esk7_2(esk49_0,esk9_2(esk49_0,esk48_0))=esk48_0|ssItem(hd(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_260]), c_0_51])]), c_0_217])).
% 49.35/49.55  cnf(c_0_273, negated_conjecture, (esk8_2(esk49_0,esk49_0)=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_262]), c_0_144])).
% 49.35/49.55  cnf(c_0_274, plain, (app(X1,app(nil,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_99]), c_0_74])])).
% 49.35/49.55  cnf(c_0_275, negated_conjecture, (app(esk8_2(esk49_0,nil),app(nil,esk9_2(esk49_0,nil)))=esk49_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_238]), c_0_238]), c_0_238])).
% 49.35/49.55  cnf(c_0_276, negated_conjecture, (ssList(esk9_2(esk49_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_239]), c_0_74]), c_0_52])])).
% 49.35/49.55  cnf(c_0_277, negated_conjecture, (esk8_2(esk49_0,esk48_0)=esk7_2(esk48_0,esk48_0)|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_169]), c_0_51]), c_0_60])]), c_0_181])).
% 49.35/49.55  cnf(c_0_278, plain, (esk7_2(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_264]), c_0_74]), c_0_265])])).
% 49.35/49.55  cnf(c_0_279, negated_conjecture, (esk49_0=nil|rearsegP(esk49_0,tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_158]), c_0_52])])).
% 49.35/49.55  cnf(c_0_280, plain, (app(esk7_2(X1,nil),nil)=X1|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_73]), c_0_74])])).
% 49.35/49.55  cnf(c_0_281, plain, (ssList(esk7_2(X1,nil))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_73]), c_0_74])])).
% 49.35/49.55  cnf(c_0_282, negated_conjecture, (~segmentP(nil,X1)|~segmentP(esk49_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_267, c_0_238]), c_0_238])])).
% 49.35/49.55  cnf(c_0_283, negated_conjecture, (segmentP(nil,esk9_2(esk49_0,nil))|ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_269]), c_0_214])])).
% 49.35/49.55  cnf(c_0_284, negated_conjecture, (segmentP(esk49_0,esk9_2(esk49_0,nil))), inference(rw,[status(thm)],[c_0_270, c_0_238])).
% 49.35/49.55  cnf(c_0_285, negated_conjecture, (esk9_2(esk49_0,esk48_0)=esk6_2(esk49_0,esk48_0)|ssItem(hd(esk49_0))), inference(spm,[status(thm)],[c_0_271, c_0_272])).
% 49.35/49.55  cnf(c_0_286, negated_conjecture, (neq(esk49_0,nil)|nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_41])).
% 49.35/49.55  cnf(c_0_287, negated_conjecture, (app(esk49_0,esk9_2(esk49_0,esk49_0))=nil|ssList(tl(esk49_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_273]), c_0_144])).
% 49.35/49.55  cnf(c_0_288, negated_conjecture, (app(nil,esk9_2(esk49_0,esk48_0))=esk49_0|ssList(tl(esk49_0))), inference(spm,[status(thm)],[c_0_124, c_0_143])).
% 49.35/49.55  cnf(c_0_289, plain, (nil=X1|rearsegP(app(X1,X2),tl(app(X1,X2)))|~ssList(tl(app(X1,X2)))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_140]), c_0_158])).
% 49.35/49.55  cnf(c_0_290, negated_conjecture, (app(esk8_2(esk49_0,nil),esk9_2(esk49_0,nil))=esk49_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_275]), c_0_276]), c_0_249])])).
% 49.35/49.55  cnf(c_0_291, negated_conjecture, (esk8_2(esk49_0,nil)=nil|ssList(tl(esk49_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_277, c_0_238]), c_0_238]), c_0_238]), c_0_278])).
% 49.35/49.55  cnf(c_0_292, plain, (X1=cons(X2,nil)|app(X1,X3)!=cons(X2,X3)|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_138]), c_0_139])).
% 49.35/49.55  cnf(c_0_293, negated_conjecture, (app(esk7_2(esk49_0,tl(esk49_0)),tl(esk49_0))=esk49_0|esk49_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_279]), c_0_52])]), c_0_144])).
% 49.35/49.55  cnf(c_0_294, negated_conjecture, (esk49_0=nil|ssList(esk7_2(esk49_0,tl(esk49_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_279]), c_0_52])]), c_0_144])).
% 49.35/49.55  cnf(c_0_295, plain, (app(tl(esk7_2(X1,nil)),nil)=tl(X1)|esk7_2(X1,nil)=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_280]), c_0_74])]), c_0_281])).
% 49.35/49.55  cnf(c_0_296, negated_conjecture, (ssItem(hd(esk49_0))|~neq(esk9_2(esk49_0,nil),nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_283]), c_0_284]), c_0_276])])).
% 49.35/49.55  cnf(c_0_297, negated_conjecture, (esk9_2(esk49_0,nil)=esk49_0|ssItem(hd(esk49_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_285, c_0_238]), c_0_238]), c_0_162])).
% 49.35/49.55  cnf(c_0_298, negated_conjecture, (neq(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_286, c_0_238])])).
% 49.35/49.55  cnf(c_0_299, negated_conjecture, (esk8_2(esk49_0,esk48_0)=nil|esk49_0!=nil|~ssList(app(esk48_0,esk9_2(esk49_0,esk48_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_62]), c_0_60])])).
% 49.35/49.55  cnf(c_0_300, negated_conjecture, (segmentP(nil,esk9_2(esk49_0,nil))|ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_287]), c_0_214])])).
% 49.35/49.55  cnf(c_0_301, negated_conjecture, (esk9_2(esk49_0,esk48_0)=esk49_0|ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_288]), c_0_59])])).
% 49.35/49.55  cnf(c_0_302, negated_conjecture, (esk8_2(esk49_0,nil)=nil|rearsegP(esk49_0,tl(esk49_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_290]), c_0_276]), c_0_249])]), c_0_291])).
% 49.35/49.55  cnf(c_0_303, negated_conjecture, (esk7_2(esk49_0,tl(esk49_0))=cons(X1,nil)|esk49_0=nil|cons(X1,tl(esk49_0))!=esk49_0|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_293]), c_0_294]), c_0_144])).
% 49.35/49.55  cnf(c_0_304, plain, (cons(hd(app(esk7_2(X1,nil),nil)),tl(X1))=app(esk7_2(X1,nil),nil)|esk7_2(X1,nil)=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_295]), c_0_74])]), c_0_281])).
% 49.35/49.55  cnf(c_0_305, negated_conjecture, (ssItem(hd(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_297]), c_0_298])])).
% 49.35/49.55  cnf(c_0_306, negated_conjecture, (esk8_2(esk49_0,esk48_0)=nil|esk49_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_56]), c_0_59]), c_0_51])])).
% 49.35/49.55  cnf(c_0_307, negated_conjecture, (~totalorderedP(X1)|~segmentP(X1,esk50_0)|~segmentP(esk51_0,X1)|~neq(esk50_0,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 49.35/49.55  fof(c_0_308, plain, (![X2]:(~ssList(X2)|segmentP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])).
% 49.35/49.55  cnf(c_0_309, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(spm,[status(thm)],[c_0_207, c_0_175])).
% 49.35/49.55  cnf(c_0_310, negated_conjecture, (ssList(tl(esk49_0))|~neq(esk9_2(esk49_0,nil),nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_300]), c_0_284]), c_0_276])])).
% 49.35/49.55  cnf(c_0_311, negated_conjecture, (esk9_2(esk49_0,nil)=esk49_0|ssList(tl(esk49_0))), inference(rw,[status(thm)],[c_0_301, c_0_238])).
% 49.35/49.55  cnf(c_0_312, negated_conjecture, (esk8_2(esk49_0,nil)=nil|ssList(esk7_2(esk49_0,tl(esk49_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_302]), c_0_52])]), c_0_291])).
% 49.35/49.55  cnf(c_0_313, negated_conjecture, (esk7_2(esk49_0,tl(esk49_0))=cons(hd(esk49_0),nil)|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_304]), c_0_112]), c_0_126]), c_0_112]), c_0_112]), c_0_126]), c_0_112]), c_0_126]), c_0_305]), c_0_52])])).
% 49.35/49.55  cnf(c_0_314, negated_conjecture, (esk8_2(esk49_0,nil)=nil|esk49_0!=nil), inference(rw,[status(thm)],[c_0_306, c_0_238])).
% 49.35/49.55  cnf(c_0_315, negated_conjecture, (~totalorderedP(X1)|~segmentP(X1,esk48_0)|~segmentP(esk49_0,X1)|~ssList(X1)|~neq(esk48_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_307, c_0_45]), c_0_45]), c_0_44])).
% 49.35/49.55  cnf(c_0_316, plain, (segmentP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_308])).
% 49.35/49.55  cnf(c_0_317, plain, (segmentP(app(X1,X2),X1)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_250]), c_0_251])).
% 49.35/49.55  cnf(c_0_318, negated_conjecture, (ssList(tl(esk49_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_311]), c_0_298])])).
% 49.35/49.55  cnf(c_0_319, negated_conjecture, (app(nil,esk9_2(esk49_0,nil))=nil|esk8_2(esk49_0,nil)!=esk49_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_171, c_0_238]), c_0_238]), c_0_238])).
% 49.35/49.55  cnf(c_0_320, negated_conjecture, (esk8_2(esk49_0,nil)=nil|ssList(cons(hd(esk49_0),nil))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_313]), c_0_314])).
% 49.35/49.55  cnf(c_0_321, negated_conjecture, (esk49_0=nil|ssList(cons(hd(esk49_0),nil))), inference(spm,[status(thm)],[c_0_294, c_0_313])).
% 49.35/49.55  cnf(c_0_322, negated_conjecture, (X1=esk6_2(esk49_0,esk8_2(esk49_0,esk48_0))|app(esk8_2(esk49_0,esk48_0),X1)!=esk49_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_85]), c_0_60]), c_0_86])])).
% 49.35/49.55  cnf(c_0_323, negated_conjecture, (~totalorderedP(X1)|~segmentP(esk49_0,X1)|~ssList(X1)|~neq(nil,X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_315, c_0_238]), c_0_238]), c_0_316])).
% 49.35/49.55  cnf(c_0_324, negated_conjecture, (esk49_0=nil|segmentP(esk49_0,esk7_2(esk49_0,tl(esk49_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_293]), c_0_318])]), c_0_294])).
% 49.35/49.55  cnf(c_0_325, plain, (segmentP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_99]), c_0_74])]), c_0_56])).
% 49.35/49.55  cnf(c_0_326, negated_conjecture, (app(nil,esk9_2(esk49_0,nil))=nil|ssList(cons(hd(esk49_0),nil))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_320]), c_0_321])).
% 49.35/49.55  cnf(c_0_327, negated_conjecture, (X1=esk6_2(esk49_0,esk8_2(esk49_0,nil))|app(esk8_2(esk49_0,nil),X1)!=esk49_0|~ssList(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_322, c_0_238]), c_0_238])).
% 49.35/49.55  cnf(c_0_328, negated_conjecture, (esk49_0=nil|~totalorderedP(esk7_2(esk49_0,tl(esk49_0)))|~neq(nil,esk7_2(esk49_0,tl(esk49_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_324]), c_0_294])).
% 49.35/49.55  fof(c_0_329, plain, (![X2]:(~ssItem(X2)|totalorderedP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])).
% 49.35/49.55  cnf(c_0_330, negated_conjecture, (segmentP(nil,esk9_2(esk49_0,nil))|ssList(cons(hd(esk49_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326]), c_0_276]), c_0_74])])).
% 49.35/49.55  cnf(c_0_331, negated_conjecture, (esk6_2(esk49_0,esk8_2(esk49_0,nil))=esk9_2(esk49_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_290]), c_0_276])])).
% 49.35/49.55  cnf(c_0_332, negated_conjecture, (esk49_0=nil|~totalorderedP(cons(hd(esk49_0),nil))|~neq(nil,cons(hd(esk49_0),nil))), inference(spm,[status(thm)],[c_0_328, c_0_313])).
% 49.35/49.55  cnf(c_0_333, plain, (totalorderedP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_329])).
% 49.35/49.55  cnf(c_0_334, negated_conjecture, (ssList(cons(hd(esk49_0),nil))|~neq(esk9_2(esk49_0,nil),nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_330]), c_0_284]), c_0_276])])).
% 49.35/49.55  cnf(c_0_335, negated_conjecture, (esk9_2(esk49_0,nil)=esk49_0|ssList(cons(hd(esk49_0),nil))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_331, c_0_320]), c_0_162])).
% 49.35/49.55  cnf(c_0_336, negated_conjecture, (totalorderedP(esk50_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 49.35/49.55  fof(c_0_337, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)!=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_39])])])).
% 49.35/49.55  cnf(c_0_338, negated_conjecture, (esk49_0=nil|~neq(nil,cons(hd(esk49_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_333]), c_0_305])])).
% 49.35/49.55  cnf(c_0_339, negated_conjecture, (ssList(cons(hd(esk49_0),nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_335]), c_0_298])])).
% 49.35/49.55  cnf(c_0_340, negated_conjecture, (totalorderedP(esk48_0)), inference(rw,[status(thm)],[c_0_336, c_0_45])).
% 49.35/49.55  cnf(c_0_341, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_337])).
% 49.35/49.55  cnf(c_0_342, negated_conjecture, (cons(hd(esk49_0),nil)=nil|esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_186]), c_0_339]), c_0_74])])).
% 49.35/49.55  cnf(c_0_343, negated_conjecture, (~neq(esk48_0,esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_175]), c_0_340]), c_0_50]), c_0_51])])).
% 49.35/49.55  cnf(c_0_344, negated_conjecture, (esk49_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_342]), c_0_74]), c_0_305])])).
% 49.35/49.55  cnf(c_0_345, negated_conjecture, (~neq(nil,nil)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_343, c_0_238]), c_0_238])).
% 49.35/49.55  cnf(c_0_346, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_298, c_0_344]), c_0_345]), ['proof']).
% 49.35/49.55  # SZS output end CNFRefutation
% 49.35/49.55  # Training examples: 0 positive, 0 negative
% 49.44/49.66  
% 49.44/49.66  # -------------------------------------------------
% 49.44/49.66  # User time                : 48.377 s
% 49.44/49.66  # System time              : 1.080 s
% 49.44/49.66  # Total time               : 49.457 s
% 49.44/49.66  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
