<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC018+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC018+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC018+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC018%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC018+1.p">SWC018+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n155.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:19 EDT 2016

% Result     : Theorem 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of formulae       :   42 ( 386 expanded)
%              Number of clauses        :   28 ( 137 expanded)
%              Number of leaves         :    7 (  93 expanded)
%              Depth                    :   10
%              Number of atoms          :  196 (5588 expanded)
%              Number of equality atoms :   68 (2100 expanded)
%              Maximal formula depth    :   24 (   5 average)
%              Maximal clause size      :   30 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ! [X5] :
                        ( ssList(X5)
                       => ( app(X3,X5) != X4
                          | ~ equalelemsP(X3)
                          | ? [X6] :
                              ( ssItem(X6)
                              & ? [X7] :
                                  ( ssList(X7)
                                  & app(cons(X6,nil),X7) = X5
                                  & ? [X8] :
                                      ( ssList(X8)
                                      & app(X8,cons(X6,nil)) = X3 ) ) ) ) )
                    | ( nil != X4
                      & nil = X3 )
                    | ( ( nil != X2
                        | nil = X1 )
                      & ( ~ neq(X2,nil)
                        | ? [X9] :
                            ( ssList(X9)
                            & neq(X9,nil)
                            & frontsegP(X2,X9)
                            & frontsegP(X1,X9) ) ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( frontsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X2,X3) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax5)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax42)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( neq(X1,X2)
          &lt;=> X1 != X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax15)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_7"></A>fof(c_0_7,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ! [X5] :
                          ( ssList(X5)
                         => ( app(X3,X5) != X4
                            | ~ equalelemsP(X3)
                            | ? [X6] :
                                ( ssItem(X6)
                                & ? [X7] :
                                    ( ssList(X7)
                                    & app(cons(X6,nil),X7) = X5
                                    & ? [X8] :
                                        ( ssList(X8)
                                        & app(X8,cons(X6,nil)) = X3 ) ) ) ) )
                      | ( nil != X4
                        & nil = X3 )
                      | ( ( nil != X2
                          | nil = X1 )
                        & ( ~ neq(X2,nil)
                          | ? [X9] :
                              ( ssList(X9)
                              & neq(X9,nil)
                              & frontsegP(X2,X9)
                              & frontsegP(X1,X9) ) ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_8"></A>fof(c_0_8,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk6_2(X4,X5))
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(X5,esk6_2(X4,X5)) = X4
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X5,X7) != X4
        | frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).

<A NAME="c_0_9"></A>fof(c_0_9,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_10"></A>fof(c_0_10,negated_conjecture,(
    ! [X15,X16,X17,X18] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ssList(esk52_0)
      & app(esk50_0,esk52_0) = esk51_0
      & equalelemsP(esk50_0)
      & ( ~ ssItem(X15)
        | ~ ssList(X16)
        | app(cons(X15,nil),X16) != esk52_0
        | ~ ssList(X17)
        | app(X17,cons(X15,nil)) != esk50_0 )
      & ( nil = esk51_0
        | nil != esk50_0 )
      & ( neq(esk49_0,nil)
        | nil = esk49_0 )
      & ( ~ ssList(X18)
        | ~ neq(X18,nil)
        | ~ frontsegP(esk49_0,X18)
        | ~ frontsegP(esk48_0,X18)
        | nil = esk49_0 )
      & ( neq(esk49_0,nil)
        | nil != esk48_0 )
      & ( ~ ssList(X18)
        | ~ neq(X18,nil)
        | ~ frontsegP(esk49_0,X18)
        | ~ frontsegP(esk48_0,X18)
        | nil != esk48_0 ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])])).

<A NAME="c_0_11"></A>fof(c_0_11,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).

<A NAME="c_0_12"></A>cnf(c_0_12,plain,
    ( frontsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X2,X3) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>])).

<A NAME="c_0_13"></A>cnf(c_0_13,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_14"></A>cnf(c_0_14,negated_conjecture,
    ( app(esk50_0,esk52_0) = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( esk49_0 = esk51_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( nil = esk49_0
    | ~ frontsegP(esk48_0,X1)
    | ~ frontsegP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,plain,
    ( frontsegP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_20"></A>cnf(c_0_20,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_12]),<A HREF="#c_0_13">c_0_13</A>])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( app(esk48_0,esk52_0) = esk49_0 ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_15]),<A HREF="#c_0_16">c_0_16</A>])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( ssList(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_23"></A>fof(c_0_23,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).

<A NAME="c_0_24"></A>cnf(c_0_24,negated_conjecture,
    ( esk49_0 = nil
    | ~ frontsegP(esk49_0,esk48_0)
    | ~ neq(esk48_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_18]),c_0_19])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( frontsegP(esk49_0,esk48_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_21]),c_0_22]),c_0_19])])).

<A NAME="c_0_26"></A>fof(c_0_26,plain,(
    ! [X3,X4] :
      ( ( ~ neq(X3,X4)
        | X3 != X4
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( X3 = X4
        | neq(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).

<A NAME="c_0_27"></A>cnf(c_0_27,negated_conjecture,
    ( nil = esk51_0
    | nil != esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_28"></A>cnf(c_0_28,negated_conjecture,
    ( nil != esk48_0
    | ~ frontsegP(esk48_0,X1)
    | ~ frontsegP(esk49_0,X1)
    | ~ neq(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_29"></A>cnf(c_0_29,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_30"></A>cnf(c_0_30,negated_conjecture,
    ( esk49_0 = nil
    | ~ neq(esk48_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_25])])).

<A NAME="c_0_31"></A>cnf(c_0_31,plain,
    ( neq(X1,X2)
    | X1 = X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_32"></A>cnf(c_0_32,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_6">c_0_6</A>])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( esk49_0 = nil
    | esk48_0 != nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_16]),<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( esk48_0 != nil
    | ~ frontsegP(esk49_0,esk48_0)
    | ~ neq(esk48_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_18]),c_0_19])])).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( esk48_0 = nil
    | esk49_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>,c_0_21]),c_0_22]),c_0_19])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( esk49_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_30">c_0_30</A>,c_0_31]),c_0_32]),c_0_19])]),<A HREF="#c_0_33">c_0_33</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,negated_conjecture,
    ( esk48_0 != nil
    | ~ neq(esk48_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_34">c_0_34</A>,c_0_25])])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( esk48_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_36])])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( neq(esk49_0,nil)
    | nil != esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( ~ neq(nil,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_38]),c_0_38])])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>,c_0_36]),c_0_38])]),<A HREF="#c_0_40">c_0_40</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC018+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n155.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:34:24 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.24  # No SInE strategy applied
% 0.02/0.24  # Trying AutoSched0 for 151 seconds
% 0.02/0.27  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.02/0.27  # and selection function SelectCQPrecWNTNp.
% 0.02/0.27  #
% 0.02/0.27  # Preprocessing time       : 0.014 s
% 0.02/0.27  # Presaturation interreduction done
% 0.02/0.27  
% 0.02/0.27  # Proof found!
% 0.02/0.27  # SZS status Theorem
% 0.02/0.27  # SZS output start CNFRefutation
% 0.02/0.27  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>((app(X3,X5)!=X4|~(equalelemsP(X3)))|?[X6]:(ssItem(X6)&?[X7]:((ssList(X7)&app(cons(X6,nil),X7)=X5)&?[X8]:(ssList(X8)&app(X8,cons(X6,nil))=X3))))))|(nil!=X4&nil=X3))|((nil!=X2|nil=X1)&(~(neq(X2,nil))|?[X9]:(((ssList(X9)&neq(X9,nil))&frontsegP(X2,X9))&frontsegP(X1,X9)))))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.02/0.27  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax5)).
% 0.02/0.27  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.02/0.27  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax42)).
% 0.02/0.27  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.02/0.27  fof(c_0_5, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)&lt;=>X1!=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax15)).
% 0.02/0.27  fof(c_0_6, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.02/0.27  fof(c_0_7, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>((((X2!=X4|X1!=X3)|![X5]:(ssList(X5)=>((app(X3,X5)!=X4|~equalelemsP(X3))|?[X6]:(ssItem(X6)&?[X7]:((ssList(X7)&app(cons(X6,nil),X7)=X5)&?[X8]:(ssList(X8)&app(X8,cons(X6,nil))=X3))))))|(nil!=X4&nil=X3))|((nil!=X2|nil=X1)&(~neq(X2,nil)|?[X9]:(((ssList(X9)&neq(X9,nil))&frontsegP(X2,X9))&frontsegP(X1,X9))))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.02/0.27  fof(c_0_8, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
% 0.02/0.27  fof(c_0_9, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.02/0.27  fof(c_0_10, negated_conjecture, (![X15]:![X16]:![X17]:![X18]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&((((esk49_0=esk51_0&esk48_0=esk50_0)&(ssList(esk52_0)&((app(esk50_0,esk52_0)=esk51_0&equalelemsP(esk50_0))&(~ssItem(X15)|((~ssList(X16)|app(cons(X15,nil),X16)!=esk52_0)|(~ssList(X17)|app(X17,cons(X15,nil))!=esk50_0))))))&(nil=esk51_0|nil!=esk50_0))&(((neq(esk49_0,nil)|nil=esk49_0)&((((~ssList(X18)|~neq(X18,nil))|~frontsegP(esk49_0,X18))|~frontsegP(esk48_0,X18))|nil=esk49_0))&((neq(esk49_0,nil)|nil!=esk48_0)&((((~ssList(X18)|~neq(X18,nil))|~frontsegP(esk49_0,X18))|~frontsegP(esk48_0,X18))|nil!=esk48_0))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])])).
% 0.02/0.27  fof(c_0_11, plain, (![X2]:(~ssList(X2)|frontsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])).
% 0.02/0.27  cnf(c_0_12, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
% 0.02/0.27  cnf(c_0_13, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
% 0.02/0.27  cnf(c_0_14, negated_conjecture, (app(esk50_0,esk52_0)=esk51_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.27  cnf(c_0_15, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.27  cnf(c_0_16, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.27  cnf(c_0_17, negated_conjecture, (nil=esk49_0|~frontsegP(esk48_0,X1)|~frontsegP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.27  cnf(c_0_18, plain, (frontsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
% 0.02/0.27  cnf(c_0_19, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.27  cnf(c_0_20, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_12]), c_0_13])).
% 0.02/0.27  cnf(c_0_21, negated_conjecture, (app(esk48_0,esk52_0)=esk49_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
% 0.02/0.27  cnf(c_0_22, negated_conjecture, (ssList(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.27  fof(c_0_23, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
% 0.02/0.27  cnf(c_0_24, negated_conjecture, (esk49_0=nil|~frontsegP(esk49_0,esk48_0)|~neq(esk48_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
% 0.02/0.27  cnf(c_0_25, negated_conjecture, (frontsegP(esk49_0,esk48_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_19])])).
% 0.02/0.27  fof(c_0_26, plain, (![X3]:![X4]:((((~neq(X3,X4)|X3!=X4)|~ssList(X4))|~ssList(X3))&(((X3=X4|neq(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
% 0.02/0.27  cnf(c_0_27, negated_conjecture, (nil=esk51_0|nil!=esk50_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.27  cnf(c_0_28, negated_conjecture, (nil!=esk48_0|~frontsegP(esk48_0,X1)|~frontsegP(esk49_0,X1)|~neq(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.27  cnf(c_0_29, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
% 0.02/0.27  cnf(c_0_30, negated_conjecture, (esk49_0=nil|~neq(esk48_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25])])).
% 0.02/0.27  cnf(c_0_31, plain, (neq(X1,X2)|X1=X2|~ssList(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
% 0.02/0.27  cnf(c_0_32, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_6])).
% 0.02/0.27  cnf(c_0_33, negated_conjecture, (esk49_0=nil|esk48_0!=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_16]), c_0_15])).
% 0.02/0.27  cnf(c_0_34, negated_conjecture, (esk48_0!=nil|~frontsegP(esk49_0,esk48_0)|~neq(esk48_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_18]), c_0_19])])).
% 0.02/0.27  cnf(c_0_35, negated_conjecture, (esk48_0=nil|esk49_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_21]), c_0_22]), c_0_19])])).
% 0.02/0.27  cnf(c_0_36, negated_conjecture, (esk49_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_19])]), c_0_33])).
% 0.02/0.27  cnf(c_0_37, negated_conjecture, (esk48_0!=nil|~neq(esk48_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_25])])).
% 0.02/0.27  cnf(c_0_38, negated_conjecture, (esk48_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36])])).
% 0.02/0.27  cnf(c_0_39, negated_conjecture, (neq(esk49_0,nil)|nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_10])).
% 0.02/0.27  cnf(c_0_40, negated_conjecture, (~neq(nil,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38]), c_0_38])])).
% 0.02/0.27  cnf(c_0_41, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_36]), c_0_38])]), c_0_40]), ['proof']).
% 0.02/0.27  # SZS output end CNFRefutation
% 0.02/0.27  # Training examples: 0 positive, 0 negative
% 0.02/0.27  
% 0.02/0.27  # -------------------------------------------------
% 0.02/0.27  # User time                : 0.033 s
% 0.02/0.27  # System time              : 0.003 s
% 0.02/0.27  # Total time               : 0.036 s
% 0.02/0.27  # Maximum resident set size: 1876 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
