<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC184+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC184+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC184+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC184%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC184+1.p">SWC184+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n156.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:58 EDT 2016

% Result     : Theorem 0.06s
% Output     : CNFRefutation 0.06s
% Verified   : 
% Statistics : Number of formulae       :  248 (65439 expanded)
%              Number of clauses        :  171 (24013 expanded)
%              Number of leaves         :   38 (16497 expanded)
%              Depth                    :   33
%              Number of atoms          :  918 (828432 expanded)
%              Number of equality atoms :  301 (305031 expanded)
%              Maximal formula depth    :   25 (   5 average)
%              Maximal clause size      :   34 (   3 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ~ ssList(X4)
                  | X2 != X4
                  | X1 != X3
                  | ! [X5] :
                      ( ssItem(X5)
                     => ! [X6] :
                          ( ssList(X6)
                         => ! [X7] :
                              ( ~ ssList(X7)
                              | app(app(X6,cons(X5,nil)),X7) != X1
                              | ( ~ memberP(X6,X5)
                                & ~ memberP(X7,X5) ) ) ) )
                  | ( ! [X8] :
                        ( ~ ssItem(X8)
                        | cons(X8,nil) != X3
                        | ~ memberP(X4,X8)
                        | ? [X9] :
                            ( ssItem(X9)
                            & X8 != X9
                            & memberP(X4,X9)
                            & leq(X9,X8) ) )
                    & ( nil != X4
                      | nil != X3 ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( frontsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X2,X3) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax5)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => cons(hd(X1),tl(X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax78)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => hd(cons(X2,X1)) = X2 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax23)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => nil != cons(X2,X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax21)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( singletonP(X1)
      &lt;=> ? [X2] :
            ( ssItem(X2)
            & cons(X2,nil) = X1 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax4)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => hd(app(X1,X2)) = hd(X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax85)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssItem(hd(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax22)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax38)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( memberP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(X3,cons(X2,X4)) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax3)).

<A NAME="c_0_17"></A>fof(c_0_17,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssList(tl(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

<A NAME="c_0_18"></A>fof(c_0_18,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax86)).

<A NAME="c_0_19"></A>fof(c_0_19,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(cons(X2,X3),X1)
              &lt;=> ( X1 = X2
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax37)).

<A NAME="c_0_20"></A>fof(c_0_20,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X2,X3) = app(X2,X1)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax80)).

<A NAME="c_0_21"></A>fof(c_0_21,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_22"></A>fof(c_0_22,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( equalelemsP(X1)
      &lt;=> ! [X2] :
            ( ssItem(X2)
           => ! [X3] :
                ( ssItem(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ! [X5] :
                        ( ssList(X5)
                       => ( app(X4,cons(X2,cons(X3,X5))) = X1
                         => X2 = X3 ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax14)).

<A NAME="c_0_23"></A>fof(c_0_23,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil = X1
        | ? [X2] :
            ( ssList(X2)
            & ? [X3] :
                ( ssItem(X3)
                & cons(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax20)).

<A NAME="c_0_24"></A>fof(c_0_24,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => equalelemsP(cons(X1,nil)) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax73)).

<A NAME="c_0_25"></A>fof(c_0_25,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( frontsegP(cons(X1,X3),cons(X2,X4))
                  &lt;=> ( X1 = X2
                      & frontsegP(X3,X4) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax44)).

<A NAME="c_0_26"></A>fof(c_0_26,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X3,X2) = app(X1,X2)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax79)).

<A NAME="c_0_27"></A>fof(c_0_27,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_28"></A>fof(c_0_28,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax42)).

<A NAME="c_0_29"></A>fof(c_0_29,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( segmentP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(app(X3,X2),X4) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax7)).

<A NAME="c_0_30"></A>fof(c_0_30,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( rearsegP(X1,X2)
               => rearsegP(app(X3,X1),X2) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax50)).

<A NAME="c_0_31"></A>fof(c_0_31,axiom,(
    ! [X1] :
      ( ssList(X1)
     => rearsegP(X1,nil) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax51)).

<A NAME="c_0_32"></A>fof(c_0_32,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( frontsegP(X1,X2)
               => frontsegP(app(X1,X3),X2) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax43)).

<A NAME="c_0_33"></A>fof(c_0_33,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,nil) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax45)).

<A NAME="c_0_34"></A>fof(c_0_34,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( rearsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax6)).

<A NAME="c_0_35"></A>fof(c_0_35,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

<A NAME="c_0_36"></A>fof(c_0_36,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( ( rearsegP(X1,X2)
              & rearsegP(X2,X1) )
           => X1 = X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax48)).

<A NAME="c_0_37"></A>fof(c_0_37,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( ( segmentP(X1,X2)
              & segmentP(X2,X1) )
           => X1 = X2 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax54)).

<A NAME="c_0_38"></A>fof(c_0_38,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ~ ssList(X4)
                    | X2 != X4
                    | X1 != X3
                    | ! [X5] :
                        ( ssItem(X5)
                       => ! [X6] :
                            ( ssList(X6)
                           => ! [X7] :
                                ( ~ ssList(X7)
                                | app(app(X6,cons(X5,nil)),X7) != X1
                                | ( ~ memberP(X6,X5)
                                  & ~ memberP(X7,X5) ) ) ) )
                    | ( ! [X8] :
                          ( ~ ssItem(X8)
                          | cons(X8,nil) != X3
                          | ~ memberP(X4,X8)
                          | ? [X9] :
                              ( ssItem(X9)
                              & X8 != X9
                              & memberP(X4,X9)
                              & leq(X9,X8) ) )
                      & ( nil != X4
                        | nil != X3 ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_39"></A>fof(c_0_39,negated_conjecture,(
    ! [X18] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ssItem(esk52_0)
      & ssList(esk53_0)
      & ssList(esk54_0)
      & app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk48_0
      & ( memberP(esk53_0,esk52_0)
        | memberP(esk54_0,esk52_0) )
      & ( nil = esk51_0
        | ssItem(esk55_0) )
      & ( nil = esk50_0
        | ssItem(esk55_0) )
      & ( nil = esk51_0
        | cons(esk55_0,nil) = esk50_0 )
      & ( nil = esk50_0
        | cons(esk55_0,nil) = esk50_0 )
      & ( nil = esk51_0
        | memberP(esk51_0,esk55_0) )
      & ( nil = esk50_0
        | memberP(esk51_0,esk55_0) )
      & ( nil = esk51_0
        | ~ ssItem(X18)
        | esk55_0 = X18
        | ~ memberP(esk51_0,X18)
        | ~ leq(X18,esk55_0) )
      & ( nil = esk50_0
        | ~ ssItem(X18)
        | esk55_0 = X18
        | ~ memberP(esk51_0,X18)
        | ~ leq(X18,esk55_0) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])])])])).

<A NAME="c_0_40"></A>cnf(c_0_40,negated_conjecture,
    ( app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_41"></A>cnf(c_0_41,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_42"></A>fof(c_0_42,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( ssList(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( ssList(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_47"></A>fof(c_0_47,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_48"></A>fof(c_0_48,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk6_2(X4,X5))
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(X5,esk6_2(X4,X5)) = X4
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X5,X7) != X4
        | frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).

<A NAME="c_0_49"></A>fof(c_0_49,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( app(esk53_0,app(cons(esk52_0,nil),esk54_0)) = esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_44]),c_0_45]),c_0_46])])).

<A NAME="c_0_51"></A>cnf(c_0_51,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_52"></A>cnf(c_0_52,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,negated_conjecture,
    ( ssItem(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_54"></A>fof(c_0_54,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).

<A NAME="c_0_55"></A>cnf(c_0_55,plain,
    ( frontsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X2,X3) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( app(esk53_0,app(cons(esk52_0,nil),esk54_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>,c_0_51]),c_0_52]),c_0_53])])).

<A NAME="c_0_58"></A>cnf(c_0_58,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>])).

<A NAME="c_0_59"></A>fof(c_0_59,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | cons(hd(X2),tl(X2)) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

<A NAME="c_0_60"></A>fof(c_0_60,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | hd(cons(X4,X3)) = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

<A NAME="c_0_61"></A>fof(c_0_61,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | nil != cons(X4,X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_62"></A>fof(c_0_62,plain,(
    ! [X3,X5] :
      ( ( ssItem(esk5_1(X3))
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( cons(esk5_1(X3),nil) = X3
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( ~ ssItem(X5)
        | cons(X5,nil) != X3
        | singletonP(X3)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).

<A NAME="c_0_63"></A>cnf(c_0_63,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_55]),<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( app(esk53_0,cons(esk52_0,esk54_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,c_0_58]),c_0_45]),c_0_53])])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,plain,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>])).

<A NAME="c_0_67"></A>cnf(c_0_67,plain,
    ( nil != cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>])).

<A NAME="c_0_68"></A>fof(c_0_68,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).

<A NAME="c_0_69"></A>cnf(c_0_69,plain,
    ( singletonP(X1)
    | ~ ssList(X1)
    | cons(X2,nil) != X1
    | ~ ssItem(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( frontsegP(esk50_0,esk53_0)
    | ~ ssList(cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>,c_0_64]),c_0_46])])).

<A NAME="c_0_72"></A>cnf(c_0_72,plain,
    ( cons(X1,tl(cons(X1,X2))) = cons(X1,X2)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_66]),c_0_51]),<A HREF="#c_0_67">c_0_67</A>])).

<A NAME="c_0_73"></A>cnf(c_0_73,plain,
    ( cons(esk5_1(X1),nil) = X1
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>])).

<A NAME="c_0_74"></A>cnf(c_0_74,plain,
    ( ssItem(esk5_1(X1))
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>])).

<A NAME="c_0_75"></A>cnf(c_0_75,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_76"></A>cnf(c_0_76,negated_conjecture,
    ( cons(esk55_0,nil) = esk50_0
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,negated_conjecture,
    ( ssItem(esk55_0)
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_78"></A>cnf(c_0_78,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>,<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_80"></A>fof(c_0_80,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | hd(app(X3,X4)) = hd(X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( app(X2,esk6_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( frontsegP(esk50_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_71">c_0_71</A>,c_0_51]),c_0_45]),c_0_53])])).

<A NAME="c_0_83"></A>cnf(c_0_83,plain,
    ( ssList(esk6_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_84"></A>fof(c_0_84,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).

<A NAME="c_0_85"></A>cnf(c_0_85,plain,
    ( cons(esk5_1(X1),tl(X1)) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_73]),c_0_52])]),<A HREF="#c_0_74">c_0_74</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( tl(esk50_0) = nil
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_75">c_0_75</A>,c_0_76]),c_0_52])]),<A HREF="#c_0_77">c_0_77</A>])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( esk50_0 = nil
    | singletonP(esk50_0) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,c_0_76]),c_0_79])]),<A HREF="#c_0_77">c_0_77</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,plain,
    ( hd(app(X1,X2)) = hd(X1)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,negated_conjecture,
    ( app(esk53_0,esk6_2(esk50_0,esk53_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_82]),c_0_46]),c_0_79])])).

<A NAME="c_0_90"></A>cnf(c_0_90,negated_conjecture,
    ( ssList(esk6_2(esk50_0,esk53_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_82]),c_0_46]),c_0_79])])).

<A NAME="c_0_91"></A>cnf(c_0_91,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_92"></A>fof(c_0_92,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssItem(hd(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).

<A NAME="c_0_93"></A>fof(c_0_93,plain,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>])).

<A NAME="c_0_94"></A>fof(c_0_94,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk3_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ssList(esk4_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6))) = X5
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(X9,cons(X6,X10)) != X5
        | memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])])).

<A NAME="c_0_95"></A>fof(c_0_95,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssList(tl(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).

<A NAME="c_0_96"></A>fof(c_0_96,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).

<A NAME="c_0_97"></A>cnf(c_0_97,negated_conjecture,
    ( cons(esk5_1(esk50_0),nil) = esk50_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_86]),c_0_79])]),<A HREF="#c_0_87">c_0_87</A>])).

<A NAME="c_0_98"></A>cnf(c_0_98,plain,
    ( cons(esk5_1(X1),X2) = app(X1,X2)
    | ~ singletonP(X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_73]),<A HREF="#c_0_74">c_0_74</A>])).

<A NAME="c_0_99"></A>cnf(c_0_99,negated_conjecture,
    ( hd(esk53_0) = hd(esk50_0)
    | esk53_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>,c_0_89]),c_0_90]),c_0_46])])).

<A NAME="c_0_100"></A>cnf(c_0_100,negated_conjecture,
    ( esk53_0 = nil
    | esk50_0 != nil
    | ~ ssList(cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,c_0_64]),c_0_46])])).

<A NAME="c_0_101"></A>cnf(c_0_101,plain,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>])).

<A NAME="c_0_102"></A>fof(c_0_102,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(cons(X5,X6),X4)
        | X4 = X5
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( X4 != X5
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).

<A NAME="c_0_103"></A>fof(c_0_103,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | ~ memberP(nil,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_93])])).

<A NAME="c_0_104"></A>cnf(c_0_104,plain,
    ( memberP(X1,X2)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | app(X3,cons(X2,X4)) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>])).

<A NAME="c_0_105"></A>cnf(c_0_105,plain,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>])).

<A NAME="c_0_106"></A>cnf(c_0_106,plain,
    ( tl(app(X1,X2)) = app(tl(X1),X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>])).

<A NAME="c_0_107"></A>cnf(c_0_107,negated_conjecture,
    ( app(esk50_0,nil) = esk50_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_98]),c_0_52]),c_0_79])]),<A HREF="#c_0_87">c_0_87</A>])).

<A NAME="c_0_108"></A>cnf(c_0_108,negated_conjecture,
    ( cons(hd(esk50_0),tl(esk53_0)) = esk53_0
    | esk53_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_99]),c_0_46])])).

<A NAME="c_0_109"></A>cnf(c_0_109,negated_conjecture,
    ( hd(esk50_0) = esk55_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>,c_0_76]),c_0_52])]),<A HREF="#c_0_77">c_0_77</A>])).

<A NAME="c_0_110"></A>cnf(c_0_110,negated_conjecture,
    ( esk53_0 = nil
    | esk50_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>,c_0_51]),c_0_45]),c_0_53])])).

<A NAME="c_0_111"></A>cnf(c_0_111,negated_conjecture,
    ( esk53_0 = nil
    | ssItem(hd(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_101">c_0_101</A>,c_0_99]),c_0_46])])).

<A NAME="c_0_112"></A>cnf(c_0_112,plain,
    ( memberP(X3,X1)
    | X1 = X2
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ memberP(cons(X2,X3),X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>])).

<A NAME="c_0_113"></A>cnf(c_0_113,plain,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_103">c_0_103</A>])).

<A NAME="c_0_114"></A>cnf(c_0_114,plain,
    ( memberP(app(X1,cons(X2,X3)),X2)
    | ~ ssList(app(X1,cons(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,plain,
    ( nil = X1
    | ssList(app(tl(X1),X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_105">c_0_105</A>,c_0_106]),c_0_56]),<A HREF="#c_0_91">c_0_91</A>])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( app(tl(esk50_0),nil) = tl(esk50_0)
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>,c_0_107]),c_0_52]),c_0_79])])).

<A NAME="c_0_117"></A>fof(c_0_117,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X5,X6) != app(X5,X4)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).

<A NAME="c_0_118"></A>fof(c_0_118,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).

<A NAME="c_0_119"></A>cnf(c_0_119,negated_conjecture,
    ( cons(esk55_0,tl(esk53_0)) = esk53_0
    | esk53_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>,c_0_109]),<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_120"></A>cnf(c_0_120,negated_conjecture,
    ( esk53_0 = nil
    | ssItem(esk55_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_111">c_0_111</A>,c_0_109]),<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_121"></A>cnf(c_0_121,plain,
    ( X1 = esk5_1(X2)
    | ~ singletonP(X2)
    | ~ memberP(X2,X1)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_73]),c_0_52])]),c_0_74]),<A HREF="#c_0_113">c_0_113</A>])).

<A NAME="c_0_122"></A>cnf(c_0_122,negated_conjecture,
    ( memberP(esk50_0,esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_114">c_0_114</A>,c_0_64]),c_0_79]),c_0_45]),c_0_46]),c_0_53])])).

<A NAME="c_0_123"></A>cnf(c_0_123,negated_conjecture,
    ( cons(esk55_0,tl(esk50_0)) = esk50_0
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_109]),c_0_79])])).

<A NAME="c_0_124"></A>cnf(c_0_124,negated_conjecture,
    ( esk50_0 = nil
    | ssList(tl(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_115">c_0_115</A>,c_0_116]),c_0_52]),c_0_79])])).

<A NAME="c_0_125"></A>cnf(c_0_125,plain,
    ( X1 = X2
    | app(X3,X1) != app(X3,X2)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_117">c_0_117</A>])).

<A NAME="c_0_126"></A>cnf(c_0_126,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_118">c_0_118</A>])).

<A NAME="c_0_127"></A>cnf(c_0_127,negated_conjecture,
    ( hd(esk53_0) = esk55_0
    | esk53_0 = nil
    | ~ ssList(tl(esk53_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>,c_0_119]),<A HREF="#c_0_120">c_0_120</A>])).

<A NAME="c_0_128"></A>cnf(c_0_128,negated_conjecture,
    ( esk5_1(esk50_0) = hd(esk50_0)
    | esk50_0 = nil
    | ~ ssItem(esk5_1(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>,c_0_97]),c_0_52])])).

<A NAME="c_0_129"></A>cnf(c_0_129,negated_conjecture,
    ( esk5_1(esk50_0) = esk52_0
    | ~ singletonP(esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_121">c_0_121</A>,c_0_122]),c_0_79]),c_0_53])])).

<A NAME="c_0_130"></A>cnf(c_0_130,negated_conjecture,
    ( esk50_0 = nil
    | X1 = esk55_0
    | memberP(tl(esk50_0),X1)
    | ~ memberP(esk50_0,X1)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,c_0_123]),c_0_77]),<A HREF="#c_0_124">c_0_124</A>])).

<A NAME="c_0_131"></A>cnf(c_0_131,plain,
    ( X1 = nil
    | app(X2,X1) != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_125">c_0_125</A>,c_0_126]),c_0_52])])).

<A NAME="c_0_132"></A>cnf(c_0_132,negated_conjecture,
    ( hd(esk53_0) = esk55_0
    | esk53_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_127">c_0_127</A>,c_0_105]),c_0_46])])).

<A NAME="c_0_133"></A>cnf(c_0_133,negated_conjecture,
    ( esk5_1(esk50_0) = hd(esk50_0)
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_128">c_0_128</A>,c_0_74]),c_0_79])]),<A HREF="#c_0_87">c_0_87</A>])).

<A NAME="c_0_134"></A>cnf(c_0_134,negated_conjecture,
    ( esk5_1(esk50_0) = esk52_0
    | esk50_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_129">c_0_129</A>,<A HREF="#c_0_87">c_0_87</A>])).

<A NAME="c_0_135"></A>cnf(c_0_135,negated_conjecture,
    ( esk52_0 = esk55_0
    | esk50_0 = nil
    | memberP(tl(esk50_0),esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_130">c_0_130</A>,c_0_122]),c_0_53])])).

<A NAME="c_0_136"></A>cnf(c_0_136,negated_conjecture,
    ( cons(esk52_0,esk54_0) = nil
    | esk53_0 != esk50_0
    | ~ ssList(cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>,c_0_64]),c_0_46])])).

<A NAME="c_0_137"></A>cnf(c_0_137,negated_conjecture,
    ( hd(esk50_0) = esk55_0
    | esk53_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,<A HREF="#c_0_132">c_0_132</A>])).

<A NAME="c_0_138"></A>cnf(c_0_138,negated_conjecture,
    ( hd(esk50_0) = esk52_0
    | esk50_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_133">c_0_133</A>,<A HREF="#c_0_134">c_0_134</A>])).

<A NAME="c_0_139"></A>cnf(c_0_139,negated_conjecture,
    ( esk50_0 = nil
    | esk52_0 = esk55_0
    | memberP(nil,esk52_0) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_135">c_0_135</A>,<A HREF="#c_0_86">c_0_86</A>])).

<A NAME="c_0_140"></A>fof(c_0_140,plain,(
    ! [X6,X7,X8,X9,X10] :
      ( ( ~ equalelemsP(X6)
        | ~ ssItem(X7)
        | ~ ssItem(X8)
        | ~ ssList(X9)
        | ~ ssList(X10)
        | app(X9,cons(X7,cons(X8,X10))) != X6
        | X7 = X8
        | ~ ssList(X6) )
      & ( ssItem(esk40_1(X6))
        | equalelemsP(X6)
        | ~ ssList(X6) )
      & ( ssItem(esk41_1(X6))
        | equalelemsP(X6)
        | ~ ssList(X6) )
      & ( ssList(esk42_1(X6))
        | equalelemsP(X6)
        | ~ ssList(X6) )
      & ( ssList(esk43_1(X6))
        | equalelemsP(X6)
        | ~ ssList(X6) )
      & ( app(esk42_1(X6),cons(esk40_1(X6),cons(esk41_1(X6),esk43_1(X6)))) = X6
        | equalelemsP(X6)
        | ~ ssList(X6) )
      & ( esk40_1(X6) != esk41_1(X6)
        | equalelemsP(X6)
        | ~ ssList(X6) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])])).

<A NAME="c_0_141"></A>cnf(c_0_141,negated_conjecture,
    ( cons(esk52_0,esk54_0) = nil
    | esk53_0 != esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_136">c_0_136</A>,c_0_51]),c_0_45]),c_0_53])])).

<A NAME="c_0_142"></A>cnf(c_0_142,negated_conjecture,
    ( esk53_0 = nil
    | esk52_0 = esk55_0 ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_137">c_0_137</A>,c_0_138]),<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_143"></A>cnf(c_0_143,negated_conjecture,
    ( esk52_0 = esk55_0
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_113">c_0_113</A>,c_0_139]),c_0_53])])).

<A NAME="c_0_144"></A>cnf(c_0_144,plain,
    ( X2 = X3
    | ~ ssList(X1)
    | app(X4,cons(X2,cons(X3,X5))) != X1
    | ~ ssList(X5)
    | ~ ssList(X4)
    | ~ ssItem(X3)
    | ~ ssItem(X2)
    | ~ equalelemsP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_140">c_0_140</A>])).

<A NAME="c_0_145"></A>fof(c_0_145,plain,(
    ! [X4] :
      ( ( ssList(esk44_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( ssItem(esk45_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( cons(esk45_1(X4),esk44_1(X4)) = X4
        | nil = X4
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).

<A NAME="c_0_146"></A>cnf(c_0_146,negated_conjecture,
    ( cons(esk52_0,esk54_0) = nil
    | esk52_0 = esk55_0 ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_141">c_0_141</A>,c_0_142]),<A HREF="#c_0_143">c_0_143</A>])).

<A NAME="c_0_147"></A>cnf(c_0_147,plain,
    ( nil = X2
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>])).

<A NAME="c_0_148"></A>cnf(c_0_148,plain,
    ( X1 = X2
    | ~ equalelemsP(app(X3,cons(X1,cons(X2,X4))))
    | ~ ssList(app(X3,cons(X1,cons(X2,X4))))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssItem(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_149"></A>cnf(c_0_149,plain,
    ( nil = X1
    | cons(esk45_1(X1),esk44_1(X1)) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_145">c_0_145</A>])).

<A NAME="c_0_150"></A>cnf(c_0_150,plain,
    ( nil = X1
    | ssItem(esk45_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_145">c_0_145</A>])).

<A NAME="c_0_151"></A>cnf(c_0_151,plain,
    ( nil = X1
    | ssList(esk44_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_145">c_0_145</A>])).

<A NAME="c_0_152"></A>cnf(c_0_152,negated_conjecture,
    ( esk52_0 = esk55_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,c_0_146]),c_0_45]),c_0_53])])).

<A NAME="c_0_153"></A>fof(c_0_153,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | equalelemsP(cons(X2,nil)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])).

<A NAME="c_0_154"></A>cnf(c_0_154,negated_conjecture,
    ( esk54_0 = nil
    | esk50_0 != nil
    | ~ ssList(app(esk53_0,cons(esk52_0,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_147">c_0_147</A>,c_0_43]),c_0_45])])).

<A NAME="c_0_155"></A>fof(c_0_155,plain,(
    ! [X5,X6,X7,X8] :
      ( ( X5 = X6
        | ~ frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) )
      & ( frontsegP(X7,X8)
        | ~ frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) )
      & ( X5 != X6
        | ~ frontsegP(X7,X8)
        | frontsegP(cons(X5,X7),cons(X6,X8))
        | ~ ssList(X8)
        | ~ ssList(X7)
        | ~ ssItem(X6)
        | ~ ssItem(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])).

<A NAME="c_0_156"></A>cnf(c_0_156,plain,
    ( X1 = esk45_1(X2)
    | nil = X2
    | ~ equalelemsP(app(X3,cons(X1,X2)))
    | ~ ssList(app(X3,cons(X1,X2)))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_148">c_0_148</A>,c_0_149]),c_0_150]),<A HREF="#c_0_151">c_0_151</A>])).

<A NAME="c_0_157"></A>cnf(c_0_157,negated_conjecture,
    ( app(esk53_0,cons(esk55_0,esk54_0)) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,<A HREF="#c_0_152">c_0_152</A>])).

<A NAME="c_0_158"></A>cnf(c_0_158,negated_conjecture,
    ( ssItem(esk55_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,<A HREF="#c_0_152">c_0_152</A>])).

<A NAME="c_0_159"></A>cnf(c_0_159,plain,
    ( equalelemsP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_153">c_0_153</A>])).

<A NAME="c_0_160"></A>cnf(c_0_160,negated_conjecture,
    ( esk54_0 = nil
    | esk50_0 != nil
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_154">c_0_154</A>,c_0_56]),c_0_46])])).

<A NAME="c_0_161"></A>fof(c_0_161,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X6,X5) != app(X4,X5)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).

<A NAME="c_0_162"></A>fof(c_0_162,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])).

<A NAME="c_0_163"></A>cnf(c_0_163,plain,
    ( frontsegP(cons(X1,X3),cons(X2,X4))
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ ssList(X4)
    | ~ frontsegP(X3,X4)
    | X1 != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_155">c_0_155</A>])).

<A NAME="c_0_164"></A>fof(c_0_164,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])).

<A NAME="c_0_165"></A>cnf(c_0_165,negated_conjecture,
    ( esk45_1(esk54_0) = esk55_0
    | esk54_0 = nil
    | ~ equalelemsP(esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_156">c_0_156</A>,c_0_157]),c_0_79]),c_0_46]),c_0_45]),c_0_158])])).

<A NAME="c_0_166"></A>cnf(c_0_166,negated_conjecture,
    ( esk50_0 = nil
    | equalelemsP(esk50_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_159">c_0_159</A>,c_0_76]),<A HREF="#c_0_77">c_0_77</A>])).

<A NAME="c_0_167"></A>cnf(c_0_167,negated_conjecture,
    ( esk54_0 = nil
    | esk50_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_160">c_0_160</A>,c_0_51]),c_0_52]),c_0_53])])).

<A NAME="c_0_168"></A>cnf(c_0_168,plain,
    ( X1 = X2
    | app(X1,X3) != app(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_161">c_0_161</A>])).

<A NAME="c_0_169"></A>cnf(c_0_169,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_162">c_0_162</A>])).

<A NAME="c_0_170"></A>fof(c_0_170,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk8_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ssList(esk9_2(X5,X6))
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6)) = X5
        | ~ segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(app(X9,X6),X10) != X5
        | segmentP(X5,X6)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])])])).

<A NAME="c_0_171"></A>cnf(c_0_171,plain,
    ( frontsegP(cons(X1,X2),cons(X1,X3))
    | ~ frontsegP(X2,X3)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_163">c_0_163</A>])).

<A NAME="c_0_172"></A>cnf(c_0_172,plain,
    ( frontsegP(X1,X1)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_164">c_0_164</A>])).

<A NAME="c_0_173"></A>cnf(c_0_173,negated_conjecture,
    ( esk45_1(esk54_0) = esk55_0
    | esk54_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_165">c_0_165</A>,c_0_166]),<A HREF="#c_0_167">c_0_167</A>])).

<A NAME="c_0_174"></A>cnf(c_0_174,plain,
    ( X1 = nil
    | app(X1,X2) != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_168">c_0_168</A>,c_0_169]),c_0_52])])).

<A NAME="c_0_175"></A>cnf(c_0_175,plain,
    ( segmentP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(app(X3,X2),X4) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_170">c_0_170</A>])).

<A NAME="c_0_176"></A>cnf(c_0_176,plain,
    ( frontsegP(cons(X1,X2),cons(X1,X2))
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_171">c_0_171</A>,<A HREF="#c_0_172">c_0_172</A>])).

<A NAME="c_0_177"></A>cnf(c_0_177,negated_conjecture,
    ( cons(esk55_0,esk44_1(esk54_0)) = esk54_0
    | esk54_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_149">c_0_149</A>,c_0_173]),c_0_45])])).

<A NAME="c_0_178"></A>fof(c_0_178,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ rearsegP(X4,X5)
      | rearsegP(app(X6,X4),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).

<A NAME="c_0_179"></A>fof(c_0_179,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | rearsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])).

<A NAME="c_0_180"></A>fof(c_0_180,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ frontsegP(X4,X5)
      | frontsegP(app(X4,X6),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])).

<A NAME="c_0_181"></A>fof(c_0_181,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])).

<A NAME="c_0_182"></A>cnf(c_0_182,plain,
    ( cons(hd(X1),tl(app(X1,X2))) = app(X1,X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_88]),c_0_56]),<A HREF="#c_0_91">c_0_91</A>])).

<A NAME="c_0_183"></A>cnf(c_0_183,negated_conjecture,
    ( esk53_0 = nil
    | app(cons(esk52_0,nil),esk54_0) != esk50_0
    | ~ ssList(app(cons(esk52_0,nil),esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_174">c_0_174</A>,c_0_57]),c_0_46])])).

<A NAME="c_0_184"></A>cnf(c_0_184,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_175">c_0_175</A>])).

<A NAME="c_0_185"></A>cnf(c_0_185,negated_conjecture,
    ( esk54_0 = nil
    | frontsegP(esk54_0,esk54_0)
    | ~ ssList(esk44_1(esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_176">c_0_176</A>,c_0_177]),c_0_158])])).

<A NAME="c_0_186"></A>cnf(c_0_186,plain,
    ( rearsegP(app(X1,X2),X3)
    | ~ rearsegP(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_178">c_0_178</A>])).

<A NAME="c_0_187"></A>cnf(c_0_187,plain,
    ( rearsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_179">c_0_179</A>])).

<A NAME="c_0_188"></A>cnf(c_0_188,plain,
    ( frontsegP(app(X1,X2),X3)
    | ~ frontsegP(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_180">c_0_180</A>])).

<A NAME="c_0_189"></A>cnf(c_0_189,plain,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_181">c_0_181</A>])).

<A NAME="c_0_190"></A>fof(c_0_190,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk7_2(X4,X5))
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(esk7_2(X4,X5),X5) = X4
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X7,X5) != X4
        | rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])])])])).

<A NAME="c_0_191"></A>cnf(c_0_191,negated_conjecture,
    ( cons(hd(esk53_0),tl(esk50_0)) = esk50_0
    | esk53_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_182">c_0_182</A>,c_0_89]),c_0_90]),c_0_46])])).

<A NAME="c_0_192"></A>cnf(c_0_192,negated_conjecture,
    ( esk53_0 = nil
    | app(cons(esk52_0,nil),esk54_0) != esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_183">c_0_183</A>,c_0_56]),c_0_45])])).

<A NAME="c_0_193"></A>cnf(c_0_193,plain,
    ( segmentP(app(X1,X2),X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_184">c_0_184</A>,c_0_126]),c_0_52])]),<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_194"></A>fof(c_0_194,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssItem(X6)
      | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])).

<A NAME="c_0_195"></A>cnf(c_0_195,negated_conjecture,
    ( esk54_0 = nil
    | frontsegP(esk54_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_185">c_0_185</A>,c_0_151]),c_0_45])])).

<A NAME="c_0_196"></A>cnf(c_0_196,plain,
    ( rearsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_186">c_0_186</A>,c_0_187]),c_0_52])])).

<A NAME="c_0_197"></A>cnf(c_0_197,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_188">c_0_188</A>,c_0_189]),c_0_52])])).

<A NAME="c_0_198"></A>fof(c_0_198,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ~ rearsegP(X3,X4)
      | ~ rearsegP(X4,X3)
      | X3 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).

<A NAME="c_0_199"></A>cnf(c_0_199,plain,
    ( rearsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X3,X2) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_190">c_0_190</A>])).

<A NAME="c_0_200"></A>cnf(c_0_200,negated_conjecture,
    ( cons(esk55_0,tl(esk50_0)) = esk50_0
    | esk53_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_191">c_0_191</A>,<A HREF="#c_0_132">c_0_132</A>])).

<A NAME="c_0_201"></A>cnf(c_0_201,negated_conjecture,
    ( esk53_0 = nil
    | app(cons(esk52_0,nil),esk54_0) != esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_192">c_0_192</A>,c_0_51]),c_0_52]),c_0_53])])).

<A NAME="c_0_202"></A>fof(c_0_202,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ~ segmentP(X3,X4)
      | ~ segmentP(X4,X3)
      | X3 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])).

<A NAME="c_0_203"></A>cnf(c_0_203,negated_conjecture,
    ( segmentP(esk50_0,cons(esk55_0,esk54_0))
    | ~ ssList(cons(esk55_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_193">c_0_193</A>,c_0_157]),c_0_46])])).

<A NAME="c_0_204"></A>cnf(c_0_204,plain,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_194">c_0_194</A>])).

<A NAME="c_0_205"></A>cnf(c_0_205,plain,
    ( app(nil,app(X1,X2)) = app(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_169]),c_0_52])])).

<A NAME="c_0_206"></A>cnf(c_0_206,negated_conjecture,
    ( app(esk54_0,esk6_2(esk54_0,esk54_0)) = esk54_0
    | esk54_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_195]),c_0_45])])).

<A NAME="c_0_207"></A>cnf(c_0_207,negated_conjecture,
    ( esk54_0 = nil
    | ssList(esk6_2(esk54_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_195]),c_0_45])])).

<A NAME="c_0_208"></A>cnf(c_0_208,negated_conjecture,
    ( rearsegP(esk50_0,nil)
    | ~ ssList(cons(esk55_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_196">c_0_196</A>,c_0_157]),c_0_46])])).

<A NAME="c_0_209"></A>cnf(c_0_209,negated_conjecture,
    ( frontsegP(esk50_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_197">c_0_197</A>,c_0_89]),c_0_90]),c_0_46])])).

<A NAME="c_0_210"></A>cnf(c_0_210,plain,
    ( X1 = X2
    | ~ rearsegP(X2,X1)
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_198">c_0_198</A>])).

<A NAME="c_0_211"></A>cnf(c_0_211,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_199]),<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_212"></A>cnf(c_0_212,negated_conjecture,
    ( app(esk53_0,app(cons(esk55_0,nil),esk54_0)) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,<A HREF="#c_0_152">c_0_152</A>])).

<A NAME="c_0_213"></A>cnf(c_0_213,negated_conjecture,
    ( cons(esk55_0,nil) = esk50_0
    | esk53_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_200">c_0_200</A>,c_0_86]),<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_214"></A>cnf(c_0_214,negated_conjecture,
    ( esk53_0 = nil
    | app(cons(esk55_0,nil),esk54_0) != esk50_0 ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_201">c_0_201</A>,c_0_143]),<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_215"></A>cnf(c_0_215,plain,
    ( X1 = X2
    | ~ segmentP(X2,X1)
    | ~ segmentP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_202">c_0_202</A>])).

<A NAME="c_0_216"></A>cnf(c_0_216,negated_conjecture,
    ( segmentP(esk50_0,cons(esk55_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_203">c_0_203</A>,c_0_51]),c_0_45]),c_0_158])])).

<A NAME="c_0_217"></A>cnf(c_0_217,negated_conjecture,
    ( cons(esk55_0,app(nil,X1)) = app(esk50_0,X1)
    | esk50_0 = nil
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_204">c_0_204</A>,c_0_76]),c_0_52])]),c_0_158])])).

<A NAME="c_0_218"></A>cnf(c_0_218,negated_conjecture,
    ( app(nil,esk54_0) = esk54_0
    | esk54_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_205">c_0_205</A>,c_0_206]),c_0_45])]),<A HREF="#c_0_207">c_0_207</A>])).

<A NAME="c_0_219"></A>cnf(c_0_219,plain,
    ( app(esk7_2(X1,X2),X2) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_190">c_0_190</A>])).

<A NAME="c_0_220"></A>cnf(c_0_220,negated_conjecture,
    ( rearsegP(esk50_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_208">c_0_208</A>,c_0_51]),c_0_45]),c_0_158])])).

<A NAME="c_0_221"></A>cnf(c_0_221,plain,
    ( ssList(esk7_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_190">c_0_190</A>])).

<A NAME="c_0_222"></A>cnf(c_0_222,negated_conjecture,
    ( app(nil,esk6_2(esk50_0,nil)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_209]),c_0_52]),c_0_79])])).

<A NAME="c_0_223"></A>cnf(c_0_223,negated_conjecture,
    ( ssList(esk6_2(esk50_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_83">c_0_83</A>,c_0_209]),c_0_52]),c_0_79])])).

<A NAME="c_0_224"></A>cnf(c_0_224,plain,
    ( app(X1,X2) = X2
    | ~ rearsegP(X2,app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_210">c_0_210</A>,c_0_211]),<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_225"></A>cnf(c_0_225,negated_conjecture,
    ( app(esk53_0,app(esk50_0,esk54_0)) = esk50_0
    | esk53_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_212">c_0_212</A>,<A HREF="#c_0_213">c_0_213</A>])).

<A NAME="c_0_226"></A>cnf(c_0_226,negated_conjecture,
    ( esk53_0 = nil
    | app(esk50_0,esk54_0) != esk50_0 ),
    inference(spm,[status(thm)],[<A HREF="#c_0_214">c_0_214</A>,<A HREF="#c_0_213">c_0_213</A>])).

<A NAME="c_0_227"></A>cnf(c_0_227,negated_conjecture,
    ( cons(esk55_0,esk54_0) = esk50_0
    | ~ segmentP(cons(esk55_0,esk54_0),esk50_0)
    | ~ ssList(cons(esk55_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_215">c_0_215</A>,c_0_216]),c_0_79])])).

<A NAME="c_0_228"></A>cnf(c_0_228,plain,
    ( segmentP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_184">c_0_184</A>,c_0_169]),c_0_52])]),<A HREF="#c_0_56">c_0_56</A>])).

<A NAME="c_0_229"></A>cnf(c_0_229,negated_conjecture,
    ( app(esk50_0,esk54_0) = cons(esk55_0,esk54_0)
    | esk54_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_217">c_0_217</A>,c_0_218]),c_0_45])]),<A HREF="#c_0_167">c_0_167</A>])).

<A NAME="c_0_230"></A>cnf(c_0_230,negated_conjecture,
    ( app(esk7_2(esk50_0,nil),nil) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_219">c_0_219</A>,c_0_220]),c_0_52]),c_0_79])])).

<A NAME="c_0_231"></A>cnf(c_0_231,negated_conjecture,
    ( ssList(esk7_2(esk50_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_221">c_0_221</A>,c_0_220]),c_0_52]),c_0_79])])).

<A NAME="c_0_232"></A>cnf(c_0_232,negated_conjecture,
    ( esk6_2(esk50_0,nil) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_169">c_0_169</A>,c_0_222]),c_0_223])])).

<A NAME="c_0_233"></A>cnf(c_0_233,negated_conjecture,
    ( esk53_0 = nil
    | ~ rearsegP(app(esk50_0,esk54_0),esk50_0)
    | ~ ssList(app(esk50_0,esk54_0)) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_224">c_0_224</A>,c_0_225]),c_0_46])]),<A HREF="#c_0_226">c_0_226</A>])).

<A NAME="c_0_234"></A>cnf(c_0_234,negated_conjecture,
    ( cons(esk55_0,esk54_0) = esk50_0
    | ~ segmentP(cons(esk55_0,esk54_0),esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_227">c_0_227</A>,c_0_51]),c_0_45]),c_0_158])])).

<A NAME="c_0_235"></A>cnf(c_0_235,negated_conjecture,
    ( esk54_0 = nil
    | segmentP(cons(esk55_0,esk54_0),esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_228">c_0_228</A>,c_0_229]),c_0_45]),c_0_79])])).

<A NAME="c_0_236"></A>cnf(c_0_236,negated_conjecture,
    ( esk54_0 = nil
    | cons(esk55_0,esk54_0) != esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>,c_0_229]),c_0_79]),c_0_45])])).

<A NAME="c_0_237"></A>cnf(c_0_237,negated_conjecture,
    ( esk7_2(esk50_0,nil) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_126">c_0_126</A>,c_0_230]),c_0_231])])).

<A NAME="c_0_238"></A>cnf(c_0_238,negated_conjecture,
    ( app(nil,esk50_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_222">c_0_222</A>,<A HREF="#c_0_232">c_0_232</A>])).

<A NAME="c_0_239"></A>cnf(c_0_239,negated_conjecture,
    ( memberP(esk54_0,esk52_0)
    | memberP(esk53_0,esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_240"></A>cnf(c_0_240,negated_conjecture,
    ( esk53_0 = nil
    | ~ rearsegP(app(esk50_0,esk54_0),esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_233">c_0_233</A>,c_0_56]),c_0_45]),c_0_79])])).

<A NAME="c_0_241"></A>cnf(c_0_241,negated_conjecture,
    ( esk54_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_234">c_0_234</A>,c_0_235]),<A HREF="#c_0_236">c_0_236</A>])).

<A NAME="c_0_242"></A>cnf(c_0_242,negated_conjecture,
    ( app(esk50_0,nil) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_230">c_0_230</A>,<A HREF="#c_0_237">c_0_237</A>])).

<A NAME="c_0_243"></A>cnf(c_0_243,negated_conjecture,
    ( rearsegP(esk50_0,esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_211">c_0_211</A>,c_0_238]),c_0_52]),c_0_79])])).

<A NAME="c_0_244"></A>cnf(c_0_244,negated_conjecture,
    ( memberP(esk54_0,esk55_0)
    | memberP(esk53_0,esk55_0) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_239">c_0_239</A>,c_0_152]),<A HREF="#c_0_152">c_0_152</A>])).

<A NAME="c_0_245"></A>cnf(c_0_245,negated_conjecture,
    ( esk53_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_240">c_0_240</A>,c_0_241]),c_0_242]),c_0_243])])).

<A NAME="c_0_246"></A>cnf(c_0_246,negated_conjecture,
    ( memberP(nil,esk55_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_244">c_0_244</A>,c_0_241]),c_0_245])])).

<A NAME="c_0_247"></A>cnf(c_0_247,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_113">c_0_113</A>,c_0_246]),c_0_158])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC184+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.22  % Computer   : n156.star.cs.uiowa.edu
% 0.03/0.22  % Model      : x86_64 x86_64
% 0.03/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.22  % Memory     : 32218.75MB
% 0.03/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.22  % CPULimit   : 300
% 0.03/0.22  % DateTime   : Tue May 10 20:36:38 CDT 2016
% 0.03/0.22  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.06/0.70  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.06/0.70  # and selection function SelectCQPrecWNTNp.
% 0.06/0.70  #
% 0.06/0.70  # Preprocessing time       : 0.014 s
% 0.06/0.70  # Presaturation interreduction done
% 0.06/0.70  
% 0.06/0.70  # Proof found!
% 0.06/0.70  # SZS status Theorem
% 0.06/0.70  # SZS output start CNFRefutation
% 0.06/0.70  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~(ssList(X4))|X2!=X4)|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssList(X6)=>![X7]:((~(ssList(X7))|app(app(X6,cons(X5,nil)),X7)!=X1)|(~(memberP(X6,X5))&~(memberP(X7,X5)))))))|(![X8]:(((~(ssItem(X8))|cons(X8,nil)!=X3)|~(memberP(X4,X8)))|?[X9]:(((ssItem(X9)&X8!=X9)&memberP(X4,X9))&leq(X9,X8)))&(nil!=X4|nil!=X3))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.06/0.70  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.06/0.70  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.06/0.70  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax5)).
% 0.06/0.70  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.06/0.70  fof(c_0_5, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.06/0.70  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.06/0.70  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax78)).
% 0.06/0.70  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax23)).
% 0.06/0.70  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>nil!=cons(X2,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax21)).
% 0.06/0.70  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)&lt;=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax4)).
% 0.06/0.70  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
% 0.06/0.70  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>hd(app(X1,X2))=hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax85)).
% 0.06/0.70  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.06/0.70  fof(c_0_14, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax22)).
% 0.06/0.70  fof(c_0_15, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax38)).
% 0.06/0.70  fof(c_0_16, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>(memberP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(X3,cons(X2,X4))=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax3)).
% 0.06/0.70  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
% 0.06/0.70  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax86)).
% 0.06/0.70  fof(c_0_19, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)&lt;=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax37)).
% 0.06/0.70  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X2,X3)=app(X2,X1)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax80)).
% 0.06/0.70  fof(c_0_21, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
% 0.06/0.70  fof(c_0_22, axiom, (![X1]:(ssList(X1)=>(equalelemsP(X1)&lt;=>![X2]:(ssItem(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssList(X4)=>![X5]:(ssList(X5)=>(app(X4,cons(X2,cons(X3,X5)))=X1=>X2=X3)))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax14)).
% 0.06/0.70  fof(c_0_23, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax20)).
% 0.06/0.70  fof(c_0_24, axiom, (![X1]:(ssItem(X1)=>equalelemsP(cons(X1,nil)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax73)).
% 0.06/0.70  fof(c_0_25, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(frontsegP(cons(X1,X3),cons(X2,X4))&lt;=>(X1=X2&frontsegP(X3,X4)))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax44)).
% 0.06/0.70  fof(c_0_26, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax79)).
% 0.06/0.70  fof(c_0_27, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
% 0.06/0.70  fof(c_0_28, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax42)).
% 0.06/0.70  fof(c_0_29, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(segmentP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(app(X3,X2),X4)=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax7)).
% 0.06/0.70  fof(c_0_30, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(rearsegP(X1,X2)=>rearsegP(app(X3,X1),X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax50)).
% 0.06/0.70  fof(c_0_31, axiom, (![X1]:(ssList(X1)=>rearsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax51)).
% 0.06/0.70  fof(c_0_32, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(frontsegP(X1,X2)=>frontsegP(app(X1,X3),X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax43)).
% 0.06/0.70  fof(c_0_33, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax45)).
% 0.06/0.70  fof(c_0_34, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax6)).
% 0.06/0.70  fof(c_0_35, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
% 0.06/0.70  fof(c_0_36, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>((rearsegP(X1,X2)&rearsegP(X2,X1))=>X1=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax48)).
% 0.06/0.70  fof(c_0_37, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>((segmentP(X1,X2)&segmentP(X2,X1))=>X1=X2)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax54)).
% 0.06/0.70  fof(c_0_38, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:((((~ssList(X4)|X2!=X4)|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssList(X6)=>![X7]:((~ssList(X7)|app(app(X6,cons(X5,nil)),X7)!=X1)|(~memberP(X6,X5)&~memberP(X7,X5))))))|(![X8]:(((~ssItem(X8)|cons(X8,nil)!=X3)|~memberP(X4,X8))|?[X9]:(((ssItem(X9)&X8!=X9)&memberP(X4,X9))&leq(X9,X8)))&(nil!=X4|nil!=X3)))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.06/0.70  fof(c_0_39, negated_conjecture, (![X18]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&((((ssList(esk51_0)&esk49_0=esk51_0)&esk48_0=esk50_0)&(ssItem(esk52_0)&(ssList(esk53_0)&((ssList(esk54_0)&app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk48_0)&(memberP(esk53_0,esk52_0)|memberP(esk54_0,esk52_0))))))&(((((nil=esk51_0|ssItem(esk55_0))&(nil=esk50_0|ssItem(esk55_0)))&((nil=esk51_0|cons(esk55_0,nil)=esk50_0)&(nil=esk50_0|cons(esk55_0,nil)=esk50_0)))&((nil=esk51_0|memberP(esk51_0,esk55_0))&(nil=esk50_0|memberP(esk51_0,esk55_0))))&((nil=esk51_0|(((~ssItem(X18)|esk55_0=X18)|~memberP(esk51_0,X18))|~leq(X18,esk55_0)))&(nil=esk50_0|(((~ssItem(X18)|esk55_0=X18)|~memberP(esk51_0,X18))|~leq(X18,esk55_0)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_38])])])])])).
% 0.06/0.70  cnf(c_0_40, negated_conjecture, (app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.70  cnf(c_0_41, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.70  fof(c_0_42, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.06/0.70  cnf(c_0_43, negated_conjecture, (app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk50_0), inference(rw,[status(thm)],[c_0_40, c_0_41])).
% 0.06/0.70  cnf(c_0_44, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
% 0.06/0.70  cnf(c_0_45, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.70  cnf(c_0_46, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.70  fof(c_0_47, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.06/0.70  fof(c_0_48, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
% 0.06/0.70  fof(c_0_49, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 0.06/0.70  cnf(c_0_50, negated_conjecture, (app(esk53_0,app(cons(esk52_0,nil),esk54_0))=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46])])).
% 0.06/0.70  cnf(c_0_51, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
% 0.06/0.70  cnf(c_0_52, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.06/0.70  cnf(c_0_53, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.70  fof(c_0_54, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
% 0.06/0.70  cnf(c_0_55, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
% 0.06/0.70  cnf(c_0_56, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
% 0.06/0.70  cnf(c_0_57, negated_conjecture, (app(esk53_0,app(cons(esk52_0,nil),esk54_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52]), c_0_53])])).
% 0.06/0.70  cnf(c_0_58, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
% 0.06/0.70  fof(c_0_59, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
% 0.06/0.70  fof(c_0_60, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
% 0.06/0.70  fof(c_0_61, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|nil!=cons(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 0.06/0.70  fof(c_0_62, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).
% 0.06/0.70  cnf(c_0_63, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_55]), c_0_56])).
% 0.06/0.70  cnf(c_0_64, negated_conjecture, (app(esk53_0,cons(esk52_0,esk54_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_45]), c_0_53])])).
% 0.06/0.70  cnf(c_0_65, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
% 0.06/0.70  cnf(c_0_66, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
% 0.06/0.70  cnf(c_0_67, plain, (nil!=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
% 0.06/0.70  fof(c_0_68, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
% 0.06/0.70  cnf(c_0_69, plain, (singletonP(X1)|~ssList(X1)|cons(X2,nil)!=X1|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
% 0.06/0.70  cnf(c_0_70, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.70  cnf(c_0_71, negated_conjecture, (frontsegP(esk50_0,esk53_0)|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_46])])).
% 0.06/0.70  cnf(c_0_72, plain, (cons(X1,tl(cons(X1,X2)))=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_51]), c_0_67])).
% 0.06/0.70  cnf(c_0_73, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
% 0.06/0.70  cnf(c_0_74, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
% 0.06/0.70  cnf(c_0_75, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
% 0.06/0.70  cnf(c_0_76, negated_conjecture, (cons(esk55_0,nil)=esk50_0|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.70  cnf(c_0_77, negated_conjecture, (ssItem(esk55_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.70  cnf(c_0_78, plain, (singletonP(cons(X1,nil))|~ssList(cons(X1,nil))|~ssItem(X1)), inference(er,[status(thm)],[c_0_69])).
% 0.06/0.70  cnf(c_0_79, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_70, c_0_41])).
% 0.06/0.70  fof(c_0_80, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|hd(app(X3,X4))=hd(X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
% 0.06/0.70  cnf(c_0_81, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
% 0.06/0.70  cnf(c_0_82, negated_conjecture, (frontsegP(esk50_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_51]), c_0_45]), c_0_53])])).
% 0.06/0.70  cnf(c_0_83, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
% 0.06/0.70  fof(c_0_84, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).
% 0.06/0.70  cnf(c_0_85, plain, (cons(esk5_1(X1),tl(X1))=X1|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_52])]), c_0_74])).
% 0.06/0.70  cnf(c_0_86, negated_conjecture, (tl(esk50_0)=nil|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_52])]), c_0_77])).
% 0.06/0.70  cnf(c_0_87, negated_conjecture, (esk50_0=nil|singletonP(esk50_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_76]), c_0_79])]), c_0_77])).
% 0.06/0.70  cnf(c_0_88, plain, (hd(app(X1,X2))=hd(X1)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
% 0.06/0.70  cnf(c_0_89, negated_conjecture, (app(esk53_0,esk6_2(esk50_0,esk53_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_46]), c_0_79])])).
% 0.06/0.70  cnf(c_0_90, negated_conjecture, (ssList(esk6_2(esk50_0,esk53_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_82]), c_0_46]), c_0_79])])).
% 0.06/0.70  cnf(c_0_91, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
% 0.06/0.70  fof(c_0_92, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])).
% 0.06/0.70  fof(c_0_93, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_15])).
% 0.06/0.70  fof(c_0_94, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk3_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&((((ssList(esk4_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&(((app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6)))=X5|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(X9,cons(X6,X10))!=X5))|memberP(X5,X6))|~ssItem(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])])).
% 0.06/0.70  fof(c_0_95, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).
% 0.06/0.70  fof(c_0_96, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
% 0.06/0.70  cnf(c_0_97, negated_conjecture, (cons(esk5_1(esk50_0),nil)=esk50_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_79])]), c_0_87])).
% 0.06/0.70  cnf(c_0_98, plain, (cons(esk5_1(X1),X2)=app(X1,X2)|~singletonP(X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_73]), c_0_74])).
% 0.06/0.70  cnf(c_0_99, negated_conjecture, (hd(esk53_0)=hd(esk50_0)|esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90]), c_0_46])])).
% 0.06/0.70  cnf(c_0_100, negated_conjecture, (esk53_0=nil|esk50_0!=nil|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_64]), c_0_46])])).
% 0.06/0.70  cnf(c_0_101, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
% 0.06/0.70  fof(c_0_102, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
% 0.06/0.70  fof(c_0_103, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_93])])).
% 0.06/0.70  cnf(c_0_104, plain, (memberP(X1,X2)|~ssList(X1)|~ssItem(X2)|app(X3,cons(X2,X4))!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_94])).
% 0.06/0.70  cnf(c_0_105, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
% 0.06/0.70  cnf(c_0_106, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
% 0.06/0.70  cnf(c_0_107, negated_conjecture, (app(esk50_0,nil)=esk50_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_52]), c_0_79])]), c_0_87])).
% 0.06/0.70  cnf(c_0_108, negated_conjecture, (cons(hd(esk50_0),tl(esk53_0))=esk53_0|esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_99]), c_0_46])])).
% 0.06/0.70  cnf(c_0_109, negated_conjecture, (hd(esk50_0)=esk55_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_76]), c_0_52])]), c_0_77])).
% 0.06/0.70  cnf(c_0_110, negated_conjecture, (esk53_0=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_51]), c_0_45]), c_0_53])])).
% 0.06/0.70  cnf(c_0_111, negated_conjecture, (esk53_0=nil|ssItem(hd(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_99]), c_0_46])])).
% 0.06/0.70  cnf(c_0_112, plain, (memberP(X3,X1)|X1=X2|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~memberP(cons(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
% 0.06/0.70  cnf(c_0_113, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
% 0.06/0.70  cnf(c_0_114, plain, (memberP(app(X1,cons(X2,X3)),X2)|~ssList(app(X1,cons(X2,X3)))|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(er,[status(thm)],[c_0_104])).
% 0.06/0.70  cnf(c_0_115, plain, (nil=X1|ssList(app(tl(X1),X2))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_56]), c_0_91])).
% 0.06/0.70  cnf(c_0_116, negated_conjecture, (app(tl(esk50_0),nil)=tl(esk50_0)|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_52]), c_0_79])])).
% 0.06/0.70  fof(c_0_117, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X5,X6)!=app(X5,X4)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
% 0.06/0.70  fof(c_0_118, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).
% 0.06/0.70  cnf(c_0_119, negated_conjecture, (cons(esk55_0,tl(esk53_0))=esk53_0|esk53_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_110])).
% 0.06/0.70  cnf(c_0_120, negated_conjecture, (esk53_0=nil|ssItem(esk55_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_109]), c_0_110])).
% 0.06/0.70  cnf(c_0_121, plain, (X1=esk5_1(X2)|~singletonP(X2)|~memberP(X2,X1)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_73]), c_0_52])]), c_0_74]), c_0_113])).
% 0.06/0.70  cnf(c_0_122, negated_conjecture, (memberP(esk50_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_64]), c_0_79]), c_0_45]), c_0_46]), c_0_53])])).
% 0.06/0.70  cnf(c_0_123, negated_conjecture, (cons(esk55_0,tl(esk50_0))=esk50_0|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_109]), c_0_79])])).
% 0.06/0.70  cnf(c_0_124, negated_conjecture, (esk50_0=nil|ssList(tl(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_52]), c_0_79])])).
% 0.06/0.70  cnf(c_0_125, plain, (X1=X2|app(X3,X1)!=app(X3,X2)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_117])).
% 0.06/0.70  cnf(c_0_126, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
% 0.06/0.70  cnf(c_0_127, negated_conjecture, (hd(esk53_0)=esk55_0|esk53_0=nil|~ssList(tl(esk53_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_119]), c_0_120])).
% 0.06/0.70  cnf(c_0_128, negated_conjecture, (esk5_1(esk50_0)=hd(esk50_0)|esk50_0=nil|~ssItem(esk5_1(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_97]), c_0_52])])).
% 0.06/0.70  cnf(c_0_129, negated_conjecture, (esk5_1(esk50_0)=esk52_0|~singletonP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_79]), c_0_53])])).
% 0.06/0.70  cnf(c_0_130, negated_conjecture, (esk50_0=nil|X1=esk55_0|memberP(tl(esk50_0),X1)|~memberP(esk50_0,X1)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_123]), c_0_77]), c_0_124])).
% 0.06/0.70  cnf(c_0_131, plain, (X1=nil|app(X2,X1)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_52])])).
% 0.06/0.70  cnf(c_0_132, negated_conjecture, (hd(esk53_0)=esk55_0|esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_105]), c_0_46])])).
% 0.06/0.70  cnf(c_0_133, negated_conjecture, (esk5_1(esk50_0)=hd(esk50_0)|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_74]), c_0_79])]), c_0_87])).
% 0.06/0.70  cnf(c_0_134, negated_conjecture, (esk5_1(esk50_0)=esk52_0|esk50_0=nil), inference(spm,[status(thm)],[c_0_129, c_0_87])).
% 0.06/0.70  cnf(c_0_135, negated_conjecture, (esk52_0=esk55_0|esk50_0=nil|memberP(tl(esk50_0),esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_122]), c_0_53])])).
% 0.06/0.70  cnf(c_0_136, negated_conjecture, (cons(esk52_0,esk54_0)=nil|esk53_0!=esk50_0|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_64]), c_0_46])])).
% 0.06/0.70  cnf(c_0_137, negated_conjecture, (hd(esk50_0)=esk55_0|esk53_0=nil), inference(spm,[status(thm)],[c_0_99, c_0_132])).
% 0.06/0.70  cnf(c_0_138, negated_conjecture, (hd(esk50_0)=esk52_0|esk50_0=nil), inference(spm,[status(thm)],[c_0_133, c_0_134])).
% 0.06/0.70  cnf(c_0_139, negated_conjecture, (esk50_0=nil|esk52_0=esk55_0|memberP(nil,esk52_0)), inference(spm,[status(thm)],[c_0_135, c_0_86])).
% 0.06/0.70  fof(c_0_140, plain, (![X6]:![X7]:![X8]:![X9]:![X10]:(((~equalelemsP(X6)|(~ssItem(X7)|(~ssItem(X8)|(~ssList(X9)|(~ssList(X10)|(app(X9,cons(X7,cons(X8,X10)))!=X6|X7=X8))))))|~ssList(X6))&(((ssItem(esk40_1(X6))|equalelemsP(X6))|~ssList(X6))&(((ssItem(esk41_1(X6))|equalelemsP(X6))|~ssList(X6))&(((ssList(esk42_1(X6))|equalelemsP(X6))|~ssList(X6))&(((ssList(esk43_1(X6))|equalelemsP(X6))|~ssList(X6))&(((app(esk42_1(X6),cons(esk40_1(X6),cons(esk41_1(X6),esk43_1(X6))))=X6|equalelemsP(X6))|~ssList(X6))&((esk40_1(X6)!=esk41_1(X6)|equalelemsP(X6))|~ssList(X6))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])])).
% 0.06/0.70  cnf(c_0_141, negated_conjecture, (cons(esk52_0,esk54_0)=nil|esk53_0!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_51]), c_0_45]), c_0_53])])).
% 0.06/0.70  cnf(c_0_142, negated_conjecture, (esk53_0=nil|esk52_0=esk55_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_110])).
% 0.06/0.70  cnf(c_0_143, negated_conjecture, (esk52_0=esk55_0|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_139]), c_0_53])])).
% 0.06/0.70  cnf(c_0_144, plain, (X2=X3|~ssList(X1)|app(X4,cons(X2,cons(X3,X5)))!=X1|~ssList(X5)|~ssList(X4)|~ssItem(X3)|~ssItem(X2)|~equalelemsP(X1)), inference(split_conjunct,[status(thm)],[c_0_140])).
% 0.06/0.70  fof(c_0_145, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])])).
% 0.06/0.70  cnf(c_0_146, negated_conjecture, (cons(esk52_0,esk54_0)=nil|esk52_0=esk55_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142]), c_0_143])).
% 0.06/0.70  cnf(c_0_147, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
% 0.06/0.70  cnf(c_0_148, plain, (X1=X2|~equalelemsP(app(X3,cons(X1,cons(X2,X4))))|~ssList(app(X3,cons(X1,cons(X2,X4))))|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_144])).
% 0.06/0.70  cnf(c_0_149, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_145])).
% 0.06/0.70  cnf(c_0_150, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_145])).
% 0.06/0.70  cnf(c_0_151, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_145])).
% 0.06/0.70  cnf(c_0_152, negated_conjecture, (esk52_0=esk55_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_146]), c_0_45]), c_0_53])])).
% 0.06/0.70  fof(c_0_153, plain, (![X2]:(~ssItem(X2)|equalelemsP(cons(X2,nil)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])).
% 0.06/0.70  cnf(c_0_154, negated_conjecture, (esk54_0=nil|esk50_0!=nil|~ssList(app(esk53_0,cons(esk52_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_43]), c_0_45])])).
% 0.06/0.70  fof(c_0_155, plain, (![X5]:![X6]:![X7]:![X8]:(((((((X5=X6|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5))&(((((frontsegP(X7,X8)|~frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))&((((((X5!=X6|~frontsegP(X7,X8))|frontsegP(cons(X5,X7),cons(X6,X8)))|~ssList(X8))|~ssList(X7))|~ssItem(X6))|~ssItem(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])).
% 0.06/0.70  cnf(c_0_156, plain, (X1=esk45_1(X2)|nil=X2|~equalelemsP(app(X3,cons(X1,X2)))|~ssList(app(X3,cons(X1,X2)))|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_150]), c_0_151])).
% 0.06/0.70  cnf(c_0_157, negated_conjecture, (app(esk53_0,cons(esk55_0,esk54_0))=esk50_0), inference(rw,[status(thm)],[c_0_64, c_0_152])).
% 0.06/0.70  cnf(c_0_158, negated_conjecture, (ssItem(esk55_0)), inference(rw,[status(thm)],[c_0_53, c_0_152])).
% 0.06/0.70  cnf(c_0_159, plain, (equalelemsP(cons(X1,nil))|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_153])).
% 0.06/0.70  cnf(c_0_160, negated_conjecture, (esk54_0=nil|esk50_0!=nil|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_56]), c_0_46])])).
% 0.06/0.70  fof(c_0_161, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
% 0.06/0.70  fof(c_0_162, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])).
% 0.06/0.70  cnf(c_0_163, plain, (frontsegP(cons(X1,X3),cons(X2,X4))|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~ssList(X4)|~frontsegP(X3,X4)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_155])).
% 0.06/0.70  fof(c_0_164, plain, (![X2]:(~ssList(X2)|frontsegP(X2,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])).
% 0.06/0.70  cnf(c_0_165, negated_conjecture, (esk45_1(esk54_0)=esk55_0|esk54_0=nil|~equalelemsP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_79]), c_0_46]), c_0_45]), c_0_158])])).
% 0.06/0.70  cnf(c_0_166, negated_conjecture, (esk50_0=nil|equalelemsP(esk50_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_76]), c_0_77])).
% 0.06/0.70  cnf(c_0_167, negated_conjecture, (esk54_0=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_51]), c_0_52]), c_0_53])])).
% 0.06/0.70  cnf(c_0_168, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_161])).
% 0.06/0.70  cnf(c_0_169, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_162])).
% 0.06/0.70  fof(c_0_170, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk8_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&((((ssList(esk9_2(X5,X6))|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))&(((app(app(esk8_2(X5,X6),X6),esk9_2(X5,X6))=X5|~segmentP(X5,X6))|~ssList(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(app(X9,X6),X10)!=X5))|segmentP(X5,X6))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])])])).
% 0.06/0.70  cnf(c_0_171, plain, (frontsegP(cons(X1,X2),cons(X1,X3))|~frontsegP(X2,X3)|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_163])).
% 0.06/0.70  cnf(c_0_172, plain, (frontsegP(X1,X1)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_164])).
% 0.06/0.70  cnf(c_0_173, negated_conjecture, (esk45_1(esk54_0)=esk55_0|esk54_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_167])).
% 0.06/0.70  cnf(c_0_174, plain, (X1=nil|app(X1,X2)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169]), c_0_52])])).
% 0.06/0.70  cnf(c_0_175, plain, (segmentP(X1,X2)|~ssList(X1)|~ssList(X2)|app(app(X3,X2),X4)!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_170])).
% 0.06/0.70  cnf(c_0_176, plain, (frontsegP(cons(X1,X2),cons(X1,X2))|~ssList(X2)|~ssItem(X1)), inference(spm,[status(thm)],[c_0_171, c_0_172])).
% 0.06/0.70  cnf(c_0_177, negated_conjecture, (cons(esk55_0,esk44_1(esk54_0))=esk54_0|esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_173]), c_0_45])])).
% 0.06/0.70  fof(c_0_178, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~rearsegP(X4,X5)|rearsegP(app(X6,X4),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
% 0.06/0.70  fof(c_0_179, plain, (![X2]:(~ssList(X2)|rearsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])).
% 0.06/0.70  fof(c_0_180, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~frontsegP(X4,X5)|frontsegP(app(X4,X6),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])).
% 0.06/0.70  fof(c_0_181, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])).
% 0.06/0.70  cnf(c_0_182, plain, (cons(hd(X1),tl(app(X1,X2)))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_88]), c_0_56]), c_0_91])).
% 0.06/0.70  cnf(c_0_183, negated_conjecture, (esk53_0=nil|app(cons(esk52_0,nil),esk54_0)!=esk50_0|~ssList(app(cons(esk52_0,nil),esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_57]), c_0_46])])).
% 0.06/0.70  cnf(c_0_184, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(er,[status(thm)],[c_0_175])).
% 0.06/0.70  cnf(c_0_185, negated_conjecture, (esk54_0=nil|frontsegP(esk54_0,esk54_0)|~ssList(esk44_1(esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177]), c_0_158])])).
% 0.06/0.70  cnf(c_0_186, plain, (rearsegP(app(X1,X2),X3)|~rearsegP(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_178])).
% 0.06/0.70  cnf(c_0_187, plain, (rearsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_179])).
% 0.06/0.70  cnf(c_0_188, plain, (frontsegP(app(X1,X2),X3)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_180])).
% 0.06/0.70  cnf(c_0_189, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_181])).
% 0.06/0.70  fof(c_0_190, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])])])])).
% 0.06/0.70  cnf(c_0_191, negated_conjecture, (cons(hd(esk53_0),tl(esk50_0))=esk50_0|esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_89]), c_0_90]), c_0_46])])).
% 0.06/0.70  cnf(c_0_192, negated_conjecture, (esk53_0=nil|app(cons(esk52_0,nil),esk54_0)!=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_56]), c_0_45])])).
% 0.06/0.70  cnf(c_0_193, plain, (segmentP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_126]), c_0_52])]), c_0_56])).
% 0.06/0.70  fof(c_0_194, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])).
% 0.06/0.70  cnf(c_0_195, negated_conjecture, (esk54_0=nil|frontsegP(esk54_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_151]), c_0_45])])).
% 0.06/0.70  cnf(c_0_196, plain, (rearsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187]), c_0_52])])).
% 0.06/0.70  cnf(c_0_197, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_189]), c_0_52])])).
% 0.06/0.70  fof(c_0_198, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|((~rearsegP(X3,X4)|~rearsegP(X4,X3))|X3=X4)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_36])])])).
% 0.06/0.70  cnf(c_0_199, plain, (rearsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X3,X2)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_190])).
% 0.06/0.70  cnf(c_0_200, negated_conjecture, (cons(esk55_0,tl(esk50_0))=esk50_0|esk53_0=nil), inference(spm,[status(thm)],[c_0_191, c_0_132])).
% 0.06/0.70  cnf(c_0_201, negated_conjecture, (esk53_0=nil|app(cons(esk52_0,nil),esk54_0)!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_51]), c_0_52]), c_0_53])])).
% 0.06/0.70  fof(c_0_202, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|((~segmentP(X3,X4)|~segmentP(X4,X3))|X3=X4)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])).
% 0.06/0.70  cnf(c_0_203, negated_conjecture, (segmentP(esk50_0,cons(esk55_0,esk54_0))|~ssList(cons(esk55_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_157]), c_0_46])])).
% 0.06/0.70  cnf(c_0_204, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_194])).
% 0.06/0.70  cnf(c_0_205, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_169]), c_0_52])])).
% 0.06/0.70  cnf(c_0_206, negated_conjecture, (app(esk54_0,esk6_2(esk54_0,esk54_0))=esk54_0|esk54_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_195]), c_0_45])])).
% 0.06/0.70  cnf(c_0_207, negated_conjecture, (esk54_0=nil|ssList(esk6_2(esk54_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_195]), c_0_45])])).
% 0.06/0.70  cnf(c_0_208, negated_conjecture, (rearsegP(esk50_0,nil)|~ssList(cons(esk55_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_157]), c_0_46])])).
% 0.06/0.70  cnf(c_0_209, negated_conjecture, (frontsegP(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_89]), c_0_90]), c_0_46])])).
% 0.06/0.70  cnf(c_0_210, plain, (X1=X2|~rearsegP(X2,X1)|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_198])).
% 0.06/0.70  cnf(c_0_211, plain, (rearsegP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_199]), c_0_56])).
% 0.06/0.70  cnf(c_0_212, negated_conjecture, (app(esk53_0,app(cons(esk55_0,nil),esk54_0))=esk50_0), inference(rw,[status(thm)],[c_0_57, c_0_152])).
% 0.06/0.70  cnf(c_0_213, negated_conjecture, (cons(esk55_0,nil)=esk50_0|esk53_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_86]), c_0_110])).
% 0.06/0.70  cnf(c_0_214, negated_conjecture, (esk53_0=nil|app(cons(esk55_0,nil),esk54_0)!=esk50_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_143]), c_0_110])).
% 0.06/0.70  cnf(c_0_215, plain, (X1=X2|~segmentP(X2,X1)|~segmentP(X1,X2)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_202])).
% 0.06/0.70  cnf(c_0_216, negated_conjecture, (segmentP(esk50_0,cons(esk55_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_51]), c_0_45]), c_0_158])])).
% 0.06/0.70  cnf(c_0_217, negated_conjecture, (cons(esk55_0,app(nil,X1))=app(esk50_0,X1)|esk50_0=nil|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_76]), c_0_52])]), c_0_158])])).
% 0.06/0.70  cnf(c_0_218, negated_conjecture, (app(nil,esk54_0)=esk54_0|esk54_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206]), c_0_45])]), c_0_207])).
% 0.06/0.70  cnf(c_0_219, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_190])).
% 0.06/0.70  cnf(c_0_220, negated_conjecture, (rearsegP(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_51]), c_0_45]), c_0_158])])).
% 0.06/0.70  cnf(c_0_221, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_190])).
% 0.06/0.70  cnf(c_0_222, negated_conjecture, (app(nil,esk6_2(esk50_0,nil))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_209]), c_0_52]), c_0_79])])).
% 0.06/0.70  cnf(c_0_223, negated_conjecture, (ssList(esk6_2(esk50_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_209]), c_0_52]), c_0_79])])).
% 0.06/0.70  cnf(c_0_224, plain, (app(X1,X2)=X2|~rearsegP(X2,app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_56])).
% 0.06/0.70  cnf(c_0_225, negated_conjecture, (app(esk53_0,app(esk50_0,esk54_0))=esk50_0|esk53_0=nil), inference(spm,[status(thm)],[c_0_212, c_0_213])).
% 0.06/0.70  cnf(c_0_226, negated_conjecture, (esk53_0=nil|app(esk50_0,esk54_0)!=esk50_0), inference(spm,[status(thm)],[c_0_214, c_0_213])).
% 0.06/0.70  cnf(c_0_227, negated_conjecture, (cons(esk55_0,esk54_0)=esk50_0|~segmentP(cons(esk55_0,esk54_0),esk50_0)|~ssList(cons(esk55_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216]), c_0_79])])).
% 0.06/0.70  cnf(c_0_228, plain, (segmentP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_169]), c_0_52])]), c_0_56])).
% 0.06/0.70  cnf(c_0_229, negated_conjecture, (app(esk50_0,esk54_0)=cons(esk55_0,esk54_0)|esk54_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_218]), c_0_45])]), c_0_167])).
% 0.06/0.70  cnf(c_0_230, negated_conjecture, (app(esk7_2(esk50_0,nil),nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_220]), c_0_52]), c_0_79])])).
% 0.06/0.70  cnf(c_0_231, negated_conjecture, (ssList(esk7_2(esk50_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_220]), c_0_52]), c_0_79])])).
% 0.06/0.70  cnf(c_0_232, negated_conjecture, (esk6_2(esk50_0,nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_222]), c_0_223])])).
% 0.06/0.70  cnf(c_0_233, negated_conjecture, (esk53_0=nil|~rearsegP(app(esk50_0,esk54_0),esk50_0)|~ssList(app(esk50_0,esk54_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_225]), c_0_46])]), c_0_226])).
% 0.06/0.70  cnf(c_0_234, negated_conjecture, (cons(esk55_0,esk54_0)=esk50_0|~segmentP(cons(esk55_0,esk54_0),esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_51]), c_0_45]), c_0_158])])).
% 0.06/0.70  cnf(c_0_235, negated_conjecture, (esk54_0=nil|segmentP(cons(esk55_0,esk54_0),esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_229]), c_0_45]), c_0_79])])).
% 0.06/0.70  cnf(c_0_236, negated_conjecture, (esk54_0=nil|cons(esk55_0,esk54_0)!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_229]), c_0_79]), c_0_45])])).
% 0.06/0.70  cnf(c_0_237, negated_conjecture, (esk7_2(esk50_0,nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_230]), c_0_231])])).
% 0.06/0.70  cnf(c_0_238, negated_conjecture, (app(nil,esk50_0)=esk50_0), inference(rw,[status(thm)],[c_0_222, c_0_232])).
% 0.06/0.70  cnf(c_0_239, negated_conjecture, (memberP(esk54_0,esk52_0)|memberP(esk53_0,esk52_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
% 0.06/0.70  cnf(c_0_240, negated_conjecture, (esk53_0=nil|~rearsegP(app(esk50_0,esk54_0),esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_56]), c_0_45]), c_0_79])])).
% 0.06/0.70  cnf(c_0_241, negated_conjecture, (esk54_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_235]), c_0_236])).
% 0.06/0.70  cnf(c_0_242, negated_conjecture, (app(esk50_0,nil)=esk50_0), inference(rw,[status(thm)],[c_0_230, c_0_237])).
% 0.06/0.70  cnf(c_0_243, negated_conjecture, (rearsegP(esk50_0,esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_238]), c_0_52]), c_0_79])])).
% 0.06/0.70  cnf(c_0_244, negated_conjecture, (memberP(esk54_0,esk55_0)|memberP(esk53_0,esk55_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_239, c_0_152]), c_0_152])).
% 0.06/0.70  cnf(c_0_245, negated_conjecture, (esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_240, c_0_241]), c_0_242]), c_0_243])])).
% 0.06/0.70  cnf(c_0_246, negated_conjecture, (memberP(nil,esk55_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_244, c_0_241]), c_0_245])])).
% 0.06/0.70  cnf(c_0_247, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_246]), c_0_158])]), ['proof']).
% 0.06/0.70  # SZS output end CNFRefutation
% 0.06/0.70  # Training examples: 0 positive, 0 negative
% 0.06/0.70  
% 0.06/0.70  # -------------------------------------------------
% 0.06/0.70  # User time                : 0.469 s
% 0.06/0.70  # System time              : 0.008 s
% 0.06/0.70  # Total time               : 0.477 s
% 0.06/0.70  # Maximum resident set size: 1872 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
