<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC068-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC068-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC068-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC068-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC068-1.p">SWC068-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n142.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:16:31 EDT 2016

% Result     : Unsatisfiable 0.02s
% Output     : CNFRefutation 0.02s
% Verified   : 
% Statistics : Number of clauses        :   27 (  74 expanded)
%              Number of leaves         :   14 (  45 expanded)
%              Depth                    :    6
%              Number of atoms          :   62 ( 135 expanded)
%              Number of equality atoms :   21 (  45 expanded)
%              Maximal clause size      :    6 (   2 average)
%              Maximal term depth       :    3 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,axiom,
    ( segmentP(X4,X2)
    | app(app(X1,X2),X3) != X4
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X4) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause173)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( app(sk4,sk4) = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_2"></A>cnf(c_0_2,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_3"></A>cnf(c_0_3,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_4"></A>cnf(c_0_4,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( ~ ssList(X1)
    | ~ neq(X1,nil)
    | ~ segmentP(sk2,X1)
    | ~ segmentP(sk1,X1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_8)).

<A NAME="c_0_6"></A>cnf(c_0_6,axiom,
    ( neq(X2,X1)
    | X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause100)).

<A NAME="c_0_7"></A>cnf(c_0_7,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_8"></A>cnf(c_0_8,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_10"></A>cnf(c_0_10,negated_conjecture,
    ( nil != sk2
    | nil != sk1 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( segmentP(X1,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause57)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( app(X1,X2) = nil
    | nil != X1
    | nil != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause145)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( nil = X2
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause119)).

<A NAME="c_0_14"></A>cnf(c_0_14,plain,
    ( segmentP(app(app(X1,X2),X3),X2)
    | ~ ssList(app(app(X1,X2),X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>])).

<A NAME="c_0_15"></A>cnf(c_0_15,negated_conjecture,
    ( app(sk4,sk4) = sk1 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_1">c_0_1</A>,<A HREF="#c_0_2">c_0_2</A>])).

<A NAME="c_0_16"></A>cnf(c_0_16,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>,<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_17"></A>cnf(c_0_17,negated_conjecture,
    ( ~ neq(X1,nil)
    | ~ segmentP(sk1,X1)
    | ~ segmentP(sk4,X1)
    | ~ ssList(X1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>,<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_18"></A>cnf(c_0_18,negated_conjecture,
    ( segmentP(app(sk1,X1),sk4)
    | ~ ssList(app(sk1,X1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>,c_0_15]),c_0_16])])).

<A NAME="c_0_19"></A>cnf(c_0_19,negated_conjecture,
    ( nil = X1
    | ~ segmentP(sk1,X1)
    | ~ segmentP(sk4,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_6]),c_0_7])])).

<A NAME="c_0_20"></A>cnf(c_0_20,negated_conjecture,
    ( segmentP(sk1,sk4) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_8]),c_0_9]),c_0_7])])).

<A NAME="c_0_21"></A>cnf(c_0_21,negated_conjecture,
    ( sk4 = nil
    | ~ segmentP(sk4,sk4) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,c_0_20]),c_0_16])])).

<A NAME="c_0_22"></A>cnf(c_0_22,negated_conjecture,
    ( sk1 != nil
    | sk4 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,<A HREF="#c_0_4">c_0_4</A>])).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_11]),c_0_16])])).

<A NAME="c_0_24"></A>cnf(c_0_24,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_12])]),c_0_7])])).

<A NAME="c_0_25"></A>cnf(c_0_25,negated_conjecture,
    ( sk1 != nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_13">c_0_13</A>,c_0_15]),c_0_16])]),<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_26"></A>cnf(c_0_26,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_15">c_0_15</A>,c_0_23]),c_0_23]),c_0_24]),<A HREF="#c_0_25">c_0_25</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC068-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.23  % Computer   : n142.star.cs.uiowa.edu
% 0.02/0.23  % Model      : x86_64 x86_64
% 0.02/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.23  % Memory     : 32218.75MB
% 0.02/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.23  % CPULimit   : 300
% 0.02/0.23  % DateTime   : Tue May 10 20:34:54 CDT 2016
% 0.02/0.23  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 0.02/0.26  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S2k
% 0.02/0.26  # and selection function SelectCQArNTEqFirst.
% 0.02/0.26  #
% 0.02/0.26  # Preprocessing time       : 0.009 s
% 0.02/0.26  # Presaturation interreduction done
% 0.02/0.26  
% 0.02/0.26  # Proof found!
% 0.02/0.26  # SZS status Unsatisfiable
% 0.02/0.26  # SZS output start CNFRefutation
% 0.02/0.26  cnf(c_0_0, axiom, (segmentP(X4,X2)|app(app(X1,X2),X3)!=X4|~ssList(X3)|~ssList(X1)|~ssList(X2)|~ssList(X4)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause173)).
% 0.02/0.26  cnf(c_0_1, negated_conjecture, (app(sk4,sk4)=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 0.02/0.26  cnf(c_0_2, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 0.02/0.26  cnf(c_0_3, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 0.02/0.26  cnf(c_0_4, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 0.02/0.26  cnf(c_0_5, negated_conjecture, (~ssList(X1)|~neq(X1,nil)|~segmentP(sk2,X1)|~segmentP(sk1,X1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_8)).
% 0.02/0.26  cnf(c_0_6, axiom, (neq(X2,X1)|X2=X1|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause100)).
% 0.02/0.26  cnf(c_0_7, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 0.02/0.26  cnf(c_0_8, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause73)).
% 0.02/0.26  cnf(c_0_9, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 0.02/0.26  cnf(c_0_10, negated_conjecture, (nil!=sk2|nil!=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
% 0.02/0.26  cnf(c_0_11, axiom, (segmentP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause57)).
% 0.02/0.26  cnf(c_0_12, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause145)).
% 0.02/0.26  cnf(c_0_13, axiom, (nil=X2|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause119)).
% 0.02/0.26  cnf(c_0_14, plain, (segmentP(app(app(X1,X2),X3),X2)|~ssList(app(app(X1,X2),X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(er,[status(thm)],[c_0_0])).
% 0.02/0.26  cnf(c_0_15, negated_conjecture, (app(sk4,sk4)=sk1), inference(rw,[status(thm)],[c_0_1, c_0_2])).
% 0.02/0.26  cnf(c_0_16, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_3, c_0_4])).
% 0.02/0.26  cnf(c_0_17, negated_conjecture, (~neq(X1,nil)|~segmentP(sk1,X1)|~segmentP(sk4,X1)|~ssList(X1)), inference(rw,[status(thm)],[c_0_5, c_0_4])).
% 0.02/0.26  cnf(c_0_18, negated_conjecture, (segmentP(app(sk1,X1),sk4)|~ssList(app(sk1,X1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])])).
% 0.02/0.26  cnf(c_0_19, negated_conjecture, (nil=X1|~segmentP(sk1,X1)|~segmentP(sk4,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_6]), c_0_7])])).
% 0.02/0.26  cnf(c_0_20, negated_conjecture, (segmentP(sk1,sk4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_8]), c_0_9]), c_0_7])])).
% 0.02/0.26  cnf(c_0_21, negated_conjecture, (sk4=nil|~segmentP(sk4,sk4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_16])])).
% 0.02/0.26  cnf(c_0_22, negated_conjecture, (sk1!=nil|sk4!=nil), inference(rw,[status(thm)],[c_0_10, c_0_4])).
% 0.02/0.26  cnf(c_0_23, negated_conjecture, (sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_11]), c_0_16])])).
% 0.02/0.26  cnf(c_0_24, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_12])]), c_0_7])])).
% 0.02/0.26  cnf(c_0_25, negated_conjecture, (sk1!=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_15]), c_0_16])]), c_0_22])).
% 0.02/0.26  cnf(c_0_26, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_23]), c_0_23]), c_0_24]), c_0_25]), ['proof']).
% 0.02/0.26  # SZS output end CNFRefutation
% 0.02/0.26  # Training examples: 0 positive, 0 negative
% 0.02/0.26  
% 0.02/0.26  # -------------------------------------------------
% 0.02/0.26  # User time                : 0.029 s
% 0.02/0.26  # System time              : 0.001 s
% 0.02/0.26  # Total time               : 0.030 s
% 0.02/0.26  # Maximum resident set size: 1860 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
