<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC295+1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC295+1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC295+1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC295%2B1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC295+1.p">SWC295+1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n153.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:23 EDT 2016

% Result     : Theorem 0.07s
% Output     : CNFRefutation 0.07s
% Verified   : 
% Statistics : Number of formulae       :  258 (1236628 expanded)
%              Number of clauses        :  195 (452775 expanded)
%              Number of leaves         :   31 (313435 expanded)
%              Depth                    :   39
%              Number of atoms          :  856 (17944406 expanded)
%              Number of equality atoms :  321 (5747003 expanded)
%              Maximal formula depth    :   29 (   4 average)
%              Maximal clause size      :   41 (   2 average)
%              Maximal term depth       :    5 (   1 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>fof(c_0_0,conjecture,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ! [X4] :
                  ( ssList(X4)
                 => ( X2 != X4
                    | X1 != X3
                    | ! [X5] :
                        ( ssItem(X5)
                       => ! [X6] :
                            ( ssList(X6)
                           => ! [X7] :
                                ( ssList(X7)
                               => ( app(app(X6,cons(X5,nil)),X7) != X1
                                  | ! [X8] :
                                      ( ssItem(X8)
                                     => ( ( ~ memberP(X6,X8)
                                          | lt(X8,X5) )
                                        & ( ~ memberP(X7,X8)
                                          | lt(X5,X8) ) ) ) ) ) ) )
                    | ( ! [X9] :
                          ( ssItem(X9)
                         => ( cons(X9,nil) != X3
                            | ~ memberP(X4,X9)
                            | ? [X10] :
                                ( ssItem(X10)
                                & X9 != X10
                                & memberP(X4,X10)
                                & leq(X9,X10) ) ) )
                      & ( nil != X4
                        | nil != X3 ) ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/theBenchmark.p',co1)).

<A NAME="c_0_1"></A>fof(c_0_1,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => app(app(X1,X2),X3) = app(X1,app(X2,X3)) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax82)).

<A NAME="c_0_2"></A>fof(c_0_2,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ssList(cons(X2,X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax16)).

<A NAME="c_0_3"></A>fof(c_0_3,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( frontsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X2,X3) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax5)).

<A NAME="c_0_4"></A>fof(c_0_4,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ssList(app(X1,X2)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax26)).

<A NAME="c_0_5"></A>fof(c_0_5,axiom,(
    ssList(nil) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax17)).

<A NAME="c_0_6"></A>fof(c_0_6,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => cons(X2,X1) = app(cons(X2,nil),X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax81)).

<A NAME="c_0_7"></A>fof(c_0_7,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => cons(hd(X1),tl(X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax78)).

<A NAME="c_0_8"></A>fof(c_0_8,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => hd(cons(X2,X1)) = X2 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax23)).

<A NAME="c_0_9"></A>fof(c_0_9,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => nil != cons(X2,X1) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax21)).

<A NAME="c_0_10"></A>fof(c_0_10,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( singletonP(X1)
      &lt;=> ? [X2] :
            ( ssItem(X2)
            & cons(X2,nil) = X1 ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax4)).

<A NAME="c_0_11"></A>fof(c_0_11,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => tl(cons(X2,X1)) = X1 ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax25)).

<A NAME="c_0_12"></A>fof(c_0_12,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => hd(app(X1,X2)) = hd(X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax85)).

<A NAME="c_0_13"></A>fof(c_0_13,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil = app(X1,X2)
          &lt;=> ( nil = X2
              & nil = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax83)).

<A NAME="c_0_14"></A>fof(c_0_14,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax38)).

<A NAME="c_0_15"></A>fof(c_0_15,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ( memberP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & ? [X4] :
                    ( ssList(X4)
                    & app(X3,cons(X2,X4)) = X1 ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax3)).

<A NAME="c_0_16"></A>fof(c_0_16,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssItem(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(cons(X2,X3),X1)
              &lt;=> ( X1 = X2
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax37)).

<A NAME="c_0_17"></A>fof(c_0_17,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssList(tl(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax24)).

<A NAME="c_0_18"></A>fof(c_0_18,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( nil != X1
           => tl(app(X1,X2)) = app(tl(X1),X2) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax86)).

<A NAME="c_0_19"></A>fof(c_0_19,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil != X1
       => ssItem(hd(X1)) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax22)).

<A NAME="c_0_20"></A>fof(c_0_20,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X2,X3) = app(X2,X1)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax80)).

<A NAME="c_0_21"></A>fof(c_0_21,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(X1,nil) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax84)).

<A NAME="c_0_22"></A>fof(c_0_22,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ( rearsegP(X1,X2)
          &lt;=> ? [X3] :
                ( ssList(X3)
                & app(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax6)).

<A NAME="c_0_23"></A>fof(c_0_23,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( app(X3,X2) = app(X1,X2)
               => X3 = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax79)).

<A NAME="c_0_24"></A>fof(c_0_24,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => cons(X3,app(X2,X1)) = app(cons(X3,X2),X1) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax27)).

<A NAME="c_0_25"></A>fof(c_0_25,axiom,(
    ! [X1] :
      ( ssList(X1)
     => app(nil,X1) = X1 ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax28)).

<A NAME="c_0_26"></A>fof(c_0_26,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( frontsegP(X1,X2)
               => frontsegP(app(X1,X3),X2) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax43)).

<A NAME="c_0_27"></A>fof(c_0_27,axiom,(
    ! [X1] :
      ( ssList(X1)
     => frontsegP(X1,nil) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax45)).

<A NAME="c_0_28"></A>fof(c_0_28,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ( nil = X1
        | ? [X2] :
            ( ssList(X2)
            & ? [X3] :
                ( ssItem(X3)
                & cons(X3,X2) = X1 ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax20)).

<A NAME="c_0_29"></A>fof(c_0_29,axiom,(
    ! [X1] :
      ( ssItem(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssList(X3)
             => ( memberP(app(X2,X3),X1)
              &lt;=> ( memberP(X2,X1)
                  | memberP(X3,X1) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax36)).

<A NAME="c_0_30"></A>fof(c_0_30,axiom,(
    ! [X1] :
      ( ssList(X1)
     => ! [X2] :
          ( ssList(X2)
         => ! [X3] :
              ( ssItem(X3)
             => ! [X4] :
                  ( ssItem(X4)
                 => ( cons(X3,X1) = cons(X4,X2)
                   => ( X3 = X4
                      & X2 = X1 ) ) ) ) ) ) ),
    file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax',ax19)).

<A NAME="c_0_31"></A>fof(c_0_31,negated_conjecture,(
    ~ ! [X1] :
        ( ssList(X1)
       => ! [X2] :
            ( ssList(X2)
           => ! [X3] :
                ( ssList(X3)
               => ! [X4] :
                    ( ssList(X4)
                   => ( X2 != X4
                      | X1 != X3
                      | ! [X5] :
                          ( ssItem(X5)
                         => ! [X6] :
                              ( ssList(X6)
                             => ! [X7] :
                                  ( ssList(X7)
                                 => ( app(app(X6,cons(X5,nil)),X7) != X1
                                    | ! [X8] :
                                        ( ssItem(X8)
                                       => ( ( ~ memberP(X6,X8)
                                            | lt(X8,X5) )
                                          & ( ~ memberP(X7,X8)
                                            | lt(X5,X8) ) ) ) ) ) ) )
                      | ( ! [X9] :
                            ( ssItem(X9)
                           => ( cons(X9,nil) != X3
                              | ~ memberP(X4,X9)
                              | ? [X10] :
                                  ( ssItem(X10)
                                  & X9 != X10
                                  & memberP(X4,X10)
                                  & leq(X9,X10) ) ) )
                        & ( nil != X4
                          | nil != X3 ) ) ) ) ) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).

<A NAME="c_0_32"></A>fof(c_0_32,negated_conjecture,(
    ! [X20] :
      ( ssList(esk48_0)
      & ssList(esk49_0)
      & ssList(esk50_0)
      & ssList(esk51_0)
      & esk49_0 = esk51_0
      & esk48_0 = esk50_0
      & ssItem(esk52_0)
      & ssList(esk53_0)
      & ssList(esk54_0)
      & app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk48_0
      & ssItem(esk55_0)
      & ( memberP(esk54_0,esk55_0)
        | memberP(esk53_0,esk55_0) )
      & ( ~ lt(esk52_0,esk55_0)
        | memberP(esk53_0,esk55_0) )
      & ( memberP(esk54_0,esk55_0)
        | ~ lt(esk55_0,esk52_0) )
      & ( ~ lt(esk52_0,esk55_0)
        | ~ lt(esk55_0,esk52_0) )
      & ( nil = esk51_0
        | ssItem(esk56_0) )
      & ( nil = esk50_0
        | ssItem(esk56_0) )
      & ( nil = esk51_0
        | cons(esk56_0,nil) = esk50_0 )
      & ( nil = esk50_0
        | cons(esk56_0,nil) = esk50_0 )
      & ( nil = esk51_0
        | memberP(esk51_0,esk56_0) )
      & ( nil = esk50_0
        | memberP(esk51_0,esk56_0) )
      & ( nil = esk51_0
        | ~ ssItem(X20)
        | esk56_0 = X20
        | ~ memberP(esk51_0,X20)
        | ~ leq(esk56_0,X20) )
      & ( nil = esk50_0
        | ~ ssItem(X20)
        | esk56_0 = X20
        | ~ memberP(esk51_0,X20)
        | ~ leq(esk56_0,X20) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])])])).

<A NAME="c_0_33"></A>cnf(c_0_33,negated_conjecture,
    ( app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk48_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_34"></A>cnf(c_0_34,negated_conjecture,
    ( esk48_0 = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_35"></A>fof(c_0_35,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(app(X4,X5),X6) = app(X4,app(X5,X6)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).

<A NAME="c_0_36"></A>cnf(c_0_36,negated_conjecture,
    ( app(app(esk53_0,cons(esk52_0,nil)),esk54_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_33">c_0_33</A>,<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_37"></A>cnf(c_0_37,plain,
    ( app(app(X1,X2),X3) = app(X1,app(X2,X3))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>])).

<A NAME="c_0_38"></A>cnf(c_0_38,negated_conjecture,
    ( ssList(esk54_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_39"></A>cnf(c_0_39,negated_conjecture,
    ( ssList(esk53_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_40"></A>fof(c_0_40,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | ssList(cons(X4,X3)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).

<A NAME="c_0_41"></A>fof(c_0_41,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk6_2(X4,X5))
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(X5,esk6_2(X4,X5)) = X4
        | ~ frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X5,X7) != X4
        | frontsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).

<A NAME="c_0_42"></A>fof(c_0_42,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | ssList(app(X3,X4)) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( app(esk53_0,app(cons(esk52_0,nil),esk54_0)) = esk50_0
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,c_0_37]),c_0_38]),c_0_39])])).

<A NAME="c_0_44"></A>cnf(c_0_44,plain,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,plain,
    ( ssList(nil) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_5">c_0_5</A>])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( ssItem(esk52_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_47"></A>fof(c_0_47,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | cons(X4,X3) = app(cons(X4,nil),X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).

<A NAME="c_0_48"></A>fof(c_0_48,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | cons(hd(X2),tl(X2)) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).

<A NAME="c_0_49"></A>fof(c_0_49,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | hd(cons(X4,X3)) = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).

<A NAME="c_0_50"></A>fof(c_0_50,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | nil != cons(X4,X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).

<A NAME="c_0_51"></A>fof(c_0_51,plain,(
    ! [X3,X5] :
      ( ( ssItem(esk5_1(X3))
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( cons(esk5_1(X3),nil) = X3
        | ~ singletonP(X3)
        | ~ ssList(X3) )
      & ( ~ ssItem(X5)
        | cons(X5,nil) != X3
        | singletonP(X3)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).

<A NAME="c_0_52"></A>cnf(c_0_52,plain,
    ( frontsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X2,X3) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( ssList(app(X1,X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_42">c_0_42</A>])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( app(esk53_0,app(cons(esk52_0,nil),esk54_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_44]),c_0_45]),c_0_46])])).

<A NAME="c_0_55"></A>cnf(c_0_55,plain,
    ( cons(X1,X2) = app(cons(X1,nil),X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_47">c_0_47</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,plain,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_48">c_0_48</A>])).

<A NAME="c_0_57"></A>cnf(c_0_57,plain,
    ( hd(cons(X1,X2)) = X1
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>])).

<A NAME="c_0_58"></A>cnf(c_0_58,plain,
    ( nil != cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_50">c_0_50</A>])).

<A NAME="c_0_59"></A>fof(c_0_59,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssItem(X4)
      | tl(cons(X4,X3)) = X3 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).

<A NAME="c_0_60"></A>cnf(c_0_60,plain,
    ( singletonP(X1)
    | ~ ssList(X1)
    | cons(X2,nil) != X1
    | ~ ssItem(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( ssList(esk48_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_62"></A>cnf(c_0_62,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_52]),<A HREF="#c_0_53">c_0_53</A>])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( app(esk53_0,cons(esk52_0,esk54_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_55]),c_0_38]),c_0_46])])).

<A NAME="c_0_64"></A>cnf(c_0_64,plain,
    ( cons(X1,tl(cons(X1,X2))) = cons(X1,X2)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_57]),c_0_44]),<A HREF="#c_0_58">c_0_58</A>])).

<A NAME="c_0_65"></A>cnf(c_0_65,plain,
    ( cons(esk5_1(X1),nil) = X1
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>])).

<A NAME="c_0_66"></A>cnf(c_0_66,plain,
    ( ssItem(esk5_1(X1))
    | ~ ssList(X1)
    | ~ singletonP(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>])).

<A NAME="c_0_67"></A>cnf(c_0_67,plain,
    ( tl(cons(X1,X2)) = X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_59">c_0_59</A>])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( cons(esk56_0,nil) = esk50_0
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_69"></A>cnf(c_0_69,negated_conjecture,
    ( ssItem(esk56_0)
    | nil = esk50_0 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_70"></A>cnf(c_0_70,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( ssList(esk50_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_61">c_0_61</A>,<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_72"></A>fof(c_0_72,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | hd(app(X3,X4)) = hd(X3) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).

<A NAME="c_0_73"></A>fof(c_0_73,plain,(
    ! [X3,X4] :
      ( ( nil = X4
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil = X3
        | nil != app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) )
      & ( nil != X4
        | nil != X3
        | nil = app(X3,X4)
        | ~ ssList(X4)
        | ~ ssList(X3) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( frontsegP(esk50_0,esk53_0)
    | ~ ssList(cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_63]),c_0_39])])).

<A NAME="c_0_75"></A>fof(c_0_75,plain,(
    ! [X1] :
      ( ssItem(X1)
     => ~ memberP(nil,X1) ) ),
    inference(fof_simplification,[status(thm)],[<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_76"></A>fof(c_0_76,plain,(
    ! [X5,X6,X9,X10] :
      ( ( ssList(esk3_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ssList(esk4_2(X5,X6))
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6))) = X5
        | ~ memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) )
      & ( ~ ssList(X9)
        | ~ ssList(X10)
        | app(X9,cons(X6,X10)) != X5
        | memberP(X5,X6)
        | ~ ssItem(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])])).

<A NAME="c_0_77"></A>cnf(c_0_77,plain,
    ( cons(esk5_1(X1),tl(X1)) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_64">c_0_64</A>,c_0_65]),c_0_45])]),<A HREF="#c_0_66">c_0_66</A>])).

<A NAME="c_0_78"></A>cnf(c_0_78,negated_conjecture,
    ( tl(esk50_0) = nil
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_67">c_0_67</A>,c_0_68]),c_0_45])]),<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,negated_conjecture,
    ( esk50_0 = nil
    | singletonP(esk50_0) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_70">c_0_70</A>,c_0_68]),c_0_71])]),<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_80"></A>cnf(c_0_80,plain,
    ( hd(app(X1,X2)) = hd(X1)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,plain,
    ( app(X2,esk6_2(X1,X2)) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( frontsegP(esk50_0,esk53_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_74">c_0_74</A>,c_0_44]),c_0_38]),c_0_46])])).

<A NAME="c_0_84"></A>cnf(c_0_84,plain,
    ( ssList(esk6_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ frontsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_41">c_0_41</A>])).

<A NAME="c_0_85"></A>fof(c_0_85,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(cons(X5,X6),X4)
        | X4 = X5
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( X4 != X5
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(cons(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssItem(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])).

<A NAME="c_0_86"></A>fof(c_0_86,plain,(
    ! [X2] :
      ( ~ ssItem(X2)
      | ~ memberP(nil,X2) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_75])])).

<A NAME="c_0_87"></A>cnf(c_0_87,plain,
    ( memberP(X1,X2)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | app(X3,cons(X2,X4)) != X1
    | ~ ssList(X4)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>])).

<A NAME="c_0_88"></A>fof(c_0_88,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssList(tl(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).

<A NAME="c_0_89"></A>fof(c_0_89,plain,(
    ! [X3,X4] :
      ( ~ ssList(X3)
      | ~ ssList(X4)
      | nil = X3
      | tl(app(X3,X4)) = app(tl(X3),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).

<A NAME="c_0_90"></A>cnf(c_0_90,plain,
    ( cons(esk5_1(X1),X2) = app(X1,X2)
    | ~ singletonP(X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_65]),<A HREF="#c_0_66">c_0_66</A>])).

<A NAME="c_0_91"></A>cnf(c_0_91,negated_conjecture,
    ( cons(esk5_1(esk50_0),nil) = esk50_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_78]),c_0_71])]),<A HREF="#c_0_79">c_0_79</A>])).

<A NAME="c_0_92"></A>cnf(c_0_92,plain,
    ( cons(hd(X1),tl(app(X1,X2))) = app(X1,X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_80]),c_0_53]),<A HREF="#c_0_81">c_0_81</A>])).

<A NAME="c_0_93"></A>cnf(c_0_93,negated_conjecture,
    ( app(esk53_0,esk6_2(esk50_0,esk53_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,c_0_83]),c_0_39]),c_0_71])])).

<A NAME="c_0_94"></A>cnf(c_0_94,negated_conjecture,
    ( ssList(esk6_2(esk50_0,esk53_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,c_0_83]),c_0_39]),c_0_71])])).

<A NAME="c_0_95"></A>cnf(c_0_95,negated_conjecture,
    ( esk53_0 = nil
    | esk50_0 != nil
    | ~ ssList(cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_63]),c_0_39])])).

<A NAME="c_0_96"></A>cnf(c_0_96,plain,
    ( memberP(X3,X1)
    | X1 = X2
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | ~ memberP(cons(X2,X3),X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>])).

<A NAME="c_0_97"></A>cnf(c_0_97,plain,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_86">c_0_86</A>])).

<A NAME="c_0_98"></A>cnf(c_0_98,plain,
    ( memberP(app(X1,cons(X2,X3)),X2)
    | ~ ssList(app(X1,cons(X2,X3)))
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(er,[status(thm)],[<A HREF="#c_0_87">c_0_87</A>])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>])).

<A NAME="c_0_100"></A>cnf(c_0_100,plain,
    ( tl(app(X1,X2)) = app(tl(X1),X2)
    | nil = X1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>])).

<A NAME="c_0_101"></A>cnf(c_0_101,negated_conjecture,
    ( app(esk50_0,nil) = esk50_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_90">c_0_90</A>,c_0_91]),c_0_45]),c_0_71])]),<A HREF="#c_0_79">c_0_79</A>])).

<A NAME="c_0_102"></A>cnf(c_0_102,negated_conjecture,
    ( cons(hd(esk53_0),tl(esk50_0)) = esk50_0
    | esk53_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_93]),c_0_94]),c_0_39])])).

<A NAME="c_0_103"></A>cnf(c_0_103,negated_conjecture,
    ( esk53_0 = nil
    | esk50_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_95">c_0_95</A>,c_0_44]),c_0_38]),c_0_46])])).

<A NAME="c_0_104"></A>fof(c_0_104,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | nil = X2
      | ssItem(hd(X2)) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).

<A NAME="c_0_105"></A>cnf(c_0_105,plain,
    ( X1 = esk5_1(X2)
    | ~ singletonP(X2)
    | ~ memberP(X2,X1)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,c_0_65]),c_0_45])]),c_0_66]),<A HREF="#c_0_97">c_0_97</A>])).

<A NAME="c_0_106"></A>cnf(c_0_106,negated_conjecture,
    ( memberP(esk50_0,esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_98">c_0_98</A>,c_0_63]),c_0_71]),c_0_38]),c_0_39]),c_0_46])])).

<A NAME="c_0_107"></A>cnf(c_0_107,negated_conjecture,
    ( hd(esk50_0) = esk56_0
    | esk50_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,c_0_68]),c_0_45])]),<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_108"></A>cnf(c_0_108,plain,
    ( nil = X1
    | ssList(app(tl(X1),X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_99">c_0_99</A>,c_0_100]),c_0_53]),<A HREF="#c_0_81">c_0_81</A>])).

<A NAME="c_0_109"></A>cnf(c_0_109,negated_conjecture,
    ( app(tl(esk50_0),nil) = tl(esk50_0)
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>,c_0_101]),c_0_45]),c_0_71])])).

<A NAME="c_0_110"></A>fof(c_0_110,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X5,X6) != app(X5,X4)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).

<A NAME="c_0_111"></A>fof(c_0_111,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(X2,nil) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).

<A NAME="c_0_112"></A>cnf(c_0_112,negated_conjecture,
    ( cons(hd(esk53_0),nil) = esk50_0
    | esk53_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>,c_0_78]),<A HREF="#c_0_103">c_0_103</A>])).

<A NAME="c_0_113"></A>cnf(c_0_113,negated_conjecture,
    ( hd(esk53_0) = hd(esk50_0)
    | esk53_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_80">c_0_80</A>,c_0_93]),c_0_94]),c_0_39])])).

<A NAME="c_0_114"></A>cnf(c_0_114,plain,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_104">c_0_104</A>])).

<A NAME="c_0_115"></A>cnf(c_0_115,negated_conjecture,
    ( esk5_1(esk50_0) = esk52_0
    | ~ singletonP(esk50_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_105">c_0_105</A>,c_0_106]),c_0_71]),c_0_46])])).

<A NAME="c_0_116"></A>cnf(c_0_116,negated_conjecture,
    ( cons(esk56_0,tl(esk50_0)) = esk50_0
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_107]),c_0_71])])).

<A NAME="c_0_117"></A>cnf(c_0_117,negated_conjecture,
    ( esk50_0 = nil
    | ssList(tl(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_108">c_0_108</A>,c_0_109]),c_0_45]),c_0_71])])).

<A NAME="c_0_118"></A>cnf(c_0_118,plain,
    ( X1 = X2
    | app(X3,X1) != app(X3,X2)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_110">c_0_110</A>])).

<A NAME="c_0_119"></A>cnf(c_0_119,plain,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_111">c_0_111</A>])).

<A NAME="c_0_120"></A>cnf(c_0_120,negated_conjecture,
    ( cons(hd(esk50_0),nil) = esk50_0
    | esk53_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_112">c_0_112</A>,<A HREF="#c_0_113">c_0_113</A>])).

<A NAME="c_0_121"></A>cnf(c_0_121,negated_conjecture,
    ( esk53_0 = nil
    | ssItem(hd(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_114">c_0_114</A>,c_0_113]),c_0_39])])).

<A NAME="c_0_122"></A>cnf(c_0_122,negated_conjecture,
    ( esk5_1(esk50_0) = esk52_0
    | esk50_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_115">c_0_115</A>,<A HREF="#c_0_79">c_0_79</A>])).

<A NAME="c_0_123"></A>cnf(c_0_123,negated_conjecture,
    ( esk50_0 = nil
    | X1 = esk56_0
    | memberP(tl(esk50_0),X1)
    | ~ memberP(esk50_0,X1)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,c_0_116]),c_0_69]),<A HREF="#c_0_117">c_0_117</A>])).

<A NAME="c_0_124"></A>cnf(c_0_124,plain,
    ( X1 = nil
    | app(X2,X1) != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_118">c_0_118</A>,c_0_119]),c_0_45])])).

<A NAME="c_0_125"></A>cnf(c_0_125,negated_conjecture,
    ( cons(esk56_0,nil) = esk50_0
    | esk53_0 = nil ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_120">c_0_120</A>,c_0_107]),<A HREF="#c_0_103">c_0_103</A>])).

<A NAME="c_0_126"></A>cnf(c_0_126,negated_conjecture,
    ( esk53_0 = nil
    | ssItem(esk56_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_121">c_0_121</A>,c_0_107]),<A HREF="#c_0_103">c_0_103</A>])).

<A NAME="c_0_127"></A>cnf(c_0_127,negated_conjecture,
    ( cons(esk52_0,nil) = esk50_0
    | esk50_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,<A HREF="#c_0_122">c_0_122</A>])).

<A NAME="c_0_128"></A>cnf(c_0_128,negated_conjecture,
    ( esk52_0 = esk56_0
    | esk50_0 = nil
    | memberP(tl(esk50_0),esk52_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_123">c_0_123</A>,c_0_106]),c_0_46])])).

<A NAME="c_0_129"></A>fof(c_0_129,plain,(
    ! [X4,X5,X7] :
      ( ( ssList(esk7_2(X4,X5))
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( app(esk7_2(X4,X5),X5) = X4
        | ~ rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) )
      & ( ~ ssList(X7)
        | app(X7,X5) != X4
        | rearsegP(X4,X5)
        | ~ ssList(X5)
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])])).

<A NAME="c_0_130"></A>cnf(c_0_130,negated_conjecture,
    ( cons(esk52_0,esk54_0) = nil
    | esk53_0 != esk50_0
    | ~ ssList(cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_124">c_0_124</A>,c_0_63]),c_0_39])])).

<A NAME="c_0_131"></A>cnf(c_0_131,negated_conjecture,
    ( hd(esk50_0) = esk56_0
    | esk53_0 = nil ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,c_0_125]),c_0_45])]),<A HREF="#c_0_126">c_0_126</A>])).

<A NAME="c_0_132"></A>cnf(c_0_132,negated_conjecture,
    ( hd(esk50_0) = esk52_0
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,c_0_127]),c_0_45]),c_0_46])])).

<A NAME="c_0_133"></A>cnf(c_0_133,negated_conjecture,
    ( esk50_0 = nil
    | esk52_0 = esk56_0
    | memberP(nil,esk52_0) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_128">c_0_128</A>,<A HREF="#c_0_78">c_0_78</A>])).

<A NAME="c_0_134"></A>cnf(c_0_134,plain,
    ( rearsegP(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | app(X3,X2) != X1
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_129">c_0_129</A>])).

<A NAME="c_0_135"></A>cnf(c_0_135,negated_conjecture,
    ( cons(esk52_0,esk54_0) = nil
    | esk53_0 != esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_130">c_0_130</A>,c_0_44]),c_0_38]),c_0_46])])).

<A NAME="c_0_136"></A>cnf(c_0_136,negated_conjecture,
    ( esk53_0 = nil
    | esk52_0 = esk56_0 ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_131">c_0_131</A>,c_0_132]),<A HREF="#c_0_103">c_0_103</A>])).

<A NAME="c_0_137"></A>cnf(c_0_137,negated_conjecture,
    ( esk52_0 = esk56_0
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_133]),c_0_46])])).

<A NAME="c_0_138"></A>cnf(c_0_138,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_134]),<A HREF="#c_0_53">c_0_53</A>])).

<A NAME="c_0_139"></A>cnf(c_0_139,negated_conjecture,
    ( cons(esk52_0,esk54_0) = nil
    | esk52_0 = esk56_0 ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_135">c_0_135</A>,c_0_136]),<A HREF="#c_0_137">c_0_137</A>])).

<A NAME="c_0_140"></A>cnf(c_0_140,negated_conjecture,
    ( rearsegP(esk50_0,esk54_0)
    | ~ ssList(app(esk53_0,cons(esk52_0,nil))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_138">c_0_138</A>,c_0_36]),c_0_38])])).

<A NAME="c_0_141"></A>cnf(c_0_141,negated_conjecture,
    ( esk52_0 = esk56_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_139]),c_0_38]),c_0_46])])).

<A NAME="c_0_142"></A>cnf(c_0_142,negated_conjecture,
    ( rearsegP(esk50_0,esk54_0)
    | ~ ssList(cons(esk52_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_140">c_0_140</A>,c_0_53]),c_0_39])])).

<A NAME="c_0_143"></A>cnf(c_0_143,negated_conjecture,
    ( X1 = esk6_2(esk50_0,esk53_0)
    | app(esk53_0,X1) != esk50_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_118">c_0_118</A>,c_0_93]),c_0_39]),c_0_94])])).

<A NAME="c_0_144"></A>cnf(c_0_144,negated_conjecture,
    ( app(esk53_0,cons(esk56_0,esk54_0)) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_63">c_0_63</A>,<A HREF="#c_0_141">c_0_141</A>])).

<A NAME="c_0_145"></A>fof(c_0_145,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | app(X6,X5) != app(X4,X5)
      | X6 = X4 ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).

<A NAME="c_0_146"></A>cnf(c_0_146,plain,
    ( app(esk7_2(X1,X2),X2) = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_129">c_0_129</A>])).

<A NAME="c_0_147"></A>cnf(c_0_147,negated_conjecture,
    ( rearsegP(esk50_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_142">c_0_142</A>,c_0_44]),c_0_45]),c_0_46])])).

<A NAME="c_0_148"></A>cnf(c_0_148,plain,
    ( ssList(esk7_2(X1,X2))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ rearsegP(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_129">c_0_129</A>])).

<A NAME="c_0_149"></A>cnf(c_0_149,negated_conjecture,
    ( esk6_2(esk50_0,esk53_0) = cons(esk56_0,esk54_0)
    | ~ ssList(cons(esk56_0,esk54_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_143">c_0_143</A>,<A HREF="#c_0_144">c_0_144</A>])).

<A NAME="c_0_150"></A>cnf(c_0_150,negated_conjecture,
    ( ssItem(esk56_0) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,<A HREF="#c_0_141">c_0_141</A>])).

<A NAME="c_0_151"></A>cnf(c_0_151,plain,
    ( X1 = X2
    | app(X1,X3) != app(X2,X3)
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_145">c_0_145</A>])).

<A NAME="c_0_152"></A>cnf(c_0_152,negated_conjecture,
    ( app(esk7_2(esk50_0,esk54_0),esk54_0) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_146">c_0_146</A>,c_0_147]),c_0_38]),c_0_71])])).

<A NAME="c_0_153"></A>cnf(c_0_153,negated_conjecture,
    ( ssList(esk7_2(esk50_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_148">c_0_148</A>,c_0_147]),c_0_38]),c_0_71])])).

<A NAME="c_0_154"></A>fof(c_0_154,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssItem(X6)
      | cons(X6,app(X5,X4)) = app(cons(X6,X5),X4) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).

<A NAME="c_0_155"></A>cnf(c_0_155,negated_conjecture,
    ( esk6_2(esk50_0,esk53_0) = cons(esk56_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_149">c_0_149</A>,c_0_44]),c_0_38]),c_0_150])])).

<A NAME="c_0_156"></A>cnf(c_0_156,negated_conjecture,
    ( app(app(esk53_0,cons(esk56_0,nil)),esk54_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_36">c_0_36</A>,<A HREF="#c_0_141">c_0_141</A>])).

<A NAME="c_0_157"></A>cnf(c_0_157,negated_conjecture,
    ( X1 = esk7_2(esk50_0,esk54_0)
    | app(X1,esk54_0) != esk50_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_151">c_0_151</A>,c_0_152]),c_0_38]),c_0_153])])).

<A NAME="c_0_158"></A>cnf(c_0_158,plain,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_154">c_0_154</A>])).

<A NAME="c_0_159"></A>cnf(c_0_159,negated_conjecture,
    ( X1 = cons(esk56_0,esk54_0)
    | app(esk53_0,X1) != esk50_0
    | ~ ssList(X1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_143">c_0_143</A>,<A HREF="#c_0_155">c_0_155</A>])).

<A NAME="c_0_160"></A>cnf(c_0_160,negated_conjecture,
    ( app(esk53_0,app(cons(esk56_0,nil),esk54_0)) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,<A HREF="#c_0_141">c_0_141</A>])).

<A NAME="c_0_161"></A>cnf(c_0_161,negated_conjecture,
    ( app(app(esk53_0,esk50_0),esk54_0) = esk50_0
    | esk50_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_156">c_0_156</A>,<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_162"></A>cnf(c_0_162,negated_conjecture,
    ( esk7_2(esk50_0,esk54_0) = app(esk53_0,cons(esk56_0,nil))
    | ~ ssList(app(esk53_0,cons(esk56_0,nil))) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_157">c_0_157</A>,<A HREF="#c_0_156">c_0_156</A>])).

<A NAME="c_0_163"></A>cnf(c_0_163,plain,
    ( nil = X2
    | ~ ssList(X1)
    | ~ ssList(X2)
    | nil != app(X1,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>])).

<A NAME="c_0_164"></A>cnf(c_0_164,negated_conjecture,
    ( app(esk53_0,cons(esk52_0,app(nil,esk54_0))) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_158]),c_0_38]),c_0_45]),c_0_46])])).

<A NAME="c_0_165"></A>cnf(c_0_165,negated_conjecture,
    ( app(cons(esk56_0,nil),esk54_0) = cons(esk56_0,esk54_0)
    | ~ ssList(app(cons(esk56_0,nil),esk54_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_159">c_0_159</A>,<A HREF="#c_0_160">c_0_160</A>])).

<A NAME="c_0_166"></A>fof(c_0_166,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | app(nil,X2) = X2 ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])).

<A NAME="c_0_167"></A>cnf(c_0_167,negated_conjecture,
    ( esk7_2(esk50_0,esk54_0) = app(esk53_0,esk50_0)
    | esk50_0 = nil
    | ~ ssList(app(esk53_0,esk50_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_157">c_0_157</A>,<A HREF="#c_0_161">c_0_161</A>])).

<A NAME="c_0_168"></A>cnf(c_0_168,negated_conjecture,
    ( esk7_2(esk50_0,esk54_0) = app(esk53_0,cons(esk56_0,nil))
    | ~ ssList(cons(esk56_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_162">c_0_162</A>,c_0_53]),c_0_39])])).

<A NAME="c_0_169"></A>cnf(c_0_169,negated_conjecture,
    ( cons(esk52_0,app(nil,esk54_0)) = nil
    | esk50_0 != nil
    | ~ ssList(cons(esk52_0,app(nil,esk54_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_163">c_0_163</A>,c_0_164]),c_0_39])])).

<A NAME="c_0_170"></A>cnf(c_0_170,negated_conjecture,
    ( app(cons(esk56_0,nil),esk54_0) = cons(esk56_0,esk54_0)
    | ~ ssList(cons(esk56_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_165">c_0_165</A>,c_0_53]),c_0_38])])).

<A NAME="c_0_171"></A>cnf(c_0_171,plain,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_166">c_0_166</A>])).

<A NAME="c_0_172"></A>cnf(c_0_172,negated_conjecture,
    ( esk7_2(esk50_0,esk54_0) = app(esk53_0,esk50_0)
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_167">c_0_167</A>,c_0_53]),c_0_71]),c_0_39])])).

<A NAME="c_0_173"></A>cnf(c_0_173,negated_conjecture,
    ( esk7_2(esk50_0,esk54_0) = app(esk53_0,cons(esk56_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_168">c_0_168</A>,c_0_44]),c_0_45]),c_0_150])])).

<A NAME="c_0_174"></A>cnf(c_0_174,negated_conjecture,
    ( cons(esk52_0,app(nil,esk54_0)) = nil
    | esk50_0 != nil
    | ~ ssList(app(nil,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_169">c_0_169</A>,c_0_44]),c_0_46])])).

<A NAME="c_0_175"></A>cnf(c_0_175,negated_conjecture,
    ( app(cons(esk56_0,nil),esk54_0) = cons(esk56_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_170">c_0_170</A>,c_0_44]),c_0_45]),c_0_150])])).

<A NAME="c_0_176"></A>fof(c_0_176,plain,(
    ! [X4,X5,X6] :
      ( ~ ssList(X4)
      | ~ ssList(X5)
      | ~ ssList(X6)
      | ~ frontsegP(X4,X5)
      | frontsegP(app(X4,X6),X5) ) ),
    inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).

<A NAME="c_0_177"></A>fof(c_0_177,plain,(
    ! [X2] :
      ( ~ ssList(X2)
      | frontsegP(X2,nil) ) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])).

<A NAME="c_0_178"></A>cnf(c_0_178,plain,
    ( app(nil,app(X1,X2)) = app(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_171]),c_0_45])])).

<A NAME="c_0_179"></A>cnf(c_0_179,negated_conjecture,
    ( app(esk53_0,cons(esk56_0,nil)) = app(esk53_0,esk50_0)
    | esk50_0 = nil ),
    inference(spm,[status(thm)],[<A HREF="#c_0_172">c_0_172</A>,<A HREF="#c_0_173">c_0_173</A>])).

<A NAME="c_0_180"></A>cnf(c_0_180,negated_conjecture,
    ( cons(esk52_0,app(nil,esk54_0)) = nil
    | esk50_0 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_174">c_0_174</A>,c_0_53]),c_0_38]),c_0_45])])).

<A NAME="c_0_181"></A>cnf(c_0_181,plain,
    ( cons(X1,app(X2,X3)) != nil
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_158]),c_0_44]),<A HREF="#c_0_58">c_0_58</A>])).

<A NAME="c_0_182"></A>cnf(c_0_182,negated_conjecture,
    ( cons(esk56_0,app(nil,esk54_0)) = cons(esk56_0,esk54_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_158">c_0_158</A>,c_0_175]),c_0_38]),c_0_45]),c_0_150])])).

<A NAME="c_0_183"></A>cnf(c_0_183,plain,
    ( frontsegP(app(X1,X2),X3)
    | ~ frontsegP(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_176">c_0_176</A>])).

<A NAME="c_0_184"></A>cnf(c_0_184,plain,
    ( frontsegP(X1,nil)
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_177">c_0_177</A>])).

<A NAME="c_0_185"></A>cnf(c_0_185,negated_conjecture,
    ( app(nil,app(esk53_0,esk50_0)) = app(esk53_0,esk50_0)
    | esk50_0 = nil
    | ~ ssList(cons(esk56_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_178">c_0_178</A>,c_0_179]),c_0_39])])).

<A NAME="c_0_186"></A>cnf(c_0_186,negated_conjecture,
    ( cons(esk56_0,app(nil,esk54_0)) = nil
    | esk50_0 != nil ),
    inference(rw,[status(thm)],[<A HREF="#c_0_180">c_0_180</A>,<A HREF="#c_0_141">c_0_141</A>])).

<A NAME="c_0_187"></A>cnf(c_0_187,negated_conjecture,
    ( cons(esk56_0,esk54_0) != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_181">c_0_181</A>,c_0_182]),c_0_38]),c_0_45]),c_0_150])])).

<A NAME="c_0_188"></A>cnf(c_0_188,plain,
    ( frontsegP(app(X1,X2),nil)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_183">c_0_183</A>,c_0_184]),c_0_45])])).

<A NAME="c_0_189"></A>cnf(c_0_189,plain,
    ( memberP(cons(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssItem(X2)
    | ~ ssList(X3)
    | X1 != X2 ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>])).

<A NAME="c_0_190"></A>fof(c_0_190,plain,(
    ! [X4] :
      ( ( ssList(esk44_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( ssItem(esk45_1(X4))
        | nil = X4
        | ~ ssList(X4) )
      & ( cons(esk45_1(X4),esk44_1(X4)) = X4
        | nil = X4
        | ~ ssList(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])])).

<A NAME="c_0_191"></A>cnf(c_0_191,negated_conjecture,
    ( rearsegP(esk50_0,cons(esk52_0,esk54_0))
    | ~ ssList(cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_138">c_0_138</A>,c_0_63]),c_0_39])])).

<A NAME="c_0_192"></A>cnf(c_0_192,negated_conjecture,
    ( app(nil,app(esk53_0,esk50_0)) = app(esk53_0,esk50_0)
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_185">c_0_185</A>,c_0_68]),c_0_71])])).

<A NAME="c_0_193"></A>cnf(c_0_193,negated_conjecture,
    ( esk50_0 != nil ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_186">c_0_186</A>,c_0_182]),<A HREF="#c_0_187">c_0_187</A>])).

<A NAME="c_0_194"></A>cnf(c_0_194,negated_conjecture,
    ( frontsegP(esk50_0,nil) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_188">c_0_188</A>,c_0_93]),c_0_94]),c_0_39])])).

<A NAME="c_0_195"></A>cnf(c_0_195,plain,
    ( memberP(cons(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(er,[status(thm)],[<A HREF="#c_0_189">c_0_189</A>])).

<A NAME="c_0_196"></A>cnf(c_0_196,plain,
    ( nil = X1
    | cons(esk45_1(X1),esk44_1(X1)) = X1
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_190">c_0_190</A>])).

<A NAME="c_0_197"></A>cnf(c_0_197,plain,
    ( nil = X1
    | ssItem(esk45_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_190">c_0_190</A>])).

<A NAME="c_0_198"></A>cnf(c_0_198,plain,
    ( nil = X1
    | ssList(esk44_1(X1))
    | ~ ssList(X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_190">c_0_190</A>])).

<A NAME="c_0_199"></A>cnf(c_0_199,negated_conjecture,
    ( rearsegP(esk50_0,cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_191">c_0_191</A>,c_0_44]),c_0_38]),c_0_46])])).

<A NAME="c_0_200"></A>cnf(c_0_200,negated_conjecture,
    ( app(tl(esk7_2(esk50_0,esk54_0)),esk54_0) = tl(esk50_0)
    | esk7_2(esk50_0,esk54_0) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>,c_0_152]),c_0_38]),c_0_153])])).

<A NAME="c_0_201"></A>cnf(c_0_201,plain,
    ( nil = X1
    | app(X2,app(X3,X1)) != nil
    | ~ ssList(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_163">c_0_163</A>,c_0_37]),<A HREF="#c_0_53">c_0_53</A>])).

<A NAME="c_0_202"></A>cnf(c_0_202,negated_conjecture,
    ( app(nil,app(esk53_0,esk50_0)) = app(esk53_0,esk50_0) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_192">c_0_192</A>,<A HREF="#c_0_193">c_0_193</A>])).

<A NAME="c_0_203"></A>cnf(c_0_203,negated_conjecture,
    ( app(nil,esk6_2(esk50_0,nil)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,c_0_194]),c_0_45]),c_0_71])])).

<A NAME="c_0_204"></A>cnf(c_0_204,negated_conjecture,
    ( ssList(esk6_2(esk50_0,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,c_0_194]),c_0_45]),c_0_71])])).

<A NAME="c_0_205"></A>cnf(c_0_205,negated_conjecture,
    ( esk50_0 = nil
    | X1 = esk56_0
    | ~ memberP(esk50_0,X1)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,c_0_68]),c_0_45])]),c_0_69]),<A HREF="#c_0_97">c_0_97</A>])).

<A NAME="c_0_206"></A>cnf(c_0_206,plain,
    ( nil = X1
    | memberP(X1,esk45_1(X1))
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_195">c_0_195</A>,c_0_196]),c_0_197]),<A HREF="#c_0_198">c_0_198</A>])).

<A NAME="c_0_207"></A>fof(c_0_207,plain,(
    ! [X4,X5,X6] :
      ( ( ~ memberP(app(X5,X6),X4)
        | memberP(X5,X4)
        | memberP(X6,X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X5,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) )
      & ( ~ memberP(X6,X4)
        | memberP(app(X5,X6),X4)
        | ~ ssList(X6)
        | ~ ssList(X5)
        | ~ ssItem(X4) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])])).

<A NAME="c_0_208"></A>cnf(c_0_208,negated_conjecture,
    ( app(esk7_2(esk50_0,cons(esk52_0,esk54_0)),cons(esk52_0,esk54_0)) = esk50_0
    | ~ ssList(cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_146">c_0_146</A>,c_0_199]),c_0_71])])).

<A NAME="c_0_209"></A>cnf(c_0_209,negated_conjecture,
    ( ssList(esk7_2(esk50_0,cons(esk52_0,esk54_0)))
    | ~ ssList(cons(esk52_0,esk54_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_148">c_0_148</A>,c_0_199]),c_0_71])])).

<A NAME="c_0_210"></A>cnf(c_0_210,negated_conjecture,
    ( app(tl(app(esk53_0,cons(esk56_0,nil))),esk54_0) = tl(esk50_0)
    | app(esk53_0,cons(esk56_0,nil)) = nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_200">c_0_200</A>,c_0_173]),<A HREF="#c_0_173">c_0_173</A>])).

<A NAME="c_0_211"></A>cnf(c_0_211,negated_conjecture,
    ( cons(esk56_0,nil) = esk50_0 ),
    inference(sr,[status(thm)],[<A HREF="#c_0_68">c_0_68</A>,<A HREF="#c_0_193">c_0_193</A>])).

<A NAME="c_0_212"></A>cnf(c_0_212,negated_conjecture,
    ( tl(esk50_0) = nil ),
    inference(sr,[status(thm)],[<A HREF="#c_0_78">c_0_78</A>,<A HREF="#c_0_193">c_0_193</A>])).

<A NAME="c_0_213"></A>cnf(c_0_213,negated_conjecture,
    ( app(esk53_0,esk50_0) != nil ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_201">c_0_201</A>,c_0_202]),c_0_71]),c_0_39]),c_0_45])]),<A HREF="#c_0_193">c_0_193</A>])).

<A NAME="c_0_214"></A>cnf(c_0_214,negated_conjecture,
    ( esk6_2(esk50_0,nil) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_171">c_0_171</A>,c_0_203]),c_0_204])])).

<A NAME="c_0_215"></A>fof(c_0_215,plain,(
    ! [X5,X6,X7,X8] :
      ( ( X7 = X8
        | cons(X7,X5) != cons(X8,X6)
        | ~ ssItem(X8)
        | ~ ssItem(X7)
        | ~ ssList(X6)
        | ~ ssList(X5) )
      & ( X6 = X5
        | cons(X7,X5) != cons(X8,X6)
        | ~ ssItem(X8)
        | ~ ssItem(X7)
        | ~ ssList(X6)
        | ~ ssList(X5) ) ) ),
    inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])).

<A NAME="c_0_216"></A>cnf(c_0_216,negated_conjecture,
    ( esk45_1(esk50_0) = esk56_0
    | esk50_0 = nil
    | ~ ssItem(esk45_1(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_205">c_0_205</A>,c_0_206]),c_0_71])])).

<A NAME="c_0_217"></A>cnf(c_0_217,plain,
    ( memberP(app(X2,X3),X1)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | ~ memberP(X3,X1) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_207">c_0_207</A>])).

<A NAME="c_0_218"></A>cnf(c_0_218,negated_conjecture,
    ( memberP(esk53_0,esk55_0)
    | memberP(esk54_0,esk55_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_219"></A>cnf(c_0_219,negated_conjecture,
    ( ssItem(esk55_0) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_220"></A>cnf(c_0_220,negated_conjecture,
    ( X1 = esk53_0
    | app(X1,esk6_2(esk50_0,esk53_0)) != esk50_0
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_151">c_0_151</A>,c_0_93]),c_0_94]),c_0_39])])).

<A NAME="c_0_221"></A>cnf(c_0_221,negated_conjecture,
    ( app(esk7_2(esk50_0,cons(esk52_0,esk54_0)),cons(esk52_0,esk54_0)) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_208">c_0_208</A>,c_0_44]),c_0_38]),c_0_46])])).

<A NAME="c_0_222"></A>cnf(c_0_222,negated_conjecture,
    ( ssList(esk7_2(esk50_0,cons(esk52_0,esk54_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_209">c_0_209</A>,c_0_44]),c_0_38]),c_0_46])])).

<A NAME="c_0_223"></A>cnf(c_0_223,negated_conjecture,
    ( app(tl(app(esk53_0,esk50_0)),esk54_0) = nil ),
    inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_210">c_0_210</A>,c_0_211]),c_0_211]),c_0_212]),<A HREF="#c_0_213">c_0_213</A>])).

<A NAME="c_0_224"></A>cnf(c_0_224,negated_conjecture,
    ( esk50_0 = nil
    | ssList(app(esk53_0,esk50_0)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_153">c_0_153</A>,<A HREF="#c_0_172">c_0_172</A>])).

<A NAME="c_0_225"></A>cnf(c_0_225,plain,
    ( app(esk7_2(app(X1,X2),X2),X2) = app(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_146">c_0_146</A>,c_0_138]),<A HREF="#c_0_53">c_0_53</A>])).

<A NAME="c_0_226"></A>cnf(c_0_226,negated_conjecture,
    ( app(nil,esk50_0) = esk50_0 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_203">c_0_203</A>,<A HREF="#c_0_214">c_0_214</A>])).

<A NAME="c_0_227"></A>cnf(c_0_227,plain,
    ( ssList(esk7_2(app(X1,X2),X2))
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_148">c_0_148</A>,c_0_138]),<A HREF="#c_0_53">c_0_53</A>])).

<A NAME="c_0_228"></A>cnf(c_0_228,plain,
    ( X2 = X1
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssItem(X3)
    | ~ ssItem(X4)
    | cons(X3,X1) != cons(X4,X2) ),
    inference(split_conjunct,[status(thm)],[<A HREF="#c_0_215">c_0_215</A>])).

<A NAME="c_0_229"></A>cnf(c_0_229,negated_conjecture,
    ( esk45_1(esk50_0) = esk56_0
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_216">c_0_216</A>,c_0_197]),c_0_71])])).

<A NAME="c_0_230"></A>cnf(c_0_230,negated_conjecture,
    ( memberP(app(X1,esk54_0),esk55_0)
    | memberP(esk53_0,esk55_0)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_217">c_0_217</A>,c_0_218]),c_0_38]),c_0_219])])).

<A NAME="c_0_231"></A>cnf(c_0_231,negated_conjecture,
    ( app(esk50_0,X1) = cons(esk56_0,X1)
    | esk50_0 = nil
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_55">c_0_55</A>,c_0_68]),<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_232"></A>cnf(c_0_232,negated_conjecture,
    ( X1 = esk53_0
    | app(X1,cons(esk56_0,esk54_0)) != esk50_0
    | ~ ssList(X1) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_220">c_0_220</A>,<A HREF="#c_0_155">c_0_155</A>])).

<A NAME="c_0_233"></A>cnf(c_0_233,negated_conjecture,
    ( app(esk7_2(esk50_0,cons(esk56_0,esk54_0)),cons(esk56_0,esk54_0)) = esk50_0 ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_221">c_0_221</A>,c_0_141]),<A HREF="#c_0_141">c_0_141</A>])).

<A NAME="c_0_234"></A>cnf(c_0_234,negated_conjecture,
    ( ssList(esk7_2(esk50_0,cons(esk56_0,esk54_0))) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_222">c_0_222</A>,<A HREF="#c_0_141">c_0_141</A>])).

<A NAME="c_0_235"></A>cnf(c_0_235,negated_conjecture,
    ( esk54_0 = nil
    | ~ ssList(tl(app(esk53_0,esk50_0))) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_163">c_0_163</A>,c_0_223]),c_0_38])])).

<A NAME="c_0_236"></A>cnf(c_0_236,negated_conjecture,
    ( ssList(app(esk53_0,esk50_0)) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_224">c_0_224</A>,<A HREF="#c_0_193">c_0_193</A>])).

<A NAME="c_0_237"></A>cnf(c_0_237,plain,
    ( X1 = nil
    | app(X1,X2) != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_151">c_0_151</A>,c_0_171]),c_0_45])])).

<A NAME="c_0_238"></A>cnf(c_0_238,negated_conjecture,
    ( app(esk7_2(esk50_0,esk50_0),esk50_0) = esk50_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_225">c_0_225</A>,c_0_226]),c_0_71]),c_0_45])])).

<A NAME="c_0_239"></A>cnf(c_0_239,negated_conjecture,
    ( ssList(esk7_2(esk50_0,esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_227">c_0_227</A>,c_0_226]),c_0_71]),c_0_45])])).

<A NAME="c_0_240"></A>cnf(c_0_240,negated_conjecture,
    ( esk50_0 = nil
    | X1 = nil
    | cons(X2,X1) != esk50_0
    | ~ ssList(X1)
    | ~ ssItem(X2) ),
    inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_228">c_0_228</A>,c_0_68]),c_0_45])]),<A HREF="#c_0_69">c_0_69</A>])).

<A NAME="c_0_241"></A>cnf(c_0_241,negated_conjecture,
    ( cons(esk56_0,esk44_1(esk50_0)) = esk50_0
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_196">c_0_196</A>,c_0_229]),c_0_71])])).

<A NAME="c_0_242"></A>cnf(c_0_242,negated_conjecture,
    ( esk50_0 = nil
    | memberP(cons(esk56_0,esk54_0),esk55_0)
    | memberP(esk53_0,esk55_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_230">c_0_230</A>,c_0_231]),c_0_71]),c_0_38])])).

<A NAME="c_0_243"></A>cnf(c_0_243,negated_conjecture,
    ( esk7_2(esk50_0,cons(esk56_0,esk54_0)) = esk53_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_232">c_0_232</A>,c_0_233]),c_0_234])])).

<A NAME="c_0_244"></A>cnf(c_0_244,negated_conjecture,
    ( esk54_0 = nil ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_235">c_0_235</A>,c_0_99]),c_0_236])]),<A HREF="#c_0_213">c_0_213</A>])).

<A NAME="c_0_245"></A>cnf(c_0_245,negated_conjecture,
    ( esk7_2(esk50_0,esk50_0) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_237">c_0_237</A>,c_0_238]),c_0_71]),c_0_239])])).

<A NAME="c_0_246"></A>cnf(c_0_246,negated_conjecture,
    ( esk44_1(esk50_0) = nil
    | esk50_0 = nil
    | ~ ssList(esk44_1(esk50_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_240">c_0_240</A>,c_0_241]),c_0_150])])).

<A NAME="c_0_247"></A>cnf(c_0_247,negated_conjecture,
    ( memberP(cons(esk56_0,esk54_0),esk55_0)
    | memberP(esk53_0,esk55_0) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_242">c_0_242</A>,<A HREF="#c_0_193">c_0_193</A>])).

<A NAME="c_0_248"></A>cnf(c_0_248,negated_conjecture,
    ( esk53_0 = nil ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_243">c_0_243</A>,c_0_244]),c_0_211]),<A HREF="#c_0_245">c_0_245</A>])).

<A NAME="c_0_249"></A>cnf(c_0_249,negated_conjecture,
    ( esk44_1(esk50_0) = nil
    | esk50_0 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_246">c_0_246</A>,c_0_198]),c_0_71])])).

<A NAME="c_0_250"></A>cnf(c_0_250,plain,
    ( X1 = esk45_1(X2)
    | nil = X2
    | memberP(esk44_1(X2),X1)
    | ~ memberP(X2,X1)
    | ~ ssList(X2)
    | ~ ssItem(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,c_0_196]),c_0_197]),<A HREF="#c_0_198">c_0_198</A>])).

<A NAME="c_0_251"></A>cnf(c_0_251,negated_conjecture,
    ( memberP(nil,esk55_0)
    | memberP(esk50_0,esk55_0) ),
    inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_247">c_0_247</A>,c_0_244]),c_0_211]),<A HREF="#c_0_248">c_0_248</A>])).

<A NAME="c_0_252"></A>cnf(c_0_252,negated_conjecture,
    ( esk45_1(esk50_0) = esk56_0 ),
    inference(sr,[status(thm)],[<A HREF="#c_0_229">c_0_229</A>,<A HREF="#c_0_193">c_0_193</A>])).

<A NAME="c_0_253"></A>cnf(c_0_253,negated_conjecture,
    ( esk44_1(esk50_0) = nil ),
    inference(sr,[status(thm)],[<A HREF="#c_0_249">c_0_249</A>,<A HREF="#c_0_193">c_0_193</A>])).

<A NAME="c_0_254"></A>cnf(c_0_254,negated_conjecture,
    ( esk55_0 = esk56_0
    | memberP(nil,esk55_0) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_250">c_0_250</A>,c_0_251]),c_0_252]),c_0_253]),c_0_71]),c_0_219])]),<A HREF="#c_0_193">c_0_193</A>])).

<A NAME="c_0_255"></A>cnf(c_0_255,negated_conjecture,
    ( esk55_0 = esk56_0 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_254]),c_0_219])])).

<A NAME="c_0_256"></A>cnf(c_0_256,negated_conjecture,
    ( memberP(nil,esk56_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_218">c_0_218</A>,c_0_244]),c_0_248]),c_0_255]),c_0_255])])).

<A NAME="c_0_257"></A>cnf(c_0_257,negated_conjecture,
    ( $false ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_97">c_0_97</A>,c_0_256]),c_0_150])]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC295+1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.03/0.23  % Computer   : n153.star.cs.uiowa.edu
% 0.03/0.23  % Model      : x86_64 x86_64
% 0.03/0.23  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.03/0.23  % Memory     : 32218.75MB
% 0.03/0.23  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.03/0.23  % CPULimit   : 300
% 0.03/0.23  % DateTime   : Tue May 10 20:40:07 CDT 2016
% 0.03/0.23  % CPUTime    : 
% 0.03/0.23  # Version: 1.9.1pre011
% 0.03/0.23  # No SInE strategy applied
% 0.03/0.23  # Trying AutoSched0 for 151 seconds
% 0.07/0.80  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4c
% 0.07/0.80  # and selection function SelectCQPrecWNTNp.
% 0.07/0.80  #
% 0.07/0.80  # Preprocessing time       : 0.014 s
% 0.07/0.80  # Presaturation interreduction done
% 0.07/0.80  
% 0.07/0.80  # Proof found!
% 0.07/0.80  # SZS status Theorem
% 0.07/0.80  # SZS output start CNFRefutation
% 0.07/0.80  fof(c_0_0, conjecture, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>(app(app(X6,cons(X5,nil)),X7)!=X1|![X8]:(ssItem(X8)=>((~(memberP(X6,X8))|lt(X8,X5))&(~(memberP(X7,X8))|lt(X5,X8)))))))))|(![X9]:(ssItem(X9)=>((cons(X9,nil)!=X3|~(memberP(X4,X9)))|?[X10]:(((ssItem(X10)&X9!=X10)&memberP(X4,X10))&leq(X9,X10))))&(nil!=X4|nil!=X3)))))))), file('/export/starexec/sandbox/benchmark/theBenchmark.p', co1)).
% 0.07/0.80  fof(c_0_1, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>app(app(X1,X2),X3)=app(X1,app(X2,X3)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax82)).
% 0.07/0.80  fof(c_0_2, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>ssList(cons(X2,X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax16)).
% 0.07/0.80  fof(c_0_3, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(frontsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X2,X3)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax5)).
% 0.07/0.80  fof(c_0_4, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>ssList(app(X1,X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax26)).
% 0.07/0.80  fof(c_0_5, axiom, (ssList(nil)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax17)).
% 0.07/0.80  fof(c_0_6, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>cons(X2,X1)=app(cons(X2,nil),X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax81)).
% 0.07/0.80  fof(c_0_7, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>cons(hd(X1),tl(X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax78)).
% 0.07/0.80  fof(c_0_8, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>hd(cons(X2,X1))=X2))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax23)).
% 0.07/0.80  fof(c_0_9, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>nil!=cons(X2,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax21)).
% 0.07/0.80  fof(c_0_10, axiom, (![X1]:(ssList(X1)=>(singletonP(X1)&lt;=>?[X2]:(ssItem(X2)&cons(X2,nil)=X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax4)).
% 0.07/0.80  fof(c_0_11, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>tl(cons(X2,X1))=X1))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax25)).
% 0.07/0.80  fof(c_0_12, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>hd(app(X1,X2))=hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax85)).
% 0.07/0.80  fof(c_0_13, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil=app(X1,X2)&lt;=>(nil=X2&nil=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax83)).
% 0.07/0.80  fof(c_0_14, axiom, (![X1]:(ssItem(X1)=>~(memberP(nil,X1)))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax38)).
% 0.07/0.80  fof(c_0_15, axiom, (![X1]:(ssList(X1)=>![X2]:(ssItem(X2)=>(memberP(X1,X2)&lt;=>?[X3]:(ssList(X3)&?[X4]:(ssList(X4)&app(X3,cons(X2,X4))=X1)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax3)).
% 0.07/0.80  fof(c_0_16, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssItem(X2)=>![X3]:(ssList(X3)=>(memberP(cons(X2,X3),X1)&lt;=>(X1=X2|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax37)).
% 0.07/0.80  fof(c_0_17, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssList(tl(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax24)).
% 0.07/0.80  fof(c_0_18, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(nil!=X1=>tl(app(X1,X2))=app(tl(X1),X2))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax86)).
% 0.07/0.80  fof(c_0_19, axiom, (![X1]:(ssList(X1)=>(nil!=X1=>ssItem(hd(X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax22)).
% 0.07/0.80  fof(c_0_20, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X2,X3)=app(X2,X1)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax80)).
% 0.07/0.80  fof(c_0_21, axiom, (![X1]:(ssList(X1)=>app(X1,nil)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax84)).
% 0.07/0.80  fof(c_0_22, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(rearsegP(X1,X2)&lt;=>?[X3]:(ssList(X3)&app(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax6)).
% 0.07/0.80  fof(c_0_23, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(app(X3,X2)=app(X1,X2)=>X3=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax79)).
% 0.07/0.80  fof(c_0_24, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>cons(X3,app(X2,X1))=app(cons(X3,X2),X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax27)).
% 0.07/0.80  fof(c_0_25, axiom, (![X1]:(ssList(X1)=>app(nil,X1)=X1)), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax28)).
% 0.07/0.80  fof(c_0_26, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(frontsegP(X1,X2)=>frontsegP(app(X1,X3),X2)))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax43)).
% 0.07/0.80  fof(c_0_27, axiom, (![X1]:(ssList(X1)=>frontsegP(X1,nil))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax45)).
% 0.07/0.80  fof(c_0_28, axiom, (![X1]:(ssList(X1)=>(nil=X1|?[X2]:(ssList(X2)&?[X3]:(ssItem(X3)&cons(X3,X2)=X1))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax20)).
% 0.07/0.80  fof(c_0_29, axiom, (![X1]:(ssItem(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>(memberP(app(X2,X3),X1)&lt;=>(memberP(X2,X1)|memberP(X3,X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax36)).
% 0.07/0.80  fof(c_0_30, axiom, (![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssItem(X3)=>![X4]:(ssItem(X4)=>(cons(X3,X1)=cons(X4,X2)=>(X3=X4&X2=X1))))))), file('/export/starexec/sandbox/benchmark/Axioms/SWC001+0.ax', ax19)).
% 0.07/0.80  fof(c_0_31, negated_conjecture, (~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(((X2!=X4|X1!=X3)|![X5]:(ssItem(X5)=>![X6]:(ssList(X6)=>![X7]:(ssList(X7)=>(app(app(X6,cons(X5,nil)),X7)!=X1|![X8]:(ssItem(X8)=>((~memberP(X6,X8)|lt(X8,X5))&(~memberP(X7,X8)|lt(X5,X8)))))))))|(![X9]:(ssItem(X9)=>((cons(X9,nil)!=X3|~memberP(X4,X9))|?[X10]:(((ssItem(X10)&X9!=X10)&memberP(X4,X10))&leq(X9,X10))))&(nil!=X4|nil!=X3))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[c_0_0])])).
% 0.07/0.80  fof(c_0_32, negated_conjecture, (![X20]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((esk49_0=esk51_0&esk48_0=esk50_0)&(ssItem(esk52_0)&(ssList(esk53_0)&(ssList(esk54_0)&(app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk48_0&(ssItem(esk55_0)&(((memberP(esk54_0,esk55_0)|memberP(esk53_0,esk55_0))&(~lt(esk52_0,esk55_0)|memberP(esk53_0,esk55_0)))&((memberP(esk54_0,esk55_0)|~lt(esk55_0,esk52_0))&(~lt(esk52_0,esk55_0)|~lt(esk55_0,esk52_0))))))))))&(((nil=esk51_0|ssItem(esk56_0))&(nil=esk50_0|ssItem(esk56_0)))&((((nil=esk51_0|cons(esk56_0,nil)=esk50_0)&(nil=esk50_0|cons(esk56_0,nil)=esk50_0))&((nil=esk51_0|memberP(esk51_0,esk56_0))&(nil=esk50_0|memberP(esk51_0,esk56_0))))&((nil=esk51_0|(((~ssItem(X20)|esk56_0=X20)|~memberP(esk51_0,X20))|~leq(esk56_0,X20)))&(nil=esk50_0|(((~ssItem(X20)|esk56_0=X20)|~memberP(esk51_0,X20))|~leq(esk56_0,X20)))))))))))), 
% inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])])])).
% 0.07/0.80  cnf(c_0_33, negated_conjecture, (app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk48_0), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.07/0.80  cnf(c_0_34, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.07/0.80  fof(c_0_35, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|app(app(X4,X5),X6)=app(X4,app(X5,X6)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])).
% 0.07/0.80  cnf(c_0_36, negated_conjecture, (app(app(esk53_0,cons(esk52_0,nil)),esk54_0)=esk50_0), inference(rw,[status(thm)],[c_0_33, c_0_34])).
% 0.07/0.80  cnf(c_0_37, plain, (app(app(X1,X2),X3)=app(X1,app(X2,X3))|~ssList(X3)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
% 0.07/0.80  cnf(c_0_38, negated_conjecture, (ssList(esk54_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.07/0.80  cnf(c_0_39, negated_conjecture, (ssList(esk53_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.07/0.80  fof(c_0_40, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|ssList(cons(X4,X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
% 0.07/0.80  fof(c_0_41, plain, (![X4]:![X5]:![X7]:(((((ssList(esk6_2(X4,X5))|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(X5,esk6_2(X4,X5))=X4|~frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X5,X7)!=X4)|frontsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
% 0.07/0.80  fof(c_0_42, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|ssList(app(X3,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
% 0.07/0.80  cnf(c_0_43, negated_conjecture, (app(esk53_0,app(cons(esk52_0,nil),esk54_0))=esk50_0|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39])])).
% 0.07/0.80  cnf(c_0_44, plain, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
% 0.07/0.80  cnf(c_0_45, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[c_0_5])).
% 0.07/0.80  cnf(c_0_46, negated_conjecture, (ssItem(esk52_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.07/0.80  fof(c_0_47, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|cons(X4,X3)=app(cons(X4,nil),X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
% 0.07/0.80  fof(c_0_48, plain, (![X2]:(~ssList(X2)|(nil=X2|cons(hd(X2),tl(X2))=X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])).
% 0.07/0.80  fof(c_0_49, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|hd(cons(X4,X3))=X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
% 0.07/0.80  fof(c_0_50, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|nil!=cons(X4,X3)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
% 0.07/0.80  fof(c_0_51, plain, (![X3]:![X5]:((((ssItem(esk5_1(X3))|~singletonP(X3))|~ssList(X3))&((cons(esk5_1(X3),nil)=X3|~singletonP(X3))|~ssList(X3)))&(((~ssItem(X5)|cons(X5,nil)!=X3)|singletonP(X3))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])])).
% 0.07/0.80  cnf(c_0_52, plain, (frontsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X2,X3)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 0.07/0.80  cnf(c_0_53, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
% 0.07/0.80  cnf(c_0_54, negated_conjecture, (app(esk53_0,app(cons(esk52_0,nil),esk54_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46])])).
% 0.07/0.80  cnf(c_0_55, plain, (cons(X1,X2)=app(cons(X1,nil),X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
% 0.07/0.80  cnf(c_0_56, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
% 0.07/0.80  cnf(c_0_57, plain, (hd(cons(X1,X2))=X1|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
% 0.07/0.80  cnf(c_0_58, plain, (nil!=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
% 0.07/0.80  fof(c_0_59, plain, (![X3]:![X4]:(~ssList(X3)|(~ssItem(X4)|tl(cons(X4,X3))=X3))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
% 0.07/0.80  cnf(c_0_60, plain, (singletonP(X1)|~ssList(X1)|cons(X2,nil)!=X1|~ssItem(X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
% 0.07/0.80  cnf(c_0_61, negated_conjecture, (ssList(esk48_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.07/0.80  cnf(c_0_62, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_52]), c_0_53])).
% 0.07/0.80  cnf(c_0_63, negated_conjecture, (app(esk53_0,cons(esk52_0,esk54_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_38]), c_0_46])])).
% 0.07/0.80  cnf(c_0_64, plain, (cons(X1,tl(cons(X1,X2)))=cons(X1,X2)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_44]), c_0_58])).
% 0.07/0.80  cnf(c_0_65, plain, (cons(esk5_1(X1),nil)=X1|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
% 0.07/0.80  cnf(c_0_66, plain, (ssItem(esk5_1(X1))|~ssList(X1)|~singletonP(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
% 0.07/0.80  cnf(c_0_67, plain, (tl(cons(X1,X2))=X2|~ssItem(X1)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
% 0.07/0.80  cnf(c_0_68, negated_conjecture, (cons(esk56_0,nil)=esk50_0|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.07/0.80  cnf(c_0_69, negated_conjecture, (ssItem(esk56_0)|nil=esk50_0), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.07/0.80  cnf(c_0_70, plain, (singletonP(cons(X1,nil))|~ssList(cons(X1,nil))|~ssItem(X1)), inference(er,[status(thm)],[c_0_60])).
% 0.07/0.80  cnf(c_0_71, negated_conjecture, (ssList(esk50_0)), inference(rw,[status(thm)],[c_0_61, c_0_34])).
% 0.07/0.80  fof(c_0_72, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|hd(app(X3,X4))=hd(X3))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
% 0.07/0.80  fof(c_0_73, plain, (![X3]:![X4]:(((((nil=X4|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3))&(((nil=X3|nil!=app(X3,X4))|~ssList(X4))|~ssList(X3)))&((((nil!=X4|nil!=X3)|nil=app(X3,X4))|~ssList(X4))|~ssList(X3)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).
% 0.07/0.80  cnf(c_0_74, negated_conjecture, (frontsegP(esk50_0,esk53_0)|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_39])])).
% 0.07/0.80  fof(c_0_75, plain, (![X1]:(ssItem(X1)=>~memberP(nil,X1))), inference(fof_simplification,[status(thm)],[c_0_14])).
% 0.07/0.80  fof(c_0_76, plain, (![X5]:![X6]:![X9]:![X10]:(((((ssList(esk3_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&((((ssList(esk4_2(X5,X6))|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))&(((app(esk3_2(X5,X6),cons(X6,esk4_2(X5,X6)))=X5|~memberP(X5,X6))|~ssItem(X6))|~ssList(X5))))&((((~ssList(X9)|(~ssList(X10)|app(X9,cons(X6,X10))!=X5))|memberP(X5,X6))|~ssItem(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])])).
% 0.07/0.80  cnf(c_0_77, plain, (cons(esk5_1(X1),tl(X1))=X1|~singletonP(X1)|~ssList(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_45])]), c_0_66])).
% 0.07/0.80  cnf(c_0_78, negated_conjecture, (tl(esk50_0)=nil|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_45])]), c_0_69])).
% 0.07/0.80  cnf(c_0_79, negated_conjecture, (esk50_0=nil|singletonP(esk50_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_68]), c_0_71])]), c_0_69])).
% 0.07/0.80  cnf(c_0_80, plain, (hd(app(X1,X2))=hd(X1)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
% 0.07/0.80  cnf(c_0_81, plain, (nil=X1|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_73])).
% 0.07/0.80  cnf(c_0_82, plain, (app(X2,esk6_2(X1,X2))=X1|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 0.07/0.80  cnf(c_0_83, negated_conjecture, (frontsegP(esk50_0,esk53_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_44]), c_0_38]), c_0_46])])).
% 0.07/0.80  cnf(c_0_84, plain, (ssList(esk6_2(X1,X2))|~ssList(X1)|~ssList(X2)|~frontsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
% 0.07/0.80  fof(c_0_85, plain, (![X4]:![X5]:![X6]:(((((~memberP(cons(X5,X6),X4)|(X4=X5|memberP(X6,X4)))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&(((((X4!=X5|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(cons(X5,X6),X4))|~ssList(X6))|~ssItem(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])).
% 0.07/0.80  fof(c_0_86, plain, (![X2]:(~ssItem(X2)|~memberP(nil,X2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_75])])).
% 0.07/0.80  cnf(c_0_87, plain, (memberP(X1,X2)|~ssList(X1)|~ssItem(X2)|app(X3,cons(X2,X4))!=X1|~ssList(X4)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_76])).
% 0.07/0.80  fof(c_0_88, plain, (![X2]:(~ssList(X2)|(nil=X2|ssList(tl(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])).
% 0.07/0.80  fof(c_0_89, plain, (![X3]:![X4]:(~ssList(X3)|(~ssList(X4)|(nil=X3|tl(app(X3,X4))=app(tl(X3),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
% 0.07/0.80  cnf(c_0_90, plain, (cons(esk5_1(X1),X2)=app(X1,X2)|~singletonP(X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_65]), c_0_66])).
% 0.07/0.80  cnf(c_0_91, negated_conjecture, (cons(esk5_1(esk50_0),nil)=esk50_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_71])]), c_0_79])).
% 0.07/0.80  cnf(c_0_92, plain, (cons(hd(X1),tl(app(X1,X2)))=app(X1,X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_80]), c_0_53]), c_0_81])).
% 0.07/0.80  cnf(c_0_93, negated_conjecture, (app(esk53_0,esk6_2(esk50_0,esk53_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_39]), c_0_71])])).
% 0.07/0.80  cnf(c_0_94, negated_conjecture, (ssList(esk6_2(esk50_0,esk53_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_83]), c_0_39]), c_0_71])])).
% 0.07/0.80  cnf(c_0_95, negated_conjecture, (esk53_0=nil|esk50_0!=nil|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_63]), c_0_39])])).
% 0.07/0.80  cnf(c_0_96, plain, (memberP(X3,X1)|X1=X2|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|~memberP(cons(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
% 0.07/0.80  cnf(c_0_97, plain, (~memberP(nil,X1)|~ssItem(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
% 0.07/0.80  cnf(c_0_98, plain, (memberP(app(X1,cons(X2,X3)),X2)|~ssList(app(X1,cons(X2,X3)))|~ssList(X3)|~ssList(X1)|~ssItem(X2)), inference(er,[status(thm)],[c_0_87])).
% 0.07/0.80  cnf(c_0_99, plain, (ssList(tl(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
% 0.07/0.80  cnf(c_0_100, plain, (tl(app(X1,X2))=app(tl(X1),X2)|nil=X1|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_89])).
% 0.07/0.80  cnf(c_0_101, negated_conjecture, (app(esk50_0,nil)=esk50_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_45]), c_0_71])]), c_0_79])).
% 0.07/0.80  cnf(c_0_102, negated_conjecture, (cons(hd(esk53_0),tl(esk50_0))=esk50_0|esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94]), c_0_39])])).
% 0.07/0.80  cnf(c_0_103, negated_conjecture, (esk53_0=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_44]), c_0_38]), c_0_46])])).
% 0.07/0.80  fof(c_0_104, plain, (![X2]:(~ssList(X2)|(nil=X2|ssItem(hd(X2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])).
% 0.07/0.80  cnf(c_0_105, plain, (X1=esk5_1(X2)|~singletonP(X2)|~memberP(X2,X1)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_65]), c_0_45])]), c_0_66]), c_0_97])).
% 0.07/0.80  cnf(c_0_106, negated_conjecture, (memberP(esk50_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_63]), c_0_71]), c_0_38]), c_0_39]), c_0_46])])).
% 0.07/0.80  cnf(c_0_107, negated_conjecture, (hd(esk50_0)=esk56_0|esk50_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_68]), c_0_45])]), c_0_69])).
% 0.07/0.80  cnf(c_0_108, plain, (nil=X1|ssList(app(tl(X1),X2))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_53]), c_0_81])).
% 0.07/0.80  cnf(c_0_109, negated_conjecture, (app(tl(esk50_0),nil)=tl(esk50_0)|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_45]), c_0_71])])).
% 0.07/0.80  fof(c_0_110, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X5,X6)!=app(X5,X4)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
% 0.07/0.80  fof(c_0_111, plain, (![X2]:(~ssList(X2)|app(X2,nil)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])).
% 0.07/0.80  cnf(c_0_112, negated_conjecture, (cons(hd(esk53_0),nil)=esk50_0|esk53_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_78]), c_0_103])).
% 0.07/0.80  cnf(c_0_113, negated_conjecture, (hd(esk53_0)=hd(esk50_0)|esk53_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_93]), c_0_94]), c_0_39])])).
% 0.07/0.80  cnf(c_0_114, plain, (ssItem(hd(X1))|nil=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
% 0.07/0.80  cnf(c_0_115, negated_conjecture, (esk5_1(esk50_0)=esk52_0|~singletonP(esk50_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_71]), c_0_46])])).
% 0.07/0.80  cnf(c_0_116, negated_conjecture, (cons(esk56_0,tl(esk50_0))=esk50_0|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_107]), c_0_71])])).
% 0.07/0.80  cnf(c_0_117, negated_conjecture, (esk50_0=nil|ssList(tl(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_45]), c_0_71])])).
% 0.07/0.80  cnf(c_0_118, plain, (X1=X2|app(X3,X1)!=app(X3,X2)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_110])).
% 0.07/0.80  cnf(c_0_119, plain, (app(X1,nil)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_111])).
% 0.07/0.80  cnf(c_0_120, negated_conjecture, (cons(hd(esk50_0),nil)=esk50_0|esk53_0=nil), inference(spm,[status(thm)],[c_0_112, c_0_113])).
% 0.07/0.80  cnf(c_0_121, negated_conjecture, (esk53_0=nil|ssItem(hd(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_113]), c_0_39])])).
% 0.07/0.80  cnf(c_0_122, negated_conjecture, (esk5_1(esk50_0)=esk52_0|esk50_0=nil), inference(spm,[status(thm)],[c_0_115, c_0_79])).
% 0.07/0.80  cnf(c_0_123, negated_conjecture, (esk50_0=nil|X1=esk56_0|memberP(tl(esk50_0),X1)|~memberP(esk50_0,X1)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_116]), c_0_69]), c_0_117])).
% 0.07/0.80  cnf(c_0_124, plain, (X1=nil|app(X2,X1)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_45])])).
% 0.07/0.80  cnf(c_0_125, negated_conjecture, (cons(esk56_0,nil)=esk50_0|esk53_0=nil), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_107]), c_0_103])).
% 0.07/0.80  cnf(c_0_126, negated_conjecture, (esk53_0=nil|ssItem(esk56_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_107]), c_0_103])).
% 0.07/0.80  cnf(c_0_127, negated_conjecture, (cons(esk52_0,nil)=esk50_0|esk50_0=nil), inference(spm,[status(thm)],[c_0_91, c_0_122])).
% 0.07/0.80  cnf(c_0_128, negated_conjecture, (esk52_0=esk56_0|esk50_0=nil|memberP(tl(esk50_0),esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_106]), c_0_46])])).
% 0.07/0.80  fof(c_0_129, plain, (![X4]:![X5]:![X7]:(((((ssList(esk7_2(X4,X5))|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4))&(((app(esk7_2(X4,X5),X5)=X4|~rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))&((((~ssList(X7)|app(X7,X5)!=X4)|rearsegP(X4,X5))|~ssList(X5))|~ssList(X4)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])])])).
% 0.07/0.80  cnf(c_0_130, negated_conjecture, (cons(esk52_0,esk54_0)=nil|esk53_0!=esk50_0|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_63]), c_0_39])])).
% 0.07/0.80  cnf(c_0_131, negated_conjecture, (hd(esk50_0)=esk56_0|esk53_0=nil), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_125]), c_0_45])]), c_0_126])).
% 0.07/0.80  cnf(c_0_132, negated_conjecture, (hd(esk50_0)=esk52_0|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_127]), c_0_45]), c_0_46])])).
% 0.07/0.80  cnf(c_0_133, negated_conjecture, (esk50_0=nil|esk52_0=esk56_0|memberP(nil,esk52_0)), inference(spm,[status(thm)],[c_0_128, c_0_78])).
% 0.07/0.80  cnf(c_0_134, plain, (rearsegP(X1,X2)|~ssList(X1)|~ssList(X2)|app(X3,X2)!=X1|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_129])).
% 0.07/0.80  cnf(c_0_135, negated_conjecture, (cons(esk52_0,esk54_0)=nil|esk53_0!=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_44]), c_0_38]), c_0_46])])).
% 0.07/0.80  cnf(c_0_136, negated_conjecture, (esk53_0=nil|esk52_0=esk56_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_103])).
% 0.07/0.80  cnf(c_0_137, negated_conjecture, (esk52_0=esk56_0|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_133]), c_0_46])])).
% 0.07/0.80  cnf(c_0_138, plain, (rearsegP(app(X1,X2),X2)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_134]), c_0_53])).
% 0.07/0.80  cnf(c_0_139, negated_conjecture, (cons(esk52_0,esk54_0)=nil|esk52_0=esk56_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_136]), c_0_137])).
% 0.07/0.80  cnf(c_0_140, negated_conjecture, (rearsegP(esk50_0,esk54_0)|~ssList(app(esk53_0,cons(esk52_0,nil)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_36]), c_0_38])])).
% 0.07/0.80  cnf(c_0_141, negated_conjecture, (esk52_0=esk56_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_139]), c_0_38]), c_0_46])])).
% 0.07/0.80  cnf(c_0_142, negated_conjecture, (rearsegP(esk50_0,esk54_0)|~ssList(cons(esk52_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_53]), c_0_39])])).
% 0.07/0.80  cnf(c_0_143, negated_conjecture, (X1=esk6_2(esk50_0,esk53_0)|app(esk53_0,X1)!=esk50_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_93]), c_0_39]), c_0_94])])).
% 0.07/0.80  cnf(c_0_144, negated_conjecture, (app(esk53_0,cons(esk56_0,esk54_0))=esk50_0), inference(rw,[status(thm)],[c_0_63, c_0_141])).
% 0.07/0.80  fof(c_0_145, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(app(X6,X5)!=app(X4,X5)|X6=X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).
% 0.07/0.80  cnf(c_0_146, plain, (app(esk7_2(X1,X2),X2)=X1|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
% 0.07/0.80  cnf(c_0_147, negated_conjecture, (rearsegP(esk50_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_44]), c_0_45]), c_0_46])])).
% 0.07/0.80  cnf(c_0_148, plain, (ssList(esk7_2(X1,X2))|~ssList(X1)|~ssList(X2)|~rearsegP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
% 0.07/0.80  cnf(c_0_149, negated_conjecture, (esk6_2(esk50_0,esk53_0)=cons(esk56_0,esk54_0)|~ssList(cons(esk56_0,esk54_0))), inference(spm,[status(thm)],[c_0_143, c_0_144])).
% 0.07/0.80  cnf(c_0_150, negated_conjecture, (ssItem(esk56_0)), inference(rw,[status(thm)],[c_0_46, c_0_141])).
% 0.07/0.80  cnf(c_0_151, plain, (X1=X2|app(X1,X3)!=app(X2,X3)|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(split_conjunct,[status(thm)],[c_0_145])).
% 0.07/0.80  cnf(c_0_152, negated_conjecture, (app(esk7_2(esk50_0,esk54_0),esk54_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_38]), c_0_71])])).
% 0.07/0.80  cnf(c_0_153, negated_conjecture, (ssList(esk7_2(esk50_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_147]), c_0_38]), c_0_71])])).
% 0.07/0.80  fof(c_0_154, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssItem(X6)|cons(X6,app(X5,X4))=app(cons(X6,X5),X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
% 0.07/0.80  cnf(c_0_155, negated_conjecture, (esk6_2(esk50_0,esk53_0)=cons(esk56_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_44]), c_0_38]), c_0_150])])).
% 0.07/0.80  cnf(c_0_156, negated_conjecture, (app(app(esk53_0,cons(esk56_0,nil)),esk54_0)=esk50_0), inference(rw,[status(thm)],[c_0_36, c_0_141])).
% 0.07/0.80  cnf(c_0_157, negated_conjecture, (X1=esk7_2(esk50_0,esk54_0)|app(X1,esk54_0)!=esk50_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_38]), c_0_153])])).
% 0.07/0.80  cnf(c_0_158, plain, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), inference(split_conjunct,[status(thm)],[c_0_154])).
% 0.07/0.80  cnf(c_0_159, negated_conjecture, (X1=cons(esk56_0,esk54_0)|app(esk53_0,X1)!=esk50_0|~ssList(X1)), inference(rw,[status(thm)],[c_0_143, c_0_155])).
% 0.07/0.80  cnf(c_0_160, negated_conjecture, (app(esk53_0,app(cons(esk56_0,nil),esk54_0))=esk50_0), inference(rw,[status(thm)],[c_0_54, c_0_141])).
% 0.07/0.80  cnf(c_0_161, negated_conjecture, (app(app(esk53_0,esk50_0),esk54_0)=esk50_0|esk50_0=nil), inference(spm,[status(thm)],[c_0_156, c_0_68])).
% 0.07/0.80  cnf(c_0_162, negated_conjecture, (esk7_2(esk50_0,esk54_0)=app(esk53_0,cons(esk56_0,nil))|~ssList(app(esk53_0,cons(esk56_0,nil)))), inference(spm,[status(thm)],[c_0_157, c_0_156])).
% 0.07/0.80  cnf(c_0_163, plain, (nil=X2|~ssList(X1)|~ssList(X2)|nil!=app(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_73])).
% 0.07/0.80  cnf(c_0_164, negated_conjecture, (app(esk53_0,cons(esk52_0,app(nil,esk54_0)))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_158]), c_0_38]), c_0_45]), c_0_46])])).
% 0.07/0.80  cnf(c_0_165, negated_conjecture, (app(cons(esk56_0,nil),esk54_0)=cons(esk56_0,esk54_0)|~ssList(app(cons(esk56_0,nil),esk54_0))), inference(spm,[status(thm)],[c_0_159, c_0_160])).
% 0.07/0.80  fof(c_0_166, plain, (![X2]:(~ssList(X2)|app(nil,X2)=X2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])).
% 0.07/0.80  cnf(c_0_167, negated_conjecture, (esk7_2(esk50_0,esk54_0)=app(esk53_0,esk50_0)|esk50_0=nil|~ssList(app(esk53_0,esk50_0))), inference(spm,[status(thm)],[c_0_157, c_0_161])).
% 0.07/0.80  cnf(c_0_168, negated_conjecture, (esk7_2(esk50_0,esk54_0)=app(esk53_0,cons(esk56_0,nil))|~ssList(cons(esk56_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_53]), c_0_39])])).
% 0.07/0.80  cnf(c_0_169, negated_conjecture, (cons(esk52_0,app(nil,esk54_0))=nil|esk50_0!=nil|~ssList(cons(esk52_0,app(nil,esk54_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164]), c_0_39])])).
% 0.07/0.80  cnf(c_0_170, negated_conjecture, (app(cons(esk56_0,nil),esk54_0)=cons(esk56_0,esk54_0)|~ssList(cons(esk56_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_53]), c_0_38])])).
% 0.07/0.80  cnf(c_0_171, plain, (app(nil,X1)=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_166])).
% 0.07/0.80  cnf(c_0_172, negated_conjecture, (esk7_2(esk50_0,esk54_0)=app(esk53_0,esk50_0)|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_53]), c_0_71]), c_0_39])])).
% 0.07/0.80  cnf(c_0_173, negated_conjecture, (esk7_2(esk50_0,esk54_0)=app(esk53_0,cons(esk56_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_44]), c_0_45]), c_0_150])])).
% 0.07/0.80  cnf(c_0_174, negated_conjecture, (cons(esk52_0,app(nil,esk54_0))=nil|esk50_0!=nil|~ssList(app(nil,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_44]), c_0_46])])).
% 0.07/0.80  cnf(c_0_175, negated_conjecture, (app(cons(esk56_0,nil),esk54_0)=cons(esk56_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_44]), c_0_45]), c_0_150])])).
% 0.07/0.80  fof(c_0_176, plain, (![X4]:![X5]:![X6]:(~ssList(X4)|(~ssList(X5)|(~ssList(X6)|(~frontsegP(X4,X5)|frontsegP(app(X4,X6),X5)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
% 0.07/0.80  fof(c_0_177, plain, (![X2]:(~ssList(X2)|frontsegP(X2,nil))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])).
% 0.07/0.80  cnf(c_0_178, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_171]), c_0_45])])).
% 0.07/0.80  cnf(c_0_179, negated_conjecture, (app(esk53_0,cons(esk56_0,nil))=app(esk53_0,esk50_0)|esk50_0=nil), inference(spm,[status(thm)],[c_0_172, c_0_173])).
% 0.07/0.80  cnf(c_0_180, negated_conjecture, (cons(esk52_0,app(nil,esk54_0))=nil|esk50_0!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_53]), c_0_38]), c_0_45])])).
% 0.07/0.80  cnf(c_0_181, plain, (cons(X1,app(X2,X3))!=nil|~ssList(X3)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_158]), c_0_44]), c_0_58])).
% 0.07/0.80  cnf(c_0_182, negated_conjecture, (cons(esk56_0,app(nil,esk54_0))=cons(esk56_0,esk54_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_175]), c_0_38]), c_0_45]), c_0_150])])).
% 0.07/0.80  cnf(c_0_183, plain, (frontsegP(app(X1,X2),X3)|~frontsegP(X1,X3)|~ssList(X2)|~ssList(X3)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_176])).
% 0.07/0.80  cnf(c_0_184, plain, (frontsegP(X1,nil)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_177])).
% 0.07/0.80  cnf(c_0_185, negated_conjecture, (app(nil,app(esk53_0,esk50_0))=app(esk53_0,esk50_0)|esk50_0=nil|~ssList(cons(esk56_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_179]), c_0_39])])).
% 0.07/0.80  cnf(c_0_186, negated_conjecture, (cons(esk56_0,app(nil,esk54_0))=nil|esk50_0!=nil), inference(rw,[status(thm)],[c_0_180, c_0_141])).
% 0.07/0.80  cnf(c_0_187, negated_conjecture, (cons(esk56_0,esk54_0)!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_182]), c_0_38]), c_0_45]), c_0_150])])).
% 0.07/0.80  cnf(c_0_188, plain, (frontsegP(app(X1,X2),nil)|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_184]), c_0_45])])).
% 0.07/0.80  cnf(c_0_189, plain, (memberP(cons(X2,X3),X1)|~ssItem(X1)|~ssItem(X2)|~ssList(X3)|X1!=X2), inference(split_conjunct,[status(thm)],[c_0_85])).
% 0.07/0.80  fof(c_0_190, plain, (![X4]:(((ssList(esk44_1(X4))|nil=X4)|~ssList(X4))&(((ssItem(esk45_1(X4))|nil=X4)|~ssList(X4))&((cons(esk45_1(X4),esk44_1(X4))=X4|nil=X4)|~ssList(X4))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])])).
% 0.07/0.80  cnf(c_0_191, negated_conjecture, (rearsegP(esk50_0,cons(esk52_0,esk54_0))|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_63]), c_0_39])])).
% 0.07/0.80  cnf(c_0_192, negated_conjecture, (app(nil,app(esk53_0,esk50_0))=app(esk53_0,esk50_0)|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_68]), c_0_71])])).
% 0.07/0.80  cnf(c_0_193, negated_conjecture, (esk50_0!=nil), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_186, c_0_182]), c_0_187])).
% 0.07/0.80  cnf(c_0_194, negated_conjecture, (frontsegP(esk50_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_93]), c_0_94]), c_0_39])])).
% 0.07/0.80  cnf(c_0_195, plain, (memberP(cons(X1,X2),X1)|~ssList(X2)|~ssItem(X1)), inference(er,[status(thm)],[c_0_189])).
% 0.07/0.80  cnf(c_0_196, plain, (nil=X1|cons(esk45_1(X1),esk44_1(X1))=X1|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_190])).
% 0.07/0.80  cnf(c_0_197, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_190])).
% 0.07/0.80  cnf(c_0_198, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_190])).
% 0.07/0.80  cnf(c_0_199, negated_conjecture, (rearsegP(esk50_0,cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_44]), c_0_38]), c_0_46])])).
% 0.07/0.80  cnf(c_0_200, negated_conjecture, (app(tl(esk7_2(esk50_0,esk54_0)),esk54_0)=tl(esk50_0)|esk7_2(esk50_0,esk54_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_152]), c_0_38]), c_0_153])])).
% 0.07/0.80  cnf(c_0_201, plain, (nil=X1|app(X2,app(X3,X1))!=nil|~ssList(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_37]), c_0_53])).
% 0.07/0.80  cnf(c_0_202, negated_conjecture, (app(nil,app(esk53_0,esk50_0))=app(esk53_0,esk50_0)), inference(sr,[status(thm)],[c_0_192, c_0_193])).
% 0.07/0.80  cnf(c_0_203, negated_conjecture, (app(nil,esk6_2(esk50_0,nil))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_194]), c_0_45]), c_0_71])])).
% 0.07/0.80  cnf(c_0_204, negated_conjecture, (ssList(esk6_2(esk50_0,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_194]), c_0_45]), c_0_71])])).
% 0.07/0.80  cnf(c_0_205, negated_conjecture, (esk50_0=nil|X1=esk56_0|~memberP(esk50_0,X1)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_68]), c_0_45])]), c_0_69]), c_0_97])).
% 0.07/0.80  cnf(c_0_206, plain, (nil=X1|memberP(X1,esk45_1(X1))|~ssList(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196]), c_0_197]), c_0_198])).
% 0.07/0.80  fof(c_0_207, plain, (![X4]:![X5]:![X6]:(((((~memberP(app(X5,X6),X4)|(memberP(X5,X4)|memberP(X6,X4)))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&(((((~memberP(X5,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))&((((~memberP(X6,X4)|memberP(app(X5,X6),X4))|~ssList(X6))|~ssList(X5))|~ssItem(X4))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])])).
% 0.07/0.80  cnf(c_0_208, negated_conjecture, (app(esk7_2(esk50_0,cons(esk52_0,esk54_0)),cons(esk52_0,esk54_0))=esk50_0|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_199]), c_0_71])])).
% 0.07/0.80  cnf(c_0_209, negated_conjecture, (ssList(esk7_2(esk50_0,cons(esk52_0,esk54_0)))|~ssList(cons(esk52_0,esk54_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_199]), c_0_71])])).
% 0.07/0.80  cnf(c_0_210, negated_conjecture, (app(tl(app(esk53_0,cons(esk56_0,nil))),esk54_0)=tl(esk50_0)|app(esk53_0,cons(esk56_0,nil))=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_200, c_0_173]), c_0_173])).
% 0.07/0.80  cnf(c_0_211, negated_conjecture, (cons(esk56_0,nil)=esk50_0), inference(sr,[status(thm)],[c_0_68, c_0_193])).
% 0.07/0.80  cnf(c_0_212, negated_conjecture, (tl(esk50_0)=nil), inference(sr,[status(thm)],[c_0_78, c_0_193])).
% 0.07/0.80  cnf(c_0_213, negated_conjecture, (app(esk53_0,esk50_0)!=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_202]), c_0_71]), c_0_39]), c_0_45])]), c_0_193])).
% 0.07/0.80  cnf(c_0_214, negated_conjecture, (esk6_2(esk50_0,nil)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_203]), c_0_204])])).
% 0.07/0.80  fof(c_0_215, plain, (![X5]:![X6]:![X7]:![X8]:((((((X7=X8|cons(X7,X5)!=cons(X8,X6))|~ssItem(X8))|~ssItem(X7))|~ssList(X6))|~ssList(X5))&(((((X6=X5|cons(X7,X5)!=cons(X8,X6))|~ssItem(X8))|~ssItem(X7))|~ssList(X6))|~ssList(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])).
% 0.07/0.80  cnf(c_0_216, negated_conjecture, (esk45_1(esk50_0)=esk56_0|esk50_0=nil|~ssItem(esk45_1(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206]), c_0_71])])).
% 0.07/0.80  cnf(c_0_217, plain, (memberP(app(X2,X3),X1)|~ssItem(X1)|~ssList(X2)|~ssList(X3)|~memberP(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_207])).
% 0.07/0.80  cnf(c_0_218, negated_conjecture, (memberP(esk53_0,esk55_0)|memberP(esk54_0,esk55_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.07/0.80  cnf(c_0_219, negated_conjecture, (ssItem(esk55_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
% 0.07/0.80  cnf(c_0_220, negated_conjecture, (X1=esk53_0|app(X1,esk6_2(esk50_0,esk53_0))!=esk50_0|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_93]), c_0_94]), c_0_39])])).
% 0.07/0.80  cnf(c_0_221, negated_conjecture, (app(esk7_2(esk50_0,cons(esk52_0,esk54_0)),cons(esk52_0,esk54_0))=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_44]), c_0_38]), c_0_46])])).
% 0.07/0.80  cnf(c_0_222, negated_conjecture, (ssList(esk7_2(esk50_0,cons(esk52_0,esk54_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_44]), c_0_38]), c_0_46])])).
% 0.07/0.80  cnf(c_0_223, negated_conjecture, (app(tl(app(esk53_0,esk50_0)),esk54_0)=nil), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_210, c_0_211]), c_0_211]), c_0_212]), c_0_213])).
% 0.07/0.80  cnf(c_0_224, negated_conjecture, (esk50_0=nil|ssList(app(esk53_0,esk50_0))), inference(spm,[status(thm)],[c_0_153, c_0_172])).
% 0.07/0.80  cnf(c_0_225, plain, (app(esk7_2(app(X1,X2),X2),X2)=app(X1,X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_138]), c_0_53])).
% 0.07/0.80  cnf(c_0_226, negated_conjecture, (app(nil,esk50_0)=esk50_0), inference(rw,[status(thm)],[c_0_203, c_0_214])).
% 0.07/0.80  cnf(c_0_227, plain, (ssList(esk7_2(app(X1,X2),X2))|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_138]), c_0_53])).
% 0.07/0.80  cnf(c_0_228, plain, (X2=X1|~ssList(X1)|~ssList(X2)|~ssItem(X3)|~ssItem(X4)|cons(X3,X1)!=cons(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_215])).
% 0.07/0.80  cnf(c_0_229, negated_conjecture, (esk45_1(esk50_0)=esk56_0|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_197]), c_0_71])])).
% 0.07/0.80  cnf(c_0_230, negated_conjecture, (memberP(app(X1,esk54_0),esk55_0)|memberP(esk53_0,esk55_0)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_218]), c_0_38]), c_0_219])])).
% 0.07/0.80  cnf(c_0_231, negated_conjecture, (app(esk50_0,X1)=cons(esk56_0,X1)|esk50_0=nil|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_68]), c_0_69])).
% 0.07/0.80  cnf(c_0_232, negated_conjecture, (X1=esk53_0|app(X1,cons(esk56_0,esk54_0))!=esk50_0|~ssList(X1)), inference(rw,[status(thm)],[c_0_220, c_0_155])).
% 0.07/0.80  cnf(c_0_233, negated_conjecture, (app(esk7_2(esk50_0,cons(esk56_0,esk54_0)),cons(esk56_0,esk54_0))=esk50_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_221, c_0_141]), c_0_141])).
% 0.07/0.80  cnf(c_0_234, negated_conjecture, (ssList(esk7_2(esk50_0,cons(esk56_0,esk54_0)))), inference(rw,[status(thm)],[c_0_222, c_0_141])).
% 0.07/0.80  cnf(c_0_235, negated_conjecture, (esk54_0=nil|~ssList(tl(app(esk53_0,esk50_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_223]), c_0_38])])).
% 0.07/0.80  cnf(c_0_236, negated_conjecture, (ssList(app(esk53_0,esk50_0))), inference(sr,[status(thm)],[c_0_224, c_0_193])).
% 0.07/0.80  cnf(c_0_237, plain, (X1=nil|app(X1,X2)!=X2|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_171]), c_0_45])])).
% 0.07/0.80  cnf(c_0_238, negated_conjecture, (app(esk7_2(esk50_0,esk50_0),esk50_0)=esk50_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_226]), c_0_71]), c_0_45])])).
% 0.07/0.80  cnf(c_0_239, negated_conjecture, (ssList(esk7_2(esk50_0,esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_226]), c_0_71]), c_0_45])])).
% 0.07/0.80  cnf(c_0_240, negated_conjecture, (esk50_0=nil|X1=nil|cons(X2,X1)!=esk50_0|~ssList(X1)|~ssItem(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_68]), c_0_45])]), c_0_69])).
% 0.07/0.80  cnf(c_0_241, negated_conjecture, (cons(esk56_0,esk44_1(esk50_0))=esk50_0|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_229]), c_0_71])])).
% 0.07/0.80  cnf(c_0_242, negated_conjecture, (esk50_0=nil|memberP(cons(esk56_0,esk54_0),esk55_0)|memberP(esk53_0,esk55_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_231]), c_0_71]), c_0_38])])).
% 0.07/0.80  cnf(c_0_243, negated_conjecture, (esk7_2(esk50_0,cons(esk56_0,esk54_0))=esk53_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_233]), c_0_234])])).
% 0.07/0.80  cnf(c_0_244, negated_conjecture, (esk54_0=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_99]), c_0_236])]), c_0_213])).
% 0.07/0.80  cnf(c_0_245, negated_conjecture, (esk7_2(esk50_0,esk50_0)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_238]), c_0_71]), c_0_239])])).
% 0.07/0.80  cnf(c_0_246, negated_conjecture, (esk44_1(esk50_0)=nil|esk50_0=nil|~ssList(esk44_1(esk50_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_241]), c_0_150])])).
% 0.07/0.80  cnf(c_0_247, negated_conjecture, (memberP(cons(esk56_0,esk54_0),esk55_0)|memberP(esk53_0,esk55_0)), inference(sr,[status(thm)],[c_0_242, c_0_193])).
% 0.07/0.80  cnf(c_0_248, negated_conjecture, (esk53_0=nil), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_243, c_0_244]), c_0_211]), c_0_245])).
% 0.07/0.80  cnf(c_0_249, negated_conjecture, (esk44_1(esk50_0)=nil|esk50_0=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_198]), c_0_71])])).
% 0.07/0.80  cnf(c_0_250, plain, (X1=esk45_1(X2)|nil=X2|memberP(esk44_1(X2),X1)|~memberP(X2,X1)|~ssList(X2)|~ssItem(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_196]), c_0_197]), c_0_198])).
% 0.07/0.80  cnf(c_0_251, negated_conjecture, (memberP(nil,esk55_0)|memberP(esk50_0,esk55_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_247, c_0_244]), c_0_211]), c_0_248])).
% 0.07/0.80  cnf(c_0_252, negated_conjecture, (esk45_1(esk50_0)=esk56_0), inference(sr,[status(thm)],[c_0_229, c_0_193])).
% 0.07/0.80  cnf(c_0_253, negated_conjecture, (esk44_1(esk50_0)=nil), inference(sr,[status(thm)],[c_0_249, c_0_193])).
% 0.07/0.80  cnf(c_0_254, negated_conjecture, (esk55_0=esk56_0|memberP(nil,esk55_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_251]), c_0_252]), c_0_253]), c_0_71]), c_0_219])]), c_0_193])).
% 0.07/0.80  cnf(c_0_255, negated_conjecture, (esk55_0=esk56_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_254]), c_0_219])])).
% 0.07/0.80  cnf(c_0_256, negated_conjecture, (memberP(nil,esk56_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_218, c_0_244]), c_0_248]), c_0_255]), c_0_255])])).
% 0.07/0.80  cnf(c_0_257, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_256]), c_0_150])]), ['proof']).
% 0.07/0.80  # SZS output end CNFRefutation
% 0.07/0.80  # Training examples: 0 positive, 0 negative
% 0.07/0.80  
% 0.07/0.80  # -------------------------------------------------
% 0.07/0.80  # User time                : 0.560 s
% 0.07/0.80  # System time              : 0.010 s
% 0.07/0.80  # Total time               : 0.570 s
% 0.07/0.80  # Maximum resident set size: 1880 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
