<!DOCTYPE html
	PUBLIC "-//W3C//DTD XHTML 1.0 Transitional//EN"
	 "http://www.w3.org/TR/xhtml1/DTD/xhtml1-transitional.dtd">
<html xmlns="http://www.w3.org/1999/xhtml" lang="en-US" xml:lang="en-US">
<head>
<title>TSTP Solution File: SWC227-1 by E---1.9.1</title>
<meta name="ROBOTS" content="NOINDEX,NOFOLLOW" />
<meta http-equiv="Content-Type" content="text/html; charset=iso-8859-1" />
</head>
<body>
<HR><CENTER>
<H2>TSTP Solution File: SWC227-1 by E---1.9.1</H2><A HREF=SeeTPTP?Category=Problems&Domain=SWC&File=SWC227-1.p>View Problem</A>
- <A HREF="http://www.cs.miami.edu/~tptp/cgi-bin/SystemOnTSTP?TPTPProblem=SWC227-1&SystemAndVersion=E---1.9.1" TARGET="_blank">Process Solution</A>
</CENTER><HR>
<pre>
%------------------------------------------------------------------------------
% File       : E---1.9.1
% Problem    : <a href="SeeTPTP?Category=Problems&Domain=SWC&File=SWC227-1.p">SWC227-1</a> : TPTP v6.4.0. Released v2.4.0.
% Transform  : none
% Format     : tptp:raw
% Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s

% Computer   : n128.star.cs.uiowa.edu
% Model      : x86_64 x86_64
% CPU        : Intel(R) Xeon(R) CPU E5-2609 0 2.40GHz
% Memory     : 32218.75MB
% OS         : Linux 3.10.0-327.10.1.el7.x86_64
% CPULimit   : 300s
% DateTime   : Wed May 11 11:17:07 EDT 2016

% Result     : Unsatisfiable 1.12s
% Output     : CNFRefutation 1.12s
% Verified   : 
% Statistics : Number of clauses        :  108 (6099 expanded)
%              Number of leaves         :   43 (4156 expanded)
%              Depth                    :   27
%              Number of atoms          :  268 (10967 expanded)
%              Number of equality atoms :   85 (3224 expanded)
%              Maximal clause size      :    7 (   2 average)
%              Maximal term depth       :    5 (   2 average)

% Comments   : 
%------------------------------------------------------------------------------
<A NAME="c_0_0"></A>cnf(c_0_0,negated_conjecture,
    ( app(app(sk6,sk3),sk7) = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_15)).

<A NAME="c_0_1"></A>cnf(c_0_1,negated_conjecture,
    ( sk1 = sk3 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_6)).

<A NAME="c_0_2"></A>cnf(c_0_2,axiom,
    ( frontsegP(X3,X1)
    | app(X1,X2) != X3
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause144)).

<A NAME="c_0_3"></A>cnf(c_0_3,axiom,
    ( ssList(app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause85)).

<A NAME="c_0_4"></A>cnf(c_0_4,axiom,
    ( app(app(X3,X2),X1) = app(X3,app(X2,X1))
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause149)).

<A NAME="c_0_5"></A>cnf(c_0_5,negated_conjecture,
    ( ssList(sk6) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_13)).

<A NAME="c_0_6"></A>cnf(c_0_6,negated_conjecture,
    ( ssList(sk1) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_1)).

<A NAME="c_0_7"></A>cnf(c_0_7,negated_conjecture,
    ( ssList(sk7) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_14)).

<A NAME="c_0_8"></A>cnf(c_0_8,negated_conjecture,
    ( ssList(sk2) ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_2)).

<A NAME="c_0_9"></A>cnf(c_0_9,negated_conjecture,
    ( sk2 = sk4 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_5)).

<A NAME="c_0_10"></A>cnf(c_0_10,axiom,
    ( X2 = X3
    | app(X1,X2) != app(X1,X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause150)).

<A NAME="c_0_11"></A>cnf(c_0_11,axiom,
    ( app(X2,skaf45(X1,X2)) = X1
    | ~ frontsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause132)).

<A NAME="c_0_12"></A>cnf(c_0_12,axiom,
    ( ssList(skaf45(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause51)).

<A NAME="c_0_13"></A>cnf(c_0_13,axiom,
    ( rearsegP(X3,X2)
    | app(X1,X2) != X3
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause143)).

<A NAME="c_0_14"></A>cnf(c_0_14,axiom,
    ( frontsegP(X1,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause61)).

<A NAME="c_0_15"></A>cnf(c_0_15,axiom,
    ( app(X1,nil) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause73)).

<A NAME="c_0_16"></A>cnf(c_0_16,axiom,
    ( ssList(nil) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause8)).

<A NAME="c_0_17"></A>cnf(c_0_17,axiom,
    ( X1 = X3
    | app(X1,X2) != app(X3,X2)
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause151)).

<A NAME="c_0_18"></A>cnf(c_0_18,axiom,
    ( app(skaf46(X1,X2),X2) = X1
    | ~ rearsegP(X1,X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause131)).

<A NAME="c_0_19"></A>cnf(c_0_19,axiom,
    ( ssList(skaf46(X1,X2)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause50)).

<A NAME="c_0_20"></A>cnf(c_0_20,axiom,
    ( frontsegP(app(X1,X3),X2)
    | ~ frontsegP(X1,X2)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause137)).

<A NAME="c_0_21"></A>cnf(c_0_21,axiom,
    ( nil = X2
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause119)).

<A NAME="c_0_22"></A>cnf(c_0_22,axiom,
    ( nil = X1
    | app(X1,X2) != nil
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause118)).

<A NAME="c_0_23"></A>cnf(c_0_23,negated_conjecture,
    ( nil != sk1 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_7)).

<A NAME="c_0_24"></A>cnf(c_0_24,axiom,
    ( cons(hd(X1),tl(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause104)).

<A NAME="c_0_25"></A>cnf(c_0_25,axiom,
    ( nil = X2
    | hd(app(X2,X1)) = hd(X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause123)).

<A NAME="c_0_26"></A>cnf(c_0_26,axiom,
    ( rearsegP(X1,X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause59)).

<A NAME="c_0_27"></A>cnf(c_0_27,axiom,
    ( ssList(tl(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause75)).

<A NAME="c_0_28"></A>cnf(c_0_28,axiom,
    ( nil = X2
    | tl(app(X2,X1)) = app(tl(X2),X1)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause133)).

<A NAME="c_0_29"></A>cnf(c_0_29,axiom,
    ( app(app(X1,cons(X2,X3)),cons(X2,X4)) != X5
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssItem(X2)
    | ~ duplicatefreeP(X5)
    | ~ ssList(X5) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause179)).

<A NAME="c_0_30"></A>cnf(c_0_30,axiom,
    ( ssItem(hd(X1))
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause76)).

<A NAME="c_0_31"></A>cnf(c_0_31,axiom,
    ( cons(X1,app(X2,X3)) = app(cons(X1,X2),X3)
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause157)).

<A NAME="c_0_32"></A>cnf(c_0_32,axiom,
    ( ssList(cons(X1,X2))
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause86)).

<A NAME="c_0_33"></A>cnf(c_0_33,axiom,
    ( app(X1,X2) = nil
    | nil != X1
    | nil != X2
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause145)).

<A NAME="c_0_34"></A>cnf(c_0_34,axiom,
    ( duplicatefreeP(X1)
    | ssItem(X2)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause72)).

<A NAME="c_0_35"></A>cnf(c_0_35,negated_conjecture,
    ( memberP(X2,sk5(X3,X2,X1))
    | ~ ssItem(X1)
    | ~ ssList(X2)
    | ~ ssList(X3)
    | app(app(X2,cons(X1,nil)),X3) != sk1 ),
    file('/export/starexec/sandbox2/benchmark/theBenchmark.p',co1_9)).

<A NAME="c_0_36"></A>cnf(c_0_36,axiom,
    ( app(nil,X1) = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause74)).

<A NAME="c_0_37"></A>cnf(c_0_37,axiom,
    ( ~ memberP(nil,X1)
    | ~ ssItem(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause71)).

<A NAME="c_0_38"></A>cnf(c_0_38,axiom,
    ( singletonP(X2)
    | cons(X1,nil) != X2
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause116)).

<A NAME="c_0_39"></A>cnf(c_0_39,axiom,
    ( cons(skaf44(X1),nil) = X1
    | ~ singletonP(X1)
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause101)).

<A NAME="c_0_40"></A>cnf(c_0_40,axiom,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssItem(X1)
    | ~ ssList(X2) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause120)).

<A NAME="c_0_41"></A>cnf(c_0_41,axiom,
    ( cons(skaf83(X1),skaf82(X1)) = X1
    | nil = X1
    | ~ ssList(X1) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause109)).

<A NAME="c_0_42"></A>cnf(c_0_42,axiom,
    ( ssList(skaf82(X1)) ),
    file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax',clause13)).

<A NAME="c_0_43"></A>cnf(c_0_43,negated_conjecture,
    ( app(app(sk6,sk1),sk7) = sk4 ),
    inference(rw,[status(thm)],[<A HREF="#c_0_0">c_0_0</A>,<A HREF="#c_0_1">c_0_1</A>])).

<A NAME="c_0_44"></A>cnf(c_0_44,plain,
    ( frontsegP(app(X1,X2),X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_2]),<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_45"></A>cnf(c_0_45,negated_conjecture,
    ( app(sk6,app(sk1,sk7)) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_43">c_0_43</A>,c_0_4]),c_0_5]),c_0_6]),c_0_7])])).

<A NAME="c_0_46"></A>cnf(c_0_46,negated_conjecture,
    ( frontsegP(sk4,sk6)
    | ~ ssList(app(sk1,sk7)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_44">c_0_44</A>,c_0_45]),c_0_5])])).

<A NAME="c_0_47"></A>cnf(c_0_47,negated_conjecture,
    ( frontsegP(sk4,sk6) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_46">c_0_46</A>,c_0_3]),c_0_6]),c_0_7])])).

<A NAME="c_0_48"></A>cnf(c_0_48,negated_conjecture,
    ( ssList(sk4) ),
    inference(rw,[status(thm)],[<A HREF="#c_0_8">c_0_8</A>,<A HREF="#c_0_9">c_0_9</A>])).

<A NAME="c_0_49"></A>cnf(c_0_49,negated_conjecture,
    ( X1 = app(sk1,sk7)
    | app(sk6,X1) != sk4
    | ~ ssList(app(sk1,sk7))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_45]),c_0_5])])).

<A NAME="c_0_50"></A>cnf(c_0_50,negated_conjecture,
    ( app(sk6,skaf45(sk4,sk6)) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_47]),c_0_5]),c_0_48])])).

<A NAME="c_0_51"></A>cnf(c_0_51,negated_conjecture,
    ( skaf45(sk4,sk6) = app(sk1,sk7)
    | ~ ssList(app(sk1,sk7)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_49">c_0_49</A>,c_0_50]),c_0_12])])).

<A NAME="c_0_52"></A>cnf(c_0_52,negated_conjecture,
    ( skaf45(sk4,sk6) = app(sk1,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_51">c_0_51</A>,c_0_3]),c_0_6]),c_0_7])])).

<A NAME="c_0_53"></A>cnf(c_0_53,plain,
    ( rearsegP(app(X1,X2),X2)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]),<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_54"></A>cnf(c_0_54,negated_conjecture,
    ( app(sk6,app(app(sk1,sk7),X1)) = app(sk4,X1)
    | ~ ssList(app(sk1,sk7))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,c_0_45]),c_0_5])])).

<A NAME="c_0_55"></A>cnf(c_0_55,negated_conjecture,
    ( ssList(app(sk1,sk7)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_12">c_0_12</A>,<A HREF="#c_0_52">c_0_52</A>])).

<A NAME="c_0_56"></A>cnf(c_0_56,negated_conjecture,
    ( rearsegP(sk4,sk7)
    | ~ ssList(app(sk6,sk1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_53">c_0_53</A>,c_0_43]),c_0_7])])).

<A NAME="c_0_57"></A>cnf(c_0_57,negated_conjecture,
    ( app(sk6,app(app(sk1,sk7),X1)) = app(sk4,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_54">c_0_54</A>,c_0_55])])).

<A NAME="c_0_58"></A>cnf(c_0_58,plain,
    ( app(X1,skaf45(X1,X1)) = X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,<A HREF="#c_0_14">c_0_14</A>])).

<A NAME="c_0_59"></A>cnf(c_0_59,negated_conjecture,
    ( rearsegP(sk4,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_56">c_0_56</A>,c_0_3]),c_0_5]),c_0_6])])).

<A NAME="c_0_60"></A>cnf(c_0_60,plain,
    ( X1 = nil
    | app(X2,X1) != X2
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_10">c_0_10</A>,c_0_15]),c_0_16])])).

<A NAME="c_0_61"></A>cnf(c_0_61,negated_conjecture,
    ( app(sk4,skaf45(app(sk1,sk7),app(sk1,sk7))) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_57">c_0_57</A>,c_0_58]),c_0_45]),c_0_12]),c_0_55])])).

<A NAME="c_0_62"></A>cnf(c_0_62,negated_conjecture,
    ( X1 = app(sk6,sk1)
    | app(X1,sk7) != sk4
    | ~ ssList(app(sk6,sk1))
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_17">c_0_17</A>,c_0_43]),c_0_7])])).

<A NAME="c_0_63"></A>cnf(c_0_63,negated_conjecture,
    ( app(skaf46(sk4,sk7),sk7) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,c_0_59]),c_0_7]),c_0_48])])).

<A NAME="c_0_64"></A>cnf(c_0_64,negated_conjecture,
    ( skaf45(app(sk1,sk7),app(sk1,sk7)) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_60">c_0_60</A>,c_0_61]),c_0_12]),c_0_48])])).

<A NAME="c_0_65"></A>cnf(c_0_65,negated_conjecture,
    ( skaf46(sk4,sk7) = app(sk6,sk1)
    | ~ ssList(app(sk6,sk1)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_62">c_0_62</A>,c_0_63]),c_0_19])])).

<A NAME="c_0_66"></A>cnf(c_0_66,plain,
    ( frontsegP(app(app(X1,X2),X3),X1)
    | ~ ssList(X3)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_20">c_0_20</A>,c_0_44]),<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_67"></A>cnf(c_0_67,negated_conjecture,
    ( app(app(sk1,sk7),nil) = app(sk1,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_58">c_0_58</A>,c_0_64]),c_0_55])])).

<A NAME="c_0_68"></A>cnf(c_0_68,negated_conjecture,
    ( skaf46(sk4,sk7) = app(sk6,sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_65">c_0_65</A>,c_0_3]),c_0_5]),c_0_6])])).

<A NAME="c_0_69"></A>cnf(c_0_69,negated_conjecture,
    ( frontsegP(app(sk1,sk7),sk1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_66">c_0_66</A>,c_0_67]),c_0_16]),c_0_6]),c_0_7])])).

<A NAME="c_0_70"></A>cnf(c_0_70,negated_conjecture,
    ( ssList(app(sk6,sk1)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_19">c_0_19</A>,<A HREF="#c_0_68">c_0_68</A>])).

<A NAME="c_0_71"></A>cnf(c_0_71,negated_conjecture,
    ( app(sk1,skaf45(app(sk1,sk7),sk1)) = app(sk1,sk7) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_11">c_0_11</A>,c_0_69]),c_0_6]),c_0_55])])).

<A NAME="c_0_72"></A>cnf(c_0_72,negated_conjecture,
    ( app(app(sk6,sk1),app(sk7,X1)) = app(sk4,X1)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,c_0_43]),c_0_7])]),c_0_70])])).

<A NAME="c_0_73"></A>cnf(c_0_73,plain,
    ( nil = X1
    | app(X2,app(X3,X1)) != nil
    | ~ ssList(X1)
    | ~ ssList(X2)
    | ~ ssList(X3) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_21">c_0_21</A>,c_0_4]),<A HREF="#c_0_3">c_0_3</A>])).

<A NAME="c_0_74"></A>cnf(c_0_74,negated_conjecture,
    ( app(sk1,sk7) != nil ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_22">c_0_22</A>,c_0_71]),c_0_12]),c_0_6])]),<A HREF="#c_0_23">c_0_23</A>])).

<A NAME="c_0_75"></A>cnf(c_0_75,negated_conjecture,
    ( app(sk4,nil) = sk4 ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_72">c_0_72</A>,c_0_15]),c_0_43]),c_0_16]),c_0_7])])).

<A NAME="c_0_76"></A>cnf(c_0_76,plain,
    ( cons(hd(X1),tl(app(X1,X2))) = app(X1,X2)
    | nil = X1
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_24">c_0_24</A>,c_0_25]),c_0_3]),<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_77"></A>cnf(c_0_77,negated_conjecture,
    ( app(X1,sk4) != nil
    | ~ ssList(X1) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_73">c_0_73</A>,c_0_45]),c_0_55]),c_0_5])]),<A HREF="#c_0_74">c_0_74</A>])).

<A NAME="c_0_78"></A>cnf(c_0_78,plain,
    ( app(skaf46(X1,X1),X1) = X1
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_18">c_0_18</A>,<A HREF="#c_0_26">c_0_26</A>])).

<A NAME="c_0_79"></A>cnf(c_0_79,plain,
    ( nil = X1
    | ssList(app(tl(X1),X2))
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_27">c_0_27</A>,c_0_28]),c_0_3]),<A HREF="#c_0_22">c_0_22</A>])).

<A NAME="c_0_80"></A>cnf(c_0_80,negated_conjecture,
    ( app(tl(sk4),nil) = tl(sk4)
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_28">c_0_28</A>,c_0_75]),c_0_48]),c_0_16])])).

<A NAME="c_0_81"></A>cnf(c_0_81,plain,
    ( ~ ssItem(X1)
    | ~ ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))
    | ~ ssList(X4)
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4))) ),
    inference(er,[status(thm)],[<A HREF="#c_0_29">c_0_29</A>])).

<A NAME="c_0_82"></A>cnf(c_0_82,negated_conjecture,
    ( cons(hd(sk4),tl(sk4)) = sk4
    | sk4 = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_76">c_0_76</A>,c_0_75]),c_0_48]),c_0_16])])).

<A NAME="c_0_83"></A>cnf(c_0_83,negated_conjecture,
    ( sk4 != nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_77">c_0_77</A>,c_0_78]),c_0_19]),c_0_48])])).

<A NAME="c_0_84"></A>cnf(c_0_84,negated_conjecture,
    ( sk4 = nil
    | ssList(tl(sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_79">c_0_79</A>,c_0_80]),c_0_48]),c_0_16])])).

<A NAME="c_0_85"></A>cnf(c_0_85,plain,
    ( nil = X1
    | ~ ssList(app(app(X2,cons(hd(X1),X3)),X1))
    | ~ ssList(X3)
    | ~ ssList(X2)
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(app(X2,cons(hd(X1),X3)),X1)) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_81">c_0_81</A>,c_0_24]),c_0_27]),<A HREF="#c_0_30">c_0_30</A>])).

<A NAME="c_0_86"></A>cnf(c_0_86,negated_conjecture,
    ( cons(hd(sk4),tl(sk4)) = sk4 ),
    inference(sr,[status(thm)],[<A HREF="#c_0_82">c_0_82</A>,<A HREF="#c_0_83">c_0_83</A>])).

<A NAME="c_0_87"></A>cnf(c_0_87,negated_conjecture,
    ( ssList(tl(sk4)) ),
    inference(sr,[status(thm)],[<A HREF="#c_0_84">c_0_84</A>,<A HREF="#c_0_83">c_0_83</A>])).

<A NAME="c_0_88"></A>cnf(c_0_88,negated_conjecture,
    ( ~ ssList(app(app(X1,sk4),sk4))
    | ~ ssList(X1)
    | ~ duplicatefreeP(app(app(X1,sk4),sk4)) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_85">c_0_85</A>,c_0_86]),c_0_87]),c_0_48])]),<A HREF="#c_0_83">c_0_83</A>])).

<A NAME="c_0_89"></A>cnf(c_0_89,plain,
    ( ssList(cons(X1,app(X2,X3)))
    | ~ ssItem(X1)
    | ~ ssList(X3)
    | ~ ssList(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_3">c_0_3</A>,c_0_31]),<A HREF="#c_0_32">c_0_32</A>])).

<A NAME="c_0_90"></A>cnf(c_0_90,plain,
    ( app(nil,nil) = nil ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_33])]),c_0_16])])).

<A NAME="c_0_91"></A>cnf(c_0_91,negated_conjecture,
    ( ~ ssList(app(sk4,sk4))
    | ~ duplicatefreeP(app(sk4,sk4)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_88">c_0_88</A>,c_0_78]),c_0_19]),c_0_48])])).

<A NAME="c_0_92"></A>cnf(c_0_92,plain,
    ( ssList(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_89">c_0_89</A>,c_0_90]),c_0_16])])).

<A NAME="c_0_93"></A>cnf(c_0_93,negated_conjecture,
    ( ssItem(X1)
    | ~ ssList(app(sk4,sk4)) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_91">c_0_91</A>,<A HREF="#c_0_34">c_0_34</A>])).

<A NAME="c_0_94"></A>cnf(c_0_94,negated_conjecture,
    ( memberP(X1,sk5(X2,X1,X3))
    | app(X1,app(cons(X3,nil),X2)) != sk1
    | ~ ssItem(X3)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_35">c_0_35</A>,c_0_4]),<A HREF="#c_0_92">c_0_92</A>])).

<A NAME="c_0_95"></A>cnf(c_0_95,negated_conjecture,
    ( ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_93">c_0_93</A>,c_0_3]),c_0_48])])).

<A NAME="c_0_96"></A>cnf(c_0_96,negated_conjecture,
    ( memberP(X1,sk5(X2,X1,X3))
    | app(X1,app(cons(X3,nil),X2)) != sk1
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_94">c_0_94</A>,c_0_95])])).

<A NAME="c_0_97"></A>cnf(c_0_97,plain,
    ( app(nil,app(X1,X2)) = app(X1,X2)
    | ~ ssList(X1)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_4">c_0_4</A>,c_0_36]),c_0_16])])).

<A NAME="c_0_98"></A>cnf(c_0_98,plain,
    ( ssList(cons(X1,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_92">c_0_92</A>,c_0_95])])).

<A NAME="c_0_99"></A>cnf(c_0_99,plain,
    ( ~ memberP(nil,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_37">c_0_37</A>,c_0_95])])).

<A NAME="c_0_100"></A>cnf(c_0_100,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1)
    | ~ ssList(cons(X1,nil)) ),
    inference(er,[status(thm)],[<A HREF="#c_0_38">c_0_38</A>])).

<A NAME="c_0_101"></A>cnf(c_0_101,negated_conjecture,
    ( app(cons(X1,nil),X2) != sk1
    | ~ ssList(X2) ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_96">c_0_96</A>,c_0_97]),c_0_16]),c_0_98])]),<A HREF="#c_0_99">c_0_99</A>])).

<A NAME="c_0_102"></A>cnf(c_0_102,plain,
    ( singletonP(cons(X1,nil))
    | ~ ssItem(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_100">c_0_100</A>,c_0_32]),c_0_16])])).

<A NAME="c_0_103"></A>cnf(c_0_103,negated_conjecture,
    ( app(X1,X2) != sk1
    | ~ singletonP(X1)
    | ~ ssList(X2)
    | ~ ssList(X1) ),
    inference(spm,[status(thm)],[<A HREF="#c_0_101">c_0_101</A>,<A HREF="#c_0_39">c_0_39</A>])).

<A NAME="c_0_104"></A>cnf(c_0_104,plain,
    ( app(cons(X1,nil),X2) = cons(X1,X2)
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_40">c_0_40</A>,c_0_95])])).

<A NAME="c_0_105"></A>cnf(c_0_105,plain,
    ( singletonP(cons(X1,nil)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[<A HREF="#c_0_102">c_0_102</A>,c_0_95])])).

<A NAME="c_0_106"></A>cnf(c_0_106,negated_conjecture,
    ( cons(X1,X2) != sk1
    | ~ ssList(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_103">c_0_103</A>,c_0_104]),c_0_105]),c_0_98])])).

<A NAME="c_0_107"></A>cnf(c_0_107,negated_conjecture,
    ( $false ),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[<A HREF="#c_0_106">c_0_106</A>,c_0_41]),c_0_42])])]),c_0_6])]),<A HREF="#c_0_23">c_0_23</A>]),
    [proof]).
%------------------------------------------------------------------------------
%----ORIGINAL SYSTEM OUTPUT
% 0.00/0.03  % Problem    : SWC227-1 : TPTP v6.4.0. Released v2.4.0.
% 0.00/0.04  % Command    : eprover --auto-schedule --definitional-cnf=24 -R --print-version --tstp-format -s --proof-object --cpu-limit=%d %s
% 0.02/0.22  % Computer   : n128.star.cs.uiowa.edu
% 0.02/0.22  % Model      : x86_64 x86_64
% 0.02/0.22  % CPU        : Intel(R) Xeon(R) CPU E5-2609 0 @ 2.40GHz
% 0.02/0.22  % Memory     : 32218.75MB
% 0.02/0.22  % OS         : Linux 3.10.0-327.10.1.el7.x86_64
% 0.02/0.22  % CPULimit   : 300
% 0.02/0.22  % DateTime   : Tue May 10 20:37:07 CDT 2016
% 0.02/0.22  % CPUTime    : 
% 0.02/0.23  # Version: 1.9.1pre011
% 0.02/0.23  # No SInE strategy applied
% 0.02/0.23  # Trying AutoSched0 for 151 seconds
% 1.12/1.31  # AutoSched0-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S4d
% 1.12/1.31  # and selection function SelectCQIPrecWNTNp.
% 1.12/1.31  #
% 1.12/1.31  # Preprocessing time       : 0.010 s
% 1.12/1.31  # Presaturation interreduction done
% 1.12/1.31  
% 1.12/1.31  # Proof found!
% 1.12/1.31  # SZS status Unsatisfiable
% 1.12/1.31  # SZS output start CNFRefutation
% 1.12/1.31  cnf(c_0_0, negated_conjecture, (app(app(sk6,sk3),sk7)=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_15)).
% 1.12/1.31  cnf(c_0_1, negated_conjecture, (sk1=sk3), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_6)).
% 1.12/1.31  cnf(c_0_2, axiom, (frontsegP(X3,X1)|app(X1,X2)!=X3|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause144)).
% 1.12/1.31  cnf(c_0_3, axiom, (ssList(app(X2,X1))|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause85)).
% 1.12/1.31  cnf(c_0_4, axiom, (app(app(X3,X2),X1)=app(X3,app(X2,X1))|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause149)).
% 1.12/1.31  cnf(c_0_5, negated_conjecture, (ssList(sk6)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_13)).
% 1.12/1.31  cnf(c_0_6, negated_conjecture, (ssList(sk1)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_1)).
% 1.12/1.31  cnf(c_0_7, negated_conjecture, (ssList(sk7)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_14)).
% 1.12/1.31  cnf(c_0_8, negated_conjecture, (ssList(sk2)), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_2)).
% 1.12/1.31  cnf(c_0_9, negated_conjecture, (sk2=sk4), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_5)).
% 1.12/1.31  cnf(c_0_10, axiom, (X2=X3|app(X1,X2)!=app(X1,X3)|~ssList(X2)|~ssList(X1)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause150)).
% 1.12/1.31  cnf(c_0_11, axiom, (app(X2,skaf45(X1,X2))=X1|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause132)).
% 1.12/1.31  cnf(c_0_12, axiom, (ssList(skaf45(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause51)).
% 1.12/1.31  cnf(c_0_13, axiom, (rearsegP(X3,X2)|app(X1,X2)!=X3|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause143)).
% 1.12/1.31  cnf(c_0_14, axiom, (frontsegP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause61)).
% 1.12/1.31  cnf(c_0_15, axiom, (app(X1,nil)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause73)).
% 1.12/1.31  cnf(c_0_16, axiom, (ssList(nil)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause8)).
% 1.12/1.31  cnf(c_0_17, axiom, (X1=X3|app(X1,X2)!=app(X3,X2)|~ssList(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause151)).
% 1.12/1.31  cnf(c_0_18, axiom, (app(skaf46(X1,X2),X2)=X1|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause131)).
% 1.12/1.31  cnf(c_0_19, axiom, (ssList(skaf46(X1,X2))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause50)).
% 1.12/1.31  cnf(c_0_20, axiom, (frontsegP(app(X1,X3),X2)|~frontsegP(X1,X2)|~ssList(X3)|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause137)).
% 1.12/1.31  cnf(c_0_21, axiom, (nil=X2|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause119)).
% 1.12/1.31  cnf(c_0_22, axiom, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause118)).
% 1.12/1.31  cnf(c_0_23, negated_conjecture, (nil!=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_7)).
% 1.12/1.31  cnf(c_0_24, axiom, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause104)).
% 1.12/1.31  cnf(c_0_25, axiom, (nil=X2|hd(app(X2,X1))=hd(X2)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause123)).
% 1.12/1.31  cnf(c_0_26, axiom, (rearsegP(X1,X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause59)).
% 1.12/1.31  cnf(c_0_27, axiom, (ssList(tl(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause75)).
% 1.12/1.31  cnf(c_0_28, axiom, (nil=X2|tl(app(X2,X1))=app(tl(X2),X1)|~ssList(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause133)).
% 1.12/1.31  cnf(c_0_29, axiom, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssList(X4)|~ssList(X3)|~ssList(X1)|~ssItem(X2)|~duplicatefreeP(X5)|~ssList(X5)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause179)).
% 1.12/1.31  cnf(c_0_30, axiom, (ssItem(hd(X1))|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause76)).
% 1.12/1.31  cnf(c_0_31, axiom, (cons(X1,app(X2,X3))=app(cons(X1,X2),X3)|~ssItem(X1)|~ssList(X2)|~ssList(X3)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause157)).
% 1.12/1.31  cnf(c_0_32, axiom, (ssList(cons(X1,X2))|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause86)).
% 1.12/1.31  cnf(c_0_33, axiom, (app(X1,X2)=nil|nil!=X1|nil!=X2|~ssList(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause145)).
% 1.12/1.31  cnf(c_0_34, axiom, (duplicatefreeP(X1)|ssItem(X2)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause72)).
% 1.12/1.31  cnf(c_0_35, negated_conjecture, (memberP(X2,sk5(X3,X2,X1))|~ssItem(X1)|~ssList(X2)|~ssList(X3)|app(app(X2,cons(X1,nil)),X3)!=sk1), file('/export/starexec/sandbox2/benchmark/theBenchmark.p', co1_9)).
% 1.12/1.31  cnf(c_0_36, axiom, (app(nil,X1)=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause74)).
% 1.12/1.31  cnf(c_0_37, axiom, (~memberP(nil,X1)|~ssItem(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause71)).
% 1.12/1.31  cnf(c_0_38, axiom, (singletonP(X2)|cons(X1,nil)!=X2|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause116)).
% 1.12/1.31  cnf(c_0_39, axiom, (cons(skaf44(X1),nil)=X1|~singletonP(X1)|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause101)).
% 1.12/1.31  cnf(c_0_40, axiom, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssItem(X1)|~ssList(X2)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause120)).
% 1.12/1.31  cnf(c_0_41, axiom, (cons(skaf83(X1),skaf82(X1))=X1|nil=X1|~ssList(X1)), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause109)).
% 1.12/1.31  cnf(c_0_42, axiom, (ssList(skaf82(X1))), file('/export/starexec/sandbox2/benchmark/Axioms/SWC001-0.ax', clause13)).
% 1.12/1.31  cnf(c_0_43, negated_conjecture, (app(app(sk6,sk1),sk7)=sk4), inference(rw,[status(thm)],[c_0_0, c_0_1])).
% 1.12/1.31  cnf(c_0_44, plain, (frontsegP(app(X1,X2),X1)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_2]), c_0_3])).
% 1.12/1.31  cnf(c_0_45, negated_conjecture, (app(sk6,app(sk1,sk7))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_4]), c_0_5]), c_0_6]), c_0_7])])).
% 1.12/1.31  cnf(c_0_46, negated_conjecture, (frontsegP(sk4,sk6)|~ssList(app(sk1,sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_5])])).
% 1.12/1.31  cnf(c_0_47, negated_conjecture, (frontsegP(sk4,sk6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_3]), c_0_6]), c_0_7])])).
% 1.12/1.31  cnf(c_0_48, negated_conjecture, (ssList(sk4)), inference(rw,[status(thm)],[c_0_8, c_0_9])).
% 1.12/1.31  cnf(c_0_49, negated_conjecture, (X1=app(sk1,sk7)|app(sk6,X1)!=sk4|~ssList(app(sk1,sk7))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_45]), c_0_5])])).
% 1.12/1.31  cnf(c_0_50, negated_conjecture, (app(sk6,skaf45(sk4,sk6))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_47]), c_0_5]), c_0_48])])).
% 1.12/1.31  cnf(c_0_51, negated_conjecture, (skaf45(sk4,sk6)=app(sk1,sk7)|~ssList(app(sk1,sk7))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_12])])).
% 1.12/1.31  cnf(c_0_52, negated_conjecture, (skaf45(sk4,sk6)=app(sk1,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_3]), c_0_6]), c_0_7])])).
% 1.12/1.31  cnf(c_0_53, plain, (rearsegP(app(X1,X2),X2)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_3])).
% 1.12/1.31  cnf(c_0_54, negated_conjecture, (app(sk6,app(app(sk1,sk7),X1))=app(sk4,X1)|~ssList(app(sk1,sk7))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_45]), c_0_5])])).
% 1.12/1.31  cnf(c_0_55, negated_conjecture, (ssList(app(sk1,sk7))), inference(spm,[status(thm)],[c_0_12, c_0_52])).
% 1.12/1.31  cnf(c_0_56, negated_conjecture, (rearsegP(sk4,sk7)|~ssList(app(sk6,sk1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_43]), c_0_7])])).
% 1.12/1.31  cnf(c_0_57, negated_conjecture, (app(sk6,app(app(sk1,sk7),X1))=app(sk4,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55])])).
% 1.12/1.31  cnf(c_0_58, plain, (app(X1,skaf45(X1,X1))=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_11, c_0_14])).
% 1.12/1.31  cnf(c_0_59, negated_conjecture, (rearsegP(sk4,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_3]), c_0_5]), c_0_6])])).
% 1.12/1.31  cnf(c_0_60, plain, (X1=nil|app(X2,X1)!=X2|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_15]), c_0_16])])).
% 1.12/1.31  cnf(c_0_61, negated_conjecture, (app(sk4,skaf45(app(sk1,sk7),app(sk1,sk7)))=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_45]), c_0_12]), c_0_55])])).
% 1.12/1.31  cnf(c_0_62, negated_conjecture, (X1=app(sk6,sk1)|app(X1,sk7)!=sk4|~ssList(app(sk6,sk1))|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_43]), c_0_7])])).
% 1.12/1.31  cnf(c_0_63, negated_conjecture, (app(skaf46(sk4,sk7),sk7)=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_59]), c_0_7]), c_0_48])])).
% 1.12/1.31  cnf(c_0_64, negated_conjecture, (skaf45(app(sk1,sk7),app(sk1,sk7))=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_12]), c_0_48])])).
% 1.12/1.31  cnf(c_0_65, negated_conjecture, (skaf46(sk4,sk7)=app(sk6,sk1)|~ssList(app(sk6,sk1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_19])])).
% 1.12/1.31  cnf(c_0_66, plain, (frontsegP(app(app(X1,X2),X3),X1)|~ssList(X3)|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_44]), c_0_3])).
% 1.12/1.31  cnf(c_0_67, negated_conjecture, (app(app(sk1,sk7),nil)=app(sk1,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_64]), c_0_55])])).
% 1.12/1.31  cnf(c_0_68, negated_conjecture, (skaf46(sk4,sk7)=app(sk6,sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_3]), c_0_5]), c_0_6])])).
% 1.12/1.31  cnf(c_0_69, negated_conjecture, (frontsegP(app(sk1,sk7),sk1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_16]), c_0_6]), c_0_7])])).
% 1.12/1.31  cnf(c_0_70, negated_conjecture, (ssList(app(sk6,sk1))), inference(spm,[status(thm)],[c_0_19, c_0_68])).
% 1.12/1.31  cnf(c_0_71, negated_conjecture, (app(sk1,skaf45(app(sk1,sk7),sk1))=app(sk1,sk7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_69]), c_0_6]), c_0_55])])).
% 1.12/1.31  cnf(c_0_72, negated_conjecture, (app(app(sk6,sk1),app(sk7,X1))=app(sk4,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_43]), c_0_7])]), c_0_70])])).
% 1.12/1.31  cnf(c_0_73, plain, (nil=X1|app(X2,app(X3,X1))!=nil|~ssList(X1)|~ssList(X2)|~ssList(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_4]), c_0_3])).
% 1.12/1.31  cnf(c_0_74, negated_conjecture, (app(sk1,sk7)!=nil), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_71]), c_0_12]), c_0_6])]), c_0_23])).
% 1.12/1.31  cnf(c_0_75, negated_conjecture, (app(sk4,nil)=sk4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_15]), c_0_43]), c_0_16]), c_0_7])])).
% 1.12/1.31  cnf(c_0_76, plain, (cons(hd(X1),tl(app(X1,X2)))=app(X1,X2)|nil=X1|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_3]), c_0_22])).
% 1.12/1.31  cnf(c_0_77, negated_conjecture, (app(X1,sk4)!=nil|~ssList(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_45]), c_0_55]), c_0_5])]), c_0_74])).
% 1.12/1.31  cnf(c_0_78, plain, (app(skaf46(X1,X1),X1)=X1|~ssList(X1)), inference(spm,[status(thm)],[c_0_18, c_0_26])).
% 1.12/1.31  cnf(c_0_79, plain, (nil=X1|ssList(app(tl(X1),X2))|~ssList(X1)|~ssList(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_3]), c_0_22])).
% 1.12/1.31  cnf(c_0_80, negated_conjecture, (app(tl(sk4),nil)=tl(sk4)|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_75]), c_0_48]), c_0_16])])).
% 1.12/1.31  cnf(c_0_81, plain, (~ssItem(X1)|~ssList(app(app(X2,cons(X1,X3)),cons(X1,X4)))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(app(app(X2,cons(X1,X3)),cons(X1,X4)))), inference(er,[status(thm)],[c_0_29])).
% 1.12/1.31  cnf(c_0_82, negated_conjecture, (cons(hd(sk4),tl(sk4))=sk4|sk4=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_75]), c_0_48]), c_0_16])])).
% 1.12/1.31  cnf(c_0_83, negated_conjecture, (sk4!=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_19]), c_0_48])])).
% 1.12/1.31  cnf(c_0_84, negated_conjecture, (sk4=nil|ssList(tl(sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_48]), c_0_16])])).
% 1.12/1.31  cnf(c_0_85, plain, (nil=X1|~ssList(app(app(X2,cons(hd(X1),X3)),X1))|~ssList(X3)|~ssList(X2)|~ssList(X1)|~duplicatefreeP(app(app(X2,cons(hd(X1),X3)),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_24]), c_0_27]), c_0_30])).
% 1.12/1.31  cnf(c_0_86, negated_conjecture, (cons(hd(sk4),tl(sk4))=sk4), inference(sr,[status(thm)],[c_0_82, c_0_83])).
% 1.12/1.31  cnf(c_0_87, negated_conjecture, (ssList(tl(sk4))), inference(sr,[status(thm)],[c_0_84, c_0_83])).
% 1.12/1.31  cnf(c_0_88, negated_conjecture, (~ssList(app(app(X1,sk4),sk4))|~ssList(X1)|~duplicatefreeP(app(app(X1,sk4),sk4))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87]), c_0_48])]), c_0_83])).
% 1.12/1.31  cnf(c_0_89, plain, (ssList(cons(X1,app(X2,X3)))|~ssItem(X1)|~ssList(X3)|~ssList(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_3, c_0_31]), c_0_32])).
% 1.12/1.31  cnf(c_0_90, plain, (app(nil,nil)=nil), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_33])]), c_0_16])])).
% 1.12/1.31  cnf(c_0_91, negated_conjecture, (~ssList(app(sk4,sk4))|~duplicatefreeP(app(sk4,sk4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_78]), c_0_19]), c_0_48])])).
% 1.12/1.31  cnf(c_0_92, plain, (ssList(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_16])])).
% 1.12/1.31  cnf(c_0_93, negated_conjecture, (ssItem(X1)|~ssList(app(sk4,sk4))), inference(spm,[status(thm)],[c_0_91, c_0_34])).
% 1.12/1.31  cnf(c_0_94, negated_conjecture, (memberP(X1,sk5(X2,X1,X3))|app(X1,app(cons(X3,nil),X2))!=sk1|~ssItem(X3)|~ssList(X2)|~ssList(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_4]), c_0_92])).
% 1.12/1.31  cnf(c_0_95, negated_conjecture, (ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_3]), c_0_48])])).
% 1.12/1.31  cnf(c_0_96, negated_conjecture, (memberP(X1,sk5(X2,X1,X3))|app(X1,app(cons(X3,nil),X2))!=sk1|~ssList(X2)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_95])])).
% 1.12/1.31  cnf(c_0_97, plain, (app(nil,app(X1,X2))=app(X1,X2)|~ssList(X1)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4, c_0_36]), c_0_16])])).
% 1.12/1.31  cnf(c_0_98, plain, (ssList(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_95])])).
% 1.12/1.31  cnf(c_0_99, plain, (~memberP(nil,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_95])])).
% 1.12/1.31  cnf(c_0_100, plain, (singletonP(cons(X1,nil))|~ssItem(X1)|~ssList(cons(X1,nil))), inference(er,[status(thm)],[c_0_38])).
% 1.12/1.31  cnf(c_0_101, negated_conjecture, (app(cons(X1,nil),X2)!=sk1|~ssList(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_16]), c_0_98])]), c_0_99])).
% 1.12/1.31  cnf(c_0_102, plain, (singletonP(cons(X1,nil))|~ssItem(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_32]), c_0_16])])).
% 1.12/1.31  cnf(c_0_103, negated_conjecture, (app(X1,X2)!=sk1|~singletonP(X1)|~ssList(X2)|~ssList(X1)), inference(spm,[status(thm)],[c_0_101, c_0_39])).
% 1.12/1.31  cnf(c_0_104, plain, (app(cons(X1,nil),X2)=cons(X1,X2)|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_95])])).
% 1.12/1.31  cnf(c_0_105, plain, (singletonP(cons(X1,nil))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_95])])).
% 1.12/1.31  cnf(c_0_106, negated_conjecture, (cons(X1,X2)!=sk1|~ssList(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_105]), c_0_98])])).
% 1.12/1.31  cnf(c_0_107, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_41]), c_0_42])])]), c_0_6])]), c_0_23]), ['proof']).
% 1.12/1.31  # SZS output end CNFRefutation
% 1.12/1.31  # Training examples: 0 positive, 0 negative
% 1.12/1.31  
% 1.12/1.31  # -------------------------------------------------
% 1.12/1.31  # User time                : 1.070 s
% 1.12/1.31  # System time              : 0.015 s
% 1.12/1.31  # Total time               : 1.085 s
% 1.12/1.31  # Maximum resident set size: 1868 pages
%------------------------------------------------------------------------------
</pre>
<HR>

</body>
</html>
